# 10. Sets

This Jupyter Notebook contains the [problems, or assignments, for Lesson 10](https://snakify.org/en/lessons/sets/problems/), gratefully imported from the online course on [Snakify](https://snakify.org). If you get stuck, you can either ask the teacher and the teaching assistents, or skim through the [theory](https://snakify.org/en/lessons/sets/) on Snakify.

```python
# Please run the contents of this cell.
# Select this cell and press Ctrl-Enter (or Command-Enter on mac).

from assertions import *
from test_cells import *
```

Output:
```
Assertions imported OK
Test cells imported OK
```

<br />
<br />
<br />
<br />

# 10.1 The number of distinct numbers

Given a list of integers. Determine (and return) how many distinct numbers there are.

This task can be solved in one line of code.

## Example input
```python
1 2 3 2 1
```

## Example return value
```
3
```

<!-- grade_id: assignment-10-01-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 10.1

def assignment_10_1(numbers):
    # YOUR CODE HERE
    return len(set(numbers))
```

```python
# Test for assignment 10.1

test_cell_10_1(assignment_10_1)
```

Output:
```
Please enter a series of integer numbers, separated by spaces (Press Enter for default values 1 2 3 2 1):  

The list [1, 2, 3, 2, 1] has 3 unique numbers.
```

<!-- grade_id: assignment-10-01-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 10.1

check_assertion(assignment_10_1)
```

Output:
```
Call assignment_10_1([1, 2, 3, 2, 1]) ... OK
Call assignment_10_1([1, 2, 3, 4, 5, 6, 7, 8, 9, 10]) ... OK
Call assignment_10_1([1, 2, 3, 4, 5, 1, 2, 1, 2, 7, 3]) ... OK
Call assignment_10_1([1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1]) ... OK

Well done! You seem to have solved assignment_10_1!
```

<br />
<br />
<br />
<br />

# 10.2 The number of equal numbers

Given two lists of numbers. Count how many unique numbers both lists have in common and return that number.

This task can be solved in one line of code.

## Example input
```python
1 3 2
4 3 2
```

## Example return value
```python
2
```

<!-- grade_id: assignment-10-02-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 10.2

def assignment_10_2(numbers1, numbers2):
    # YOUR CODE HERE
    return len(set(numbers1) & set(numbers2))
```

```python
# Test for assignment 10.2

test_cell_10_2(assignment_10_2)
```

Output:
```
Please enter a series of integer numbers, separated by spaces (Press Enter for default values 1 3 2):  
Please enter a series of integer numbers, separated by spaces (Press Enter for default values 4 3 2):  

The lists [1, 3, 2] and [4, 3, 2] have 2 numbers in common.
```

<!-- grade_id: assignment-10-02-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 10.2

check_assertion(assignment_10_2)
```

Output:
```
Call assignment_10_2([1, 3, 2], [4, 3, 2]) ... OK
Call assignment_10_2([1, 2, 6, 4, 5, 7], [10, 2, 3, 4, 8]) ... OK
Call assignment_10_2([1, 7, 3, 8, 10, 2, 5], [6, 5, 2, 8, 4, 3, 7]) ... OK
Call assignment_10_2([75, 8, 65, 66, 3, 30, 38, 45, 82, 93, 89, 81, 68, 20, 10, 99, 84, 71, 59, 36, 54, 50, 14, 86, 91, 16, 96, 37, 80, 79, 69, 52, 64, 88, 32, 28, 53, 2, 55, 95, 21, 62, 31, 5, 85, 22, 39, 78, 7, 83, 43, 18, 48, 33, 97, 11, 19, 100, 41, 12, 29, 58, 74], [26, 94, 17, 71, 84, 75, 5, 29, 21, 38, 2, 35, 6, 78, 89, 53, 58, 86, 93, 63, 36, 74, 27, 72, 76, 49, 52, 64, 65, 100, 28, 69, 8, 1, 48, 39, 50, 87, 73, 90, 70, 13, 11, 62, 66, 80, 37, 15, 14, 54, 67, 34, 59, 44, 43, 20, 57, 82, 45, 95, 92, 56, 47, 51, 33, 10, 98]) ... OK

Well done! You seem to have solved assignment_10_2!
```

<br />
<br />
<br />
<br />

# 10.3 The intersection of sets

Given two lists of numbers. Find all the numbers that occur in both lists and return them in a new list, in ascending order.

Even this task can be solved in one line of code, using the `sorted` builtin function.

## Example input
```python
1 3 2
4 3 2
```

## Example output
```python
[2, 3]
```

<!-- grade_id: assignment-10-03-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 10.3

def assignment_10_3(numbers1, numbers2):
    # YOUR CODE HERE
    return sorted(list(set(numbers1) & set(numbers2)))
```

```python
# Test for assignment 10.3

test_cell_10_3(assignment_10_3)
```

Output:
```
Please enter a series of integer numbers, separated by spaces (Press Enter for default values 1 3 2):  
Please enter a series of integer numbers, separated by spaces (Press Enter for default values 4 3 2):  

The lists [1, 3, 2] and [4, 3, 2] both contain [2, 3].
```

<!-- grade_id: assignment-10-03-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 10.3

check_assertion(assignment_10_3)
```

Output:
```
Call assignment_10_3([1, 3, 2], [4, 3, 2]) ... OK
Call assignment_10_3([1, 2, 6, 4, 5, 7], [10, 2, 3, 4, 8]) ... OK
Call assignment_10_3([82, 54, 91, 100, 70, 33, 88, 14, 52, 48, 56, 20, 63, 16, 22, 23, 30, 8, 84, 75, 45], [10, 44, 77, 90, 43, 75, 25, 24, 5, 21, 71, 70, 83, 68, 18, 92, 81, 57, 27, 67, 48, 6]) ... OK

Well done! You seem to have solved assignment_10_3!
```

<br />
<br />
<br />
<br />

# 10.4 Has the number been encountered before?

Given a list of numbers, determine if the next number has already been encountered. Create a list with booleans. For each number, append  `True` if this number has already been encountered, and append `False`, if it has not already been encountered. The first element of your resulting list will always be `False`, of course, since no numbers were encountered at that point.

## Example input
```python
1 2 3 2 3 4
```

## Example return value
```python
[False, False, False, True, True, False]
```


<!-- grade_id: assignment-10-04-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 10.4

def assignment_10_4(numbers):
    result = [False]
    encountered_numbers = {numbers[0]} 
    for number in numbers[1:]: 
        if number in encountered_numbers: 
            result.append(True)
        else: 
            encountered_numbers.add(number)
            result.append(False)
    
    return result
```

```python
# Test for assignment 10.4

test_cell_10_4(assignment_10_4)
```

Output:
```
Please enter a series of integer numbers, separated by spaces (Press Enter for default values 1 2 3 2 3 4):  

The list [1, 2, 3, 2, 3, 4] has the following result [False, False, False, True, True, False].
```

<!-- grade_id: assignment-10-04-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 10.4

check_assertion(assignment_10_4)
```

Output:
```
Call assignment_10_4([1, 2, 3, 2, 3, 4]) ... OK
Call assignment_10_4([1, 1, 1, 1, 1, 1, 1, 1, 1, 1]) ... OK
Call assignment_10_4([1, 2, 3, 4, 5, 6, 7, 8, 9, 10]) ... OK
Call assignment_10_4([10, 10, 4, 2, 5, 5, 10, 4, 1, 10]) ... OK

Well done! You seem to have solved assignment_10_4!
```

<br />
<br />
<br />
<br />

# 10.5 Blocks

Alice and Bob like to play with colored blocks. Each child has its own set of blocks and each block has a distinct color. They want to know which block colors they have in common, and which block colors they have individually.

Finalize the function `assignment_10_5(alice, bob)` that takes two lists with colors (a single word each). Make the function return  three lists with the:
1. The colors of the blocks in both collections
1. The colors of the blocks which are in Alice's collection but not in Bob's
1. The colors of the blocks in Bob's collection but not in Alice's.

**NB** Sort the colors in each list.

## Example input
```python
red green orange yellow
green blue red
```

## Example return value
```python
(['green', 'red'], ['orange', 'yellow'], ['blue'])
```

<!-- grade_id: assignment-10-05-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 10.5

def assignment_10_5(alice, bob):
    alice = set(alice) 
    bob = set(bob) 
    
    both = sorted(list(alice & bob))
    unique_alice = sorted(list(alice - bob))
    unique_bob = sorted(list(bob - alice))
    return (both, unique_alice, unique_bob) 
```

```python
# Test for assignment 10.5

test_cell_10_5(assignment_10_5)
```

Output:
```
Please enter the single-word colors of each block in Alice's collection (Press Enter for default "red green orange yellow"):  
Please enter the single-word colors of each block in Bob's collection (Press Enter for default "green blue red"):  

The colors both children have in common: ['green', 'red']
The unqiue colors Alice has: ['orange', 'yellow']
The unqiue colors Bob has: ['blue']
```

<!-- grade_id: assignment-10-05-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 10.5

check_assertion(assignment_10_5)
```

Output:
```
Call assignment_10_5(['blue', 'pink', 'gray'], ['red', 'blue', 'yellow', 'brown', 'white', 'black']) ... OK
Call assignment_10_5(['red', 'green', 'orange', 'pink', 'white', 'black', 'gray'], ['red', 'green', 'yellow', 'pink']) ... OK
Call assignment_10_5(['blue', 'yellow', 'orange', 'pink'], ['red', 'green', 'yellow', 'brown', 'pink', 'black']) ... OK
Call assignment_10_5(['red', 'brown', 'pink', 'gray'], ['red', 'blue', 'yellow', 'orange', 'pink', 'black', 'gray']) ... OK
Call assignment_10_5(['red', 'green', 'blue', 'orange', 'pink'], ['red', 'blue', 'yellow', 'brown', 'pink', 'black']) ... OK
Call assignment_10_5(['red', 'green', 'yellow', 'orange', 'pink', 'black'], ['red', 'yellow', 'orange', 'brown', 'pink', 'white', 'black']) ... OK

Well done! You seem to have solved assignment_10_5!
```

<br />
<br />
<br />
<br />

# 10.6 The number of distinct words in some text

Given some lines of text, return the number of distinct words that appear in the text.

For this, we define a word to be a sequence of non-whitespace characters, seperated by one or more whitespace characters. Note that punctuation marks, such as commas and periods, are part of a word according to this definition.

## Example input
```
She sells sea shells on the sea shore;
The shells that she sells are sea shells I'm sure.
So if she sells sea shells on the sea shore,
I'm sure that the shells are sea shore shells.
```

## Example return value
```python
19
```

<!-- grade_id: assignment-10-06-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 10.6

def assignment_10_6(string):
    return len(set(string.split()))
    
```

```python
# Test for assignment 10.6

test_cell_10_6(assignment_10_6)
```

Output:
```
Please enter a string (Press Enter for default "
    
She sells sea shells on the sea shore;
The shells that she sells are sea shells I'm sure.
So if she sells sea shells on the sea shore,
I'm sure that the shells are sea shore shells.

"):  

This text has 19 unique words.
```

<!-- grade_id: assignment-10-06-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 10.6

check_assertion(assignment_10_6)
```

Output:
```
Call assignment_10_6('She sells sea shells on the sea shore; The shells that she sells are sea shells I'm sure. So if she sells sea shells on the sea shore, I'm sure that the shells are sea shore shells. ') ... OK
Call assignment_10_6(' The other two, slight air, and purging fire, Are both with thee, wherever I abide, The first my thought, the other my desire, These present-absent with swift motion slide. For when these quicker elements are gone In tender embassy of love to thee, My life being made of four, with two alone, Sinks down to death, oppressed with melancholy. Until life's composition be recured, By those swift messengers returned from thee, Who even but now come back again assured, Of thy fair health, recounting it to me. This told, I joy, but then no longer glad, I send them back again and straight grow sad. ') ... OK
Call assignment_10_6('a a a a a a a a a a a a a a a a a a a a a a a a a a a a a a a a a a a a a') ... OK
Call assignment_10_6('a b a c a b a d a b a c a b a')
```

<br />
<br />
<br />
<br />

# 10.7 Guess the number

Augustus and Beatrice play the following game. Augustus thinks of a secret integer number from 1 to $n$, where $n$ is a agreed upon beforehand. Beatrice tries to guess the number by providing a set of integers. Augustus answers `YES` if his secret number exists in the provided set, or `NO`, if his number does not exist in the provided numbers. Then after a few questions Beatrice, totally confused, asks you to `HELP` her determine Augustus's secret number.

Finalize the function `assignment_10_7(n, turns)` which takes two arguments `n` and `turns`. The argument `n` corresponds to the highest possible number $n$ in Augustus' secret number range. The argument `turns` is a list of lists; each of these lists contains the guessed numbers for that turn including Augustus's answer: `True` if the secret number is present, or `False` if it is not.

The function should return an **ordered** list of the remaining possible numbers in Augustus' secret range.

If you press Enter without values, the default example input will be used in the test cell.

## Example input
```python
10
1 2 3 4 5
YES
2 4 6 8 10
NO
HELP
```

This will be converted in the test cell to this specific function call:
```python
assignment_10_7(10, [[1, 2, 3, 4, 5, True], [2, 4, 6, 8, 10, False]])
```

## Example return value
```python
[1, 3, 5]
```

<!-- grade_id: assignment-10-07-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 10.7

def assignment_10_7(n, turns):
    result = {x for x in range(1, n +1)}

    for turn in turns:
        if turn[-1]:
            result = result & set(turn[:-1])
        else:
            result = result - set(turn[:-1])
    
    return sorted(list(result))
    
    
```

```python
# Test cell for assignment 10.7

test_cell_10_7(assignment_10_7)
```

Output:
```
Would you like to use the default example input (Y/n)?  

The numbers still possible in the secret range of Augustus between 1 and 10 are [1, 3, 5].
```

<!-- grade_id: assignment-10-07-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 10.7

check_assertion(assignment_10_7)
```

Output:
```
Call assignment_10_7(10, [[1, 2, 3, 4, 5, True], [2, 4, 6, 8, 10, False]]) ... OK
Call assignment_10_7(10, [[1, 2, 3, 4, 5, 6, 7, 8, 9, 10, True], [1, False], [2, False], [3, False], [4, False], [6, False], [7, False], [8, False], [9, False], [10, False]]) ... OK
Call assignment_10_7(100, [[3, 5, 8, 10, 11, 13, 14, 15, 16, 20, 21, 23, 25, 27, 29, 31, 32, 33, 35, 42, 43, 46, 47, 48, 51, 52, 57, 58, 67, 68, 72, 74, 75, 76, 79, 80, 81, 82, 83, 84, 86, 89, 90, 91, 92, 93, 94, 95, 96, 99, True], [1, 4, 7, 9, 14, 16, 17, 19, 20, 23, 25, 27, 28, 29, 30, 31, 32, 33, 41, 42, 45, 46, 48, 50, 52, 54, 59, 60, 61, 63, 66, 67, 68, 69, 70, 71, 73, 74, 79, 81, 82, 83, 89, 90, 91, 92, 93, 94, 95, 100, True], [3, 5, 6, 15, 16, 18, 19, 20, 21, 22, 23, 24, 26, 28, 29, 35, 37, 38, 39, 40, 41, 43, 44, 45, 46, 47, 53, 55, 56, 60, 61, 63, 64, 65, 68, 70, 71, 73, 76, 79, 80, 81, 82, 83, 86, 90, 91, 95, 96, 99, True], [1, 2, 3, 5, 11, 16, 17, 23, 25, 31, 32, 33, 34, 35, 38, 39, 41, 43, 44, 46, 47, 49, 51, 53, 5
```

<br />
<br />
<br />
<br />

# 10.8 Polyglots

Each student at a certain school speaks a number of languages. We need to determine which languges are spoken by all the students, as well as which languages are spoken by at least one student.

Finalize the function `assignment_10_8(polyglots)` in such a way that it returns two lists:
1. a list with languages spoken by **all** students
2. a list with languages spoken by **at least one** of the students

**NB** Return the languages in alphabetic order in the lists.

`polyglots` will be created by the test cell from space separated language names, one student per line. Or, if you start with just an Enter on the first input, you will be using the default example values.

## Example input
```python
Russian English Japanese
Russian English
English
```

## Example return value
```python
(['English'], ['English', 'Japanese', 'Russian'])
```

<!-- grade_id: assignment-10-08-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 10.8

def assignment_10_8(polyglots):
    all_languages = set()
    for language in polyglots: 
        all_languages.update(language)
    
    all_students_speak = set(polyglots[0])

    for language in polyglots[1:]: 
        current_student = set(language) 
        all_students_speak = all_students_speak & current_student
        
    return (sorted(list(all_students_speak)), sorted(list(all_languages)))
```

```python
# Test for assignment 10.8

test_cell_10_8(assignment_10_8)
```

Output:
```
Enter the languages separated by spaces spoken by student 1,
or an empty line after all students have been entered.
 

The result of your function is: (['English'], ['English', 'Japanese', 'Russian'])

This means that ['English'] is/are spoken by all students,
and ['English', 'Japanese', 'Russian'] is/are spoken by at least one student.
```

<!-- grade_id: assignment-10-08-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 10.8

check_assertion(assignment_10_8)
```

Output:
```
Call assignment_10_8([['Russian', 'English', 'Japanese'], ['Russian', 'English'], ['English']]) ... OK
Call assignment_10_8([['German'], ['German'], ['German'], ['German'], ['German']]) ... OK
Call assignment_10_8([['Russian', 'Spanish', 'German'], ['Russian', 'German', 'Spanish'], ['Spanish', 'Russian', 'German'], ['German', 'Russian', 'Spanish']]) ... OK

Well done! You seem to have solved assignment_10_8!
```