# 2. Integer and float numbers

This Jupyter Notebook contains the [problems, or assignments, for Lesson 2](https://snakify.org/en/lessons/integer_float_numbers/problems/), gratefully imported from the online course on [Snakify](https://snakify.org). If you get stuck, you can either ask the teacher and the teaching assistents, or skim through the [theory](https://snakify.org/en/lessons/integer_float_numbers/) and [explanation of the steps](https://snakify.org/en/lessons/integer_float_numbers/steps/1/) on Snakify.

```python
# Please run the contents of this cell.
# Select this cell and press Ctrl-Enter (or Command-Enter on mac).

from assertions import *
from test_cells import *
```

Output:
```
Assertions imported OK
Test cells imported OK
```

<br />
<br />
<br />
<br />

# 2.1 Last digit of integer

Given a positive integer `number`, return its last digit.

## Example input

```python
179
```

## Example return value

```python
9
```

<!-- grade_id: assignment-02-01-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 2.1

def assignment_2_1(number):
    # Hint: The remainder operator % could be used
    # YOUR CODE HERE
    return number % 10

```

```python
# Test for assignment 2.1

test_cell_2_1(assignment_2_1)
```

Output:
```
Please enter an integer value (Press Enter for default value 179):  

The last digit of 179 is: 
9
```

<!-- grade_id: assignment-02-01-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 2.1

check_assertion(assignment_2_1)
```

Output:
```
Call assignment_2_1(179) ... OK
Call assignment_2_1(40) ... OK
Call assignment_2_1(101) ... OK
Call assignment_2_1(222) ... OK
Call assignment_2_1(1923) ... OK
Call assignment_2_1(74) ... OK
Call assignment_2_1(505) ... OK
Call assignment_2_1(996) ... OK
Call assignment_2_1(3487) ... OK
Call assignment_2_1(308) ... OK
Call assignment_2_1(1) ... OK
Call assignment_2_1(2) ... OK
Call assignment_2_1(3) ... OK
Call assignment_2_1(4) ... OK
Call assignment_2_1(5) ... OK
Call assignment_2_1(6) ... OK
Call assignment_2_1(7) ... OK
Call assignment_2_1(8) ... OK
Call assignment_2_1(9) ... OK
Call assignment_2_1(10) ... OK

Well done! You seem to have solved assignment_2_1!
```

<br />
<br />
<br />
<br />

# 2.2 Tens digit

Given a positive integer `number`. Return the digit that refers to the tens: 7 in 73, 8 in 86 and 2 in 123, for example.

## Example input

```python
73
```

## Example return value

```python
7
```

<!-- grade_id: assignment-02-02-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 2.2

def assignment_2_2(number):
    # Hint: The remainder operator % and integer division // could be used
    # YOUR CODE HERE
    return (number % 100) // 10
```

```python
# Test for assignment 2.2

test_cell_2_2(assignment_2_2)
```

Output:
```
Please enter an integer value (Press Enter for default value 73):  

The tens digit of 73 is: 
7
```

<!-- grade_id: assignment-02-02-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 2.2

check_assertion(assignment_2_2)
```

Output:
```
Call assignment_2_2(73) ... OK
Call assignment_2_2(1234) ... OK
Call assignment_2_2(10) ... OK
Call assignment_2_2(29) ... OK
Call assignment_2_2(37) ... OK
Call assignment_2_2(48) ... OK
Call assignment_2_2(50) ... OK
Call assignment_2_2(67) ... OK
Call assignment_2_2(81) ... OK
Call assignment_2_2(99) ... OK
Call assignment_2_2(179) ... OK
Call assignment_2_2(100) ... OK
Call assignment_2_2(999) ... OK
Call assignment_2_2(854345) ... OK
Call assignment_2_2(1000000) ... OK
Call assignment_2_2(9999) ... OK
Call assignment_2_2(9) ... OK
Call assignment_2_2(1) ... OK
Call assignment_2_2(0) ... OK

Well done! You seem to have solved assignment_2_2!
```

<br />
<br />
<br />
<br />

# 2.3 Sum of digits

Given a positive three-digit integer `number`. Find the sum of its digits and return it.

## Example input

```python
678
```

## Example return value

```python
21
```

<!-- grade_id: assignment-02-03-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 2.3

def assignment_2_3(number):
    # Hint: The remainder operator % and integer division // could be used
    # YOUR CODE HERE
    return (number % 10) + (number // 10 % 10) + (number // 100) 
```

```python
# Test for assignment 2.3

test_cell_2_3(assignment_2_3)
```

Output:
```
Please enter a positive 3-digit number (Press Enter for default value 678):  590

The sum of the digits in 590 is:
14
```

<!-- grade_id: assignment-02-03-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 2.3

check_assertion(assignment_2_3)
```

Output:
```
Call assignment_2_3(179) ... OK
Call assignment_2_3(829) ... OK
Call assignment_2_3(204) ... OK
Call assignment_2_3(100) ... OK
Call assignment_2_3(999) ... OK
Call assignment_2_3(483) ... OK

Well done! You seem to have solved assignment_2_3!
```

<br />
<br />
<br />
<br />

# 2.4 Fractional part

Given a positive real value `number`, return its fractional part.

## Example input

```python
17.9
```

## Example return value

```python
0.9
```

<!-- grade_id: assignment-02-04-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 2.4

def assignment_2_4(number):
    # Hint: use integer division // or the int() function to be able to subtract the integer part
    # YOUR CODE HERE
    return number - int(number)
```

```python
# Test for assignment 2.4

test_cell_2_4(assignment_2_4)
```

Output:
```
Please enter a value (Press Enter for default value 17.9):  

The fractional part of 17.9 is:
0.900
```

<!-- grade_id: assignment-02-04-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 2.4

check_assertion(assignment_2_4)
```

Output:
```
Call assignment_2_4(17.9) ... OK
Call assignment_2_4(10.34) ... OK
Call assignment_2_4(0.001) ... OK
Call assignment_2_4(179.001) ... OK
Call assignment_2_4(123.4) ... OK

Well done! You seem to have solved assignment_2_4!
```

<br />
<br />
<br />
<br />

# 2.5 First digit after decimal point

Given a positive real value `number`, return its first digit to the right of the decimal point. You may assume that the given `number` argument is always a floating point number, so it will always have a decimal point.

## Example input

```python
1.79
```

## Example return value

```python
7
```

<!-- grade_id: assignment-02-05-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 2.5

def assignment_2_5(number):
    # YOUR CODE HERE
    return int(number * 10) % 10
```

```python
# Test for assignment 2.5

test_cell_2_5(assignment_2_5)
```

Output:
```
Please enter a value (Press Enter for default value 1.79):  

The first digit after the decimal point of 1.79 is:
7
```

<!-- grade_id: assignment-02-05-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 2.5

check_assertion(assignment_2_5)
```

Output:
```
Call assignment_2_5(1.79) ... OK
Call assignment_2_5(10.34) ... OK
Call assignment_2_5(0.001) ... OK
Call assignment_2_5(19.99) ... OK
Call assignment_2_5(179.12) ... OK
Call assignment_2_5(5.29) ... OK
Call assignment_2_5(0.31) ... OK
Call assignment_2_5(12.45) ... OK
Call assignment_2_5(18.58) ... OK
Call assignment_2_5(0.83) ... OK
Call assignment_2_5(999.99) ... OK
Call assignment_2_5(146.67) ... OK
Call assignment_2_5(1293.73) ... OK
Call assignment_2_5(0.09999) ... OK
Call assignment_2_5(312.19999) ... OK
Call assignment_2_5(901.2999) ... OK
Call assignment_2_5(3.39999) ... OK
Call assignment_2_5(2371.49999) ... OK
Call assignment_2_5(290.59999) ... OK
Call assignment_2_5(90291.69999) ... OK
Call assignment_2_5(412.79999) ... OK
Call assignment_2_5(1.89999) ... OK
Call assignment_2_5(0.999999) ... OK

Well done! You seem to have solved assignment_2_5!
```

<br />
<br />
<br />
<br />

# 2.6 Digital clock

Given the integer `minutes_since_midnight` - the number of minutes that have passed since midnight - how many hours and minutes are displayed on the 24h digital clock?

The function below should return two numbers: the number of hours (between 0 and 23) and the number of minutes (between 0 and 59).

For example, if `minutes_since_midnight` = 150, then 150 minutes have passed since midnight - i.e. now is 2:30 am. So the program should return (2, 30).


## Example input

```python
150
```

## Example return value

```python
(2, 30)
```

<!-- grade_id: assignment-02-06-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 2.6

def assignment_2_6(minutes_since_midnight):
    # YOUR CODE HERE
    return (minutes_since_midnight // 60, minutes_since_midnight - ((minutes_since_midnight // 60) * 60))
```

```python
# Test for assignment 2.6

test_cell_2_6(assignment_2_6)
```

Output:
```
Please enter the number of minutes since midnight (Press Enter for default value 150):  567

The time at 567 minutes past midnight is:
(9, 27)
```

<!-- grade_id: assignment-02-06-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 2.6

check_assertion(assignment_2_6)
```

Output:
```
Call assignment_2_6(150) ... OK
Call assignment_2_6(180) ... OK
Call assignment_2_6(444) ... OK
Call assignment_2_6(1111) ... OK
Call assignment_2_6(1439) ... OK

Well done! You seem to have solved assignment_2_6!
```

<br />
<br />
<br />
<br />

# 2.7 Total cost

A cupcake costs `dollars` dollars and `cents` cents. Determine, how many dollars and cents should one pay for `number_cupcakes` cupcakes. The function below uses three integer arguments, `dollars`, `cents`, `number_cupcakes`. It should return two integer numbers: total cost in dollars and cents.

## Example input

```python
2
50
5
```

## Example return value

```python
(12, 50)
```

<!-- grade_id: assignment-02-07-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 2.7

def assignment_2_7(dollars, cents, number_cupcakes):
    # Hint: counting money should be easy for business administration students :-) 
    # YOUR CODE HERE
    dollars = dollars * 100
    total_cost = (dollars + cents) * number_cupcakes  
    return (total_cost // 100, total_cost % 100) 
```

```python
# Test for assignment 2.7

test_cell_2_7(assignment_2_7)
```

Output:
```
Please enter the amount of dollars (Press Enter for default value 2):  
Please enter the amount of cents (Press Enter for default value 50):  
Please enter number of cupcakes (Press Enter for default value 5):  

The total cost of 5 cupcakes at 2 dollars 50 cents each is:
(12, 50)
```

<!-- grade_id: assignment-02-07-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 2.7
check_assertion(assignment_2_7)
```

Output:
```
Call assignment_2_7(2, 50, 4) ... OK
Call assignment_2_7(10, 15, 2) ... OK
Call assignment_2_7(3000, 99, 3000) ... OK
Call assignment_2_7(2029, 34, 1848) ... OK
Call assignment_2_7(1886, 73, 295) ... OK
Call assignment_2_7(1069, 40, 1967) ... OK
Call assignment_2_7(905, 79, 496) ... OK
Call assignment_2_7(1967, 91, 926) ... OK
Call assignment_2_7(2255, 76, 1090) ... OK
Call assignment_2_7(2435, 6, 1965) ... OK

Well done! You seem to have solved assignment_2_7!
```

<br />
<br />
<br />
<br />

# 2.8 Clock face - 1

`hour` hours, `minutes` minutes and `seconds` seconds have passed since midnight (0 ≤ `hours` < 12, 0 ≤ `minutes` < 60, 0 ≤ `seconds` < 60). Determine angle (in degrees) between the hour hand and the top of the clock and return that value.

## Example input

```python
1
2
6
```

## Example return value

```python
31.05
```

<!-- grade_id: assignment-02-08-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 2.8

def assignment_2_8(hours, minutes, seconds):
    # Again, do not worry about rounding errors in this exercise!
    # YOUR CODE HERE
    minutes_in_hours = minutes / 60 
    seconds_in_hours = seconds / 3600 
    total_hours = hours + seconds_in_hours + minutes_in_hours 
    return (30 * total_hours)

```

```python
# Test for assignment 2.8

test_cell_2_8(assignment_2_8)
```

Output:
```
Please enter number of hours (Press Enter for default value 1):  
Please enter number of minutes (Press Enter for default value 2):  
Please enter number of seconds (Press Enter for default value 6):  

The angle (in degrees) at 01:02:06 is:
31.050000000000004
```

<!-- grade_id: assignment-02-08-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 2.8

check_assertion(assignment_2_8)
```

Output:
```
Call assignment_2_8(1, 2, 6) ... OK
Call assignment_2_8(1, 0, 0) ... OK
Call assignment_2_8(0, 1, 0) ... OK
Call assignment_2_8(0, 2, 0) ... OK
Call assignment_2_8(0, 2, 1) ... OK
Call assignment_2_8(0, 0, 1) ... OK
Call assignment_2_8(11, 59, 59) ... OK
Call assignment_2_8(7, 18, 49) ... OK
Call assignment_2_8(4, 54, 55) ... OK
Call assignment_2_8(9, 9, 37) ... OK

Well done! You seem to have solved assignment_2_8!
```

<br />
<br />
<br />
<br />

# 2.9 Clock face - 2

Hour hand turned by `degrees` degrees since midnight. Determine the angle by which minute hand turned since the start of the current hour and return that value. Input and output in this problem are floating-point numbers.

## Example input

```python
190.0
```

## Example return value

```python
120.0
```


<!-- grade_id: assignment-02-09-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 2.9
def assignment_2_9(degrees):
    # YOUR CODE HERE
    angle_past_hour = degrees % 30 
    angle_minutes_hand = angle_past_hour * 12
    return(angle_minutes_hand)
    
    
```

```python
# Test for assignment 2.9

test_cell_2_9(assignment_2_9)
```

Output:
```
Please enter hour hand degrees after midnight (Press Enter for default value 190.0):  

For an hour hand angle of 190.0 degrees, the number of degrees for the minute hand are:
120.0
```

<!-- grade_id: assignment-02-09-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 2.9

check_assertion(assignment_2_9)
```

Output:
```
Call assignment_2_9(190.0) ... OK
Call assignment_2_9(5.0) ... OK
Call assignment_2_9(10.0) ... OK
Call assignment_2_9(15.0) ... OK
Call assignment_2_9(20.0) ... OK
Call assignment_2_9(25.0) ... OK
Call assignment_2_9(29.0) ... OK
Call assignment_2_9(360.0) ... OK
Call assignment_2_9(700.0) ... OK

Well done! You seem to have solved assignment_2_9!
```