# 11. Dictionaries

This Jupyter Notebook contains the [problems, or assignments, for Lesson 11](https://snakify.org/en/lessons/dictionaries_dicts/problems/), gratefully imported from the online course on [Snakify](https://snakify.org). If you get stuck, you can either ask the teacher and the teaching assistents, or skim through the [theory](https://snakify.org/en/lessons/dictionaries_dicts/) on Snakify.

```python
# Please run the contents of this cell.
# Select this cell and press Ctrl-Enter (or Command-Enter on mac).

from assertions import *
from test_cells import *
```

Output:
```
Assertions imported OK
Test cells imported OK
```

<br />
<br />
<br />
<br />

# 11.1 Number of occurrences

The text is given in a single line. Return an alphabatically sorted dictionary with the word count for each unique word. Remember that in recent versions of Python, a dictionary will retain the order in which you added items to it.

A word is a sequence of non-whitespace characters. Two consecutive words are separated by one or more spaces. Punctiation marks are a part of a word, by this definition.

**Hint** You can use the power of `set()` to find unique words and turn it into a sorted list with `sorted()` and `list()`:
```python
unique_words = sorted(set(string.split()))
```
You can use `split()` without arguments to obtain a list of all the words in the string:
```python
all_words = string.split()
```
Build your dictionary with these two variables.

## Example input
```python
one two one two three
```

## Example return value
```python
{'one': 2, 'three': 1, 'two': 2}
```

<!-- grade_id: assignment-11-01-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 11.1

def assignment_11_1(string):
    words = string.split()
    word_count = { word: words.count(word) for word in words}
    
    return word_count
```

```python
# Test for assignment 11.1

test_cell_11_1(assignment_11_1)
```

Output:
```
Please enter a string (Press Enter for default "one two one two three"):  

The resulting dictionary is: 
{'one': 2, 'two': 2, 'three': 1}
```

<!-- grade_id: assignment-11-01-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 11.1

check_assertion(assignment_11_1)
```

Output:
```
Call assignment_11_1('one two one two three') ... OK
Call assignment_11_1('She sells sea shells on the sea shore; The shells that she sells are sea shells I'm sure. So if she sells sea shells on the sea shore, I'm sure that the shells are sea shore shells.') ... OK
Call assignment_11_1('aba aba; aba @?"') ... OK
Call assignment_11_1('AA aa Aa aA') ... OK
Call assignment_11_1('a a a a a a a a a a a a a a a a a a a a a a a a a a a a a a a a a a a a a') ... OK
Call assignment_11_1('a b a c a b a d a b a c a b a') ... OK
Call assignment_11_1('dsgsdf ' sdfg gdf ghgh sdfg sdh ghhfgjgj sdfg fhdh dfh dfh dfh') ... OK
Call assignment_11_1('word') ... OK
Call assignment_11_1('Lady of luck come out of your hiden course bless your light upon me as the light of the moon shines above and in the light of luck will be blessed i, when the moon is next to be full') ... OK
Call assignment_11_1('As I prick this candle, I prick at thee Broken hearts unhappy be May you part another day Soon to go your separ
```

<br />
<br />
<br />
<br />

# 11.2 Dictionary of synonyms

The function `assignment_11_2(synonyms, word)` takes two parameters: `synonyms` is a dictionary with synonyms. `synonyms` is built up by entering word pairs. `word` is the word that the synonym is requested for. 

In the test cell, you enter synonyms by entering **exactly** two words on a single line. The test cell will create a dictionary `synonyms` of the pairs for you.

As last line of the input of the test cell, enter **a single word** - this will count as the `word` for which we are looking for a synonym.

**Please note** that `word` can be both the key *and* the value of the dictionary.

## Example input
```python
Hello Hi
Bye Goodbye
List Array
Goodbye
```

## Example return value
```python
'Bye'
```

<!-- grade_id: assignment-11-02-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 11.2

def assignment_11_2(synonyms, word):
    for first_word, synonym in synonyms.items(): 
        if word == first_word: 
            return synonym
        elif word == synonym: 
            return first_word
```

```python
# Test for assignment 11.2

test_cell_11_2(assignment_11_2)
```

Output:
```
Please enter 2 words for a synonym or 1 word as lookup (or empty for default):  

Using default:
Hello Hi
Bye Goodbye
List Array
Goodbye

Your function returned: 'Bye'
```

<!-- grade_id: assignment-11-02-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 11.2

check_assertion(assignment_11_2)
```

Output:
```
Call assignment_11_2({'Hello': 'Hi', 'Bye': 'Goodbye', 'List': 'Array'}, 'Goodbye') ... OK
Call assignment_11_2({'beep': 'Car'}, 'Car') ... OK
Call assignment_11_2({'1234': '4321'}, '1234') ... OK
Call assignment_11_2({'a': '1', 'b': '2', 'c': '3', 'd': '4', 'e': '5', 'f': '6', 'g': '7', 'h': '8', 'i': '9', 'j': '10'}, 'e') ... OK
Call assignment_11_2({'a': '1', 'b': '2', 'c': '3', 'd': '4', 'e': '5', 'f': '6', 'g': '7', 'h': '8', 'i': '9', 'j': '10'}, '5') ... OK
Call assignment_11_2({'95UUSA4P3K': '140YAZQF1W2', 'JIGCEL82Y3BXJ': '277WMID9GFAWGC', 'X36U1924TQ': 'GUDPNJ60NR3A', 'TR7JOPXQ0JK64CM': 'O6BJSY7MHMAI8S', '17HZOQO20CDQ1UG': 'VVIPAVTVY202R', '8T1A0XRRKKXEV': '7L2UKWCJ917G1Q', 'WLPOJYRN0G3MC': '5QWJPM4FG58', 'Z8TN35IQAF': 'GXUPQY7I3AJCX3', '5F382242HD25Q': '3IV9JL29DNP64', 'KZWNAOSV8H8APVF': 'LDAIALO4YMXE7'}, 'WLPOJYRN0G3MC') ... OK

Well done! You seem to have solved assignment_11_2!
```

<br />
<br />
<br />
<br />

# 11.3 Shopping items

At the end of the week, a grocery shop owner wants to create an overview of the total number of fruit items he has sold. From his administration, he gets an overview of the fruit items he has sold to individual customers. This will be fed into the function as a *list of tuples*:
```
groceries = [('Oranges', 10), ('Oranges', 5), ('Kiwis', 9), ('Kiwis', 8), ('Oranges', 1)]
```
to `assignment_11_3(groceries)`. 

Finalize the function in such a way that it returns a **dictionary** with the fruit types as keys and the total number sold for of each of the fruit types as values.

## Example input
```python
Oranges 10
Oranges 5
Kiwis 9
Kiwis 8
Oranges 1
```

## Example return value
```python
{'Oranges': 16, 'Kiwis': 17}
```

<!-- grade_id: assignment-11-03-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 11.3

def assignment_11_3(groceries):
    result = {}
    for product, amount in groceries: 
        if product in result.keys(): 
            result[product] += amount 
        else: 
            result[product] = amount
            
    return result
```

```python
# Test for assignment 11.3

test_cell_11_3(assignment_11_3)
```

Output:
```
Please enter combination 1 of a fruit type with the number sold (empty = default):  

Using default:  [('Oranges', 10), ('Oranges', 5), ('Kiwis', 9), ('Kiwis', 8), ('Oranges', 1)]

Your function returned: {'Oranges': 16, 'Kiwis': 17}
```

<!-- grade_id: assignment-11-03-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 11.3

check_assertion(assignment_11_3)
```

Output:
```
Call assignment_11_3((('Oranges', 10), ('Oranges', 5), ('Kiwis', 9), ('Kiwis', 8), ('Oranges', 1), ('Oranges', 16), ('Kiwis', 17))) ... OK
Call assignment_11_3((('Oranges', 10), ('Oranges', 3), ('Kiwis', 19), ('Kiwis', 2), ('Oranges', 7), ('Oranges', 2), ('Kiwis', 6), ('Kiwis', 10), ('Oranges', 11), ('Oranges', 5), ('Kiwis', 3), ('Kiwis', 12), ('Oranges', 13), ('Oranges', 51), ('Kiwis', 52))) ... OK
Call assignment_11_3((('Kiwis', 1), ('Kiwis', 1))) ... OK
Call assignment_11_3((('Apples', 2), ('Bananas', 1), ('Pears', 100000), ('Apples', 1), ('Apples', 1), ('Apples', 0), ('Bananas', 100), ('Bananas', 101), ('Apples', 4), ('Pears', 100000))) ... OK
Call assignment_11_3((('Kiwis', 1), ('Kiwis', 1), ('Kiwis', 1), ('Kiwis', 1), ('Kiwis', 1), ('Kiwis', 1), ('Kiwis', 1), ('Kiwis', 1), ('Kiwis', 1), ('Kiwis', 1), ('Kiwis', 1), ('Kiwis', 1), ('Kiwis', 1), ('Kiwis', 1), ('Kiwis', 1), ('Kiwis', 1), ('Kiwis', 1), ('Kiwis', 1), ('Kiwis', 1), ('Kiwis', 1), ('Kiwis', 1), ('Kiwis', 1), ('Kiwis', 1), 
```

<br />
<br />
<br />
<br />

# 11.4 The most frequent word

Given a multiline text as list of strings `text`, return the word in the text that occurs most often. If there are many such words, print the one that appears first in alphabetical order.

In the test cell, you can enter lines of text until you enter an empty line.

## Example input
```python
apple orange banana
banana orange
```

## Example return value
```python
'banana'
```


<!-- grade_id: assignment-11-04-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 11.4

def assignment_11_4(text):
    words = []
    for item in text: 
        splitting = item.split() 
        words.extend(splitting) 
    
    words_unique = set(words)
    word_count = {word: words.count(word) for word in words_unique}
    
    maximum_occurrence = max(word_count.values())
    
    frequent_words = []
    
    for word in word_count.keys():
        if word_count[word] == maximum_occurrence: 
            frequent_words.append(word)
        
    frequent_words = sorted(frequent_words)
    
    return frequent_words[0]
        
        

```

```python
# Test for assignment 11.4

test_cell_11_4(assignment_11_4)
```

Output:
```
Please enter a string with words (empty line is default):  

Using default: ['apple orange banana', 'banana orange']

The word that occurs most often is: 
'banana'
```

<!-- grade_id: assignment-11-04-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 11.4

check_assertion(assignment_11_4)
```

Output:
```
Call assignment_11_4(('apple orange banana banana orange',)) ... OK
Call assignment_11_4(('q w e r t y u i o p', 'a s d f g h j k l', 'z x c v b n m')) ... OK
Call assignment_11_4(('vqcg vqcg vqcg vqcg vqcg vqcg vqcg', 'vqcg vqcg', 'vqcg')) ... OK
Call assignment_11_4(('taia ikm ikm ikm taia taia taia', 'ikm ikm ikm')) ... OK
Call assignment_11_4(('zc gdwbsv uwsz j bqz cytzpsdpl q zczwrq e zvnisf cytzpsdpl gdwbsv cytzpsdpl e bqz bqz zczwrq gdwbsv gdwbsv e',)) ... OK
Call assignment_11_4(('iawkvxlhqwjm rwcahhaesy kbfahrmup huq w rzllq pq ufszcjb', 'wmjdbhcwhxzlhlf', 'jeeyfhbsauapwnpmskgh rzllq rzllq wmjdbhcwhxzlhlf rzllq rzllq jeeyfhbsauapwnpmskgh wmjdbhcwhxzlhlf rzllq jeeyfhbsauapwnpmskgh jeeyfhbsauapwnpmskgh')) ... OK
Call assignment_11_4(('enagtgufzhnmyzkf kaphgquvoanw tkvj sylaejugfsv vuszbsdsv kmyyed rnhcdrpx b ejqa vykiu fjmdqlgxmameddhiytv yy owmldhdcjnoi durqhg ftjnehww aq ou tqxlcnxdpnmo ssdhbcdzzp twjdcvcojgtj',)) ... OK
Call assignment_11_4(('Death there mirth way the noisy m
```

<br />
<br />
<br />
<br />

# 11.5 Access rights

On many computers, every file has three types of "permissions", which can be set in any combination:
* read `R`
* write `W`
* execute `X`

The function `assignment_11_5(files, actions)` takes a list of files called `files`. Each item of the list is a tuple. The first element is the filename, the following elements are any combination of `R`, `W` and `X`. For example:
```python
[('info.csv', 'R'), ('job.sh', 'R', 'X'), ('thesis.doc', 'R', 'W')]
```

`actions` is a list of actions that need to be tested. This is also a list of tuples. The first element in the tuple is a keyword for the preferred action 'read', 'write' or 'execute'. The second element is the filename. For example:
```python
[('read', 'info.csv'), ('write', 'job.sh'), ('execute', 'thesis.doc')]
```

The function `assignment_11_5(files, actions)` should return a list of the same length as `actions` and contain the word `OK` if the test action succeeded or `Access denied` if it failed. For the example above:
```python
['OK', 'Access denied', 'Access denied']
```

In the test cell, first type in the files with their respective actions `R`, `W` and/or `X` separated by spaces, followed by an empty line. 

After the empty line, type in the test actions. First the action requested `read`, `write` or `execute` followed by the filename on which to apply the action.

**Hints** 
1. Inside your function, initialize a dictionary `permissions` as follows:
```python
permissions = {'R': [], 'W': [], 'X': []}
```
1. Loop through your `files` and append the filename to each corresponding permission list
1. Then, loop through your `actions` and check whether the filename is on the permissions list for that file


## Example input
```python
info.txt R
job.sh R X
thesis.doc R W

read info.txt
write job.sh
execute thesis.doc
```

## Example return value
```python
['OK', 'Access denied', 'Access denied']
```

<!-- grade_id: assignment-11-05-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 11.5

def assignment_11_5(files, actions):
    result = []
    files_permissions = {"R": [], "W": [], "X": []} 
    for file in files:
        for permission in file[1:]: 
            files_permissions[permission].append(file[0])
        
    for action, file_name in actions: 
        if action == "read": 
            if file_name in files_permissions["R"]:
                result.append("OK")
            else:
                result.append("Access denied")
        if action == "write": 
            if file_name in files_permissions["W"]:
                result.append("OK")
            else:
                result.append("Access denied")
        if action == "execute": 
            if file_name in files_permissions["X"]:
                result.append("OK")
            else:
                result.append("Access denied")
    return result
```

```python
# Test for assignment 11.5

test_cell_11_5(assignment_11_5)
```

Output:
```
Please enter Enter 1st filename with any combination of R, W and X separated by spaces. Empty for default:  

Using default files: [('info.txt', 'R'), ('job.sh', 'R', 'X'), ('thesis.doc', 'R', 'W')]
Using default actions: [('read', 'info.txt'), ('write', 'job.sh'), ('execute', 'thesis.doc')]

The result for the actions [('read', 'info.txt'), ('write', 'job.sh'), ('execute', 'thesis.doc')] are:
['OK', 'Access denied', 'Access denied']
```

<!-- grade_id: assignment-11-05-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 11.5

check_assertion(assignment_11_5)
```

Output:
```
Call assignment_11_5([('helloworld.exe', 'R', 'X'), ('pinglog', 'W', 'R'), ('nya', 'R'), ('goodluck', 'X', 'W', 'R')], [('read', 'nya'), ('write', 'helloworld.exe'), ('execute', 'nya'), ('read', 'pinglog'), ('write', 'pinglog')]) ... OK
Call assignment_11_5([('abacaba', 'X')], [('read', 'abacaba'), ('write', 'abacaba'), ('execute', 'abacaba')]) ... OK
Call assignment_11_5([('tmp_909925047', 'W', 'X', 'R')], [('execute', 'tmp_909925047'), ('read', 'tmp_909925047'), ('write', 'tmp_909925047'), ('read', 'tmp_909925047'), ('execute', 'tmp_909925047'), ('execute', 'tmp_909925047'), ('read', 'tmp_909925047')]) ... OK
Call assignment_11_5([('tmp_1017722015', 'W'), ('tmp_897110090', 'X', 'W', 'R'), ('tmp_651548400', 'W', 'X'), ('tmp_422551574', 'X', 'R', 'W'), ('tmp_477658548', 'W')], [('write', 'tmp_897110090')]) ... OK
Call assignment_11_5([('tmp_584361681', 'R', 'X'), ('tmp_70361076', 'X')], [('read', 'tmp_70361076'), ('write', 'tmp_70361076'), ('write', 'tmp_70361076')]) ... OK
Call assign
```

<br />
<br />
<br />
<br />

# 11.6 Countries and cities

For this assignment, the function `assignment_11_6(cities, requested)` takes two lists of strings. `cities` contains the strings in which the first word represents a country, while the following words are the cities that are located in that particular country. `requested` is a list of single city names for which the country needs to be looked up in your function. 

Your function should return a list, having the same length as `cities`, with the country name for each city.

In the test cell, the cities per country are given on a single line. The first word on the line is the name of the country, each next word is the name of a city. For this assignment, only use single-word countries and city names, and make sure that you don't have the same city occur in multiple countries. The entry of countries and cities is ended with an empty line. After that, enter single city names to be looked up.

## Example input
```
USA Boston Pittsburgh Washington Seattle
UK London Edinburgh Cardiff Belfast

Cardiff
Seattle
London
```

## Example return value
```python
['UK', 'USA', 'UK']
```

<!-- grade_id: assignment-11-06-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 11.6

def assignment_11_6(cities, requested):
    result= []
    
    
    country_city_dict = {item.split(maxsplit = 1)[0]: item.split()[1:] for item in cities}

    for city in requested:
        for key in country_city_dict: 
            if city in country_city_dict[key]:
                result.append(key)
                
    
    return result 
        
    
```

```python
# Test for assignment 11.6

test_cell_11_6(assignment_11_6)
```

Output:
```
Please enter first country, followed by a number of city names (empty for default):  

Using default cities: ['USA Boston Pittsburgh Washington Seattle', 'UK London Edinburgh Cardiff Belfast']
Using default requested: ['Cardiff', 'Seattle', 'London']

The result for cities ['Cardiff', 'Seattle', 'London'] are:
['UK', 'USA', 'UK']
```

<!-- grade_id: assignment-11-06-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 11.6

check_assertion(assignment_11_6)
```

Output:
```
Call assignment_11_6(['USA Boston Pittsburgh Washington Seattle', 'UK London Edinburgh Cardiff Belfast'], ['Cardiff', 'Seattle', 'London']) ... OK
Call assignment_11_6(['Russia Moscow Petersburg Novgorod Kaluga', 'Ukraine Kiev Donetsk Odessa'], ['Odessa', 'Moscow', 'Novgorod']) ... OK
Call assignment_11_6(['A B', 'C D', 'E F', 'G H', 'I J'], ['J', 'H', 'F', 'D', 'B']) ... OK
Call assignment_11_6(['A B', 'C D E', 'F G H I', 'J K L M N', 'O P Q R S T'], ['T', 'S', 'N', 'R', 'M', 'I', 'Q', 'L', 'H', 'E', 'P', 'K', 'G', 'D', 'B']) ... OK

Well done! You seem to have solved assignment_11_6!
```

<br />
<br />
<br />
<br />

# 11.7 Frequency analysis

Given a list of strings called `text`, finalize `assignment_11_7(text)` in such a way that it returns a list with all words encountered in `text`, sorted in descending order according to their number of occurrences. All words with the same number of occurrences should be printed in alphabetical order.

**Hint** After you create a dictionary of the words and their numbers of occurrences, you would like to sort it according to the occurrences. This can be achieved if you create a list whose elements are tuples of two elements: first the occurrences of a word and then the word itself. 

This works because the standard list `sort()` will handle a list of tuple by comparing the first elements of the tuples, and if those are equal comparing the second element, and so on. A trick to sort in *descending* order, is to multiply the count of words by $-1$.

For example:
```python
my_list = [(-2, 'hi'), (-1, 'why'), (-3, 'is'), (-1, 'what')]
my_list.sort()
# Now the list is [(-3, 'is'), (-2, 'hi'), (-1, 'what'), (-1, 'why')]
```

Alternatively, you can use the `reverse` keyword-argument of the `sort()` function:
```python
my_list = [(2, 'hi'), (1, 'why'), (3, 'is'), (1, 'what')]
my_list.sort(reverse=True)
# Now the list is [(3, 'is'), (2, 'hi'), (1, 'what'), (1, 'why')]
``` 


## Example input
```
hi
hi
what is your name
my name is bond
james bond
my name is damme
van damme
claude van damme
jean claude van damme
```

## Example return value
```python
[ 'damme', 'is', 'name', 'van', 'bond', 'claude', 'hi', 'my', 'james', 'jean', 'what', 'your']
```

<!-- grade_id: assignment-11-07-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 11.7

def assignment_11_7(text):
    text = [word for index in range(len(text)) for word in text[index].split() ]
    unique_words = set(text) 
    word_count = {word: -text.count(word) for word in unique_words} 
    
    pairs = [(occurrence, word) for word, occurrence in word_count.items()]
    pairs.sort()
    result = [word[1] for word in pairs]
    
    return result
    
```

```python
# Test for assignment 11.7

test_cell_11_7(assignment_11_7)
```

Output:
```
Please enter a string (empty to use default):  

Using default: ['hi', 'hi', 'what is your name', 'my name is bond', 'james bond', 'my name is damme', 'van damme', 'claude van damme', 'jean claude van damme']

The resulting list is: 
['damme', 'is', 'name', 'van', 'bond', 'claude', 'hi', 'my', 'james', 'jean', 'what', 'your']
```

<!-- grade_id: assignment-11-07-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 11.7

check_assertion(assignment_11_7)
```

Output:
```
Call assignment_11_7(['hi', 'hi', 'what is your name', 'my name is bond', 'james bond', 'my name is damme', 'van damme', 'claude van damme', 'jean claude van damme']) ... OK
Call assignment_11_7(['ai ai ai ai ai ai ai ai ai ai']) ... OK
Call assignment_11_7(['iovjxotfvt h h iovjxotfvt h iovjxotfvt iovjxotfvt h', 'h iovjxotfvt']) ... OK
Call assignment_11_7(['eqlbahbovv pzpumhz mlcgtbbnfr', 'axsjbontg emojwajtfi', 'basdu nzyh wpyirkmz xxkmrr blnlqwcpur eqlbahbovv nzyh mlcgtbbnfr', 'mlcgtbbnfr axsjbontg mlcgtbbnfr mlcgtbbnfr mlcgtbbnfr blnlqwcpur mlcgtbbnfr']) ... OK
Call assignment_11_7(['That thou hast her it is not all my grief,', 'And yet it may be said I loved her dearly,', 'That she hath thee is of my wailing chief,', 'A loss in love that touches me more nearly.', 'Loving offenders thus I will excuse ye,', "Thou dost love her, because thou know'st I love her,", 'And for my sake even so doth she abuse me,', "Suff'ring my friend for my sake to approve her.", "If I lose thee, my loss 
```

<br />
<br />
<br />
<br />

# 11.8 English-Latin dictionary

One day, going through old books in the attic, a student called Bob found an English-Latin dictionary. By that time he spoke English fluently, and his dream was to learn Latin. So finding the dictionary was a stroke of luck.

Unfortunately, full-fledged language studying requires also another type of dictionary: Latin-English. Since Bob did not posess one of those, he decided to make that second dictionary using the first one.

As you know, the dictionary consists of words, each of which contains multiple translations. For each Latin word that appears anywhere in the dictionary, Bob has to find all of its translations (that is, all English words, for which our Latin word is among its translations), and take those as the translations of this Latin word.

Help Bob to create a Latin-English dictionary.

Each entry is contained on a separate line, which contains first the English word, then the translations of this English word in Latin. All the words consist only of lowercase English letters.

Return a list of strings with the corresponding Latin-English dictionary. Make sure that the first words (Latin) are sorted alphabetically and also that the English translation(s) are also sorted alphabetically.

## Example input
```python
apple malum pomum popula
fruit baca bacca popum
punishment malum multa
```

## Example return value
```python
['baca fruit', 'bacca fruit', 'malum apple punishment', 'multa punishment', \
 'pomum apple', 'popula apple', 'popum fruit']
```

which corresponds to:

```
baca fruit
bacca fruit
malum apple punishment
multa punishment
pomum apple
popula apple
popum fruit
```

<!-- grade_id: assignment-11-08-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 11.8

def assignment_11_8(english_latin):
    list_of_words = [string.split() for string in english_latin] 
    dictionary = {latin_word: [] for words in list_of_words for latin_word in words[1:]}
    
    for word in dictionary.keys(): 
        for step in list_of_words: 
            if word in step[1:]: 
                dictionary[word].append(step[0])
    
    dictionary = {k: sorted(v) for k, v in sorted(dictionary.items())}
    result = [] 
    
    for k, v in dictionary.items(): 
        string = k + " " + " ".join(v)
        result.append(string)
    
    return result
```

```python
# Test for assignment 11.8

test_cell_11_8(assignment_11_8)
```

Output:
```
Please enter 1st English word followed by Latin words (empty line for default):  

Using default: ['apple malum pomum popula', 'fruit baca bacca popum', 'punishment malum multa']

The resulting Latin-English dictionary is: 
baca fruit
bacca fruit
malum apple punishment
multa punishment
pomum apple
popula apple
popum fruit
```

<!-- grade_id: assignment-11-08-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 11.8

check_assertion(assignment_11_8)
```

Output:
```
Call assignment_11_8(['apple malum pomum popula', 'fruit baca bacca popum', 'punishment malum multa']) ... OK
Call assignment_11_8(['school schola']) ... OK
Call assignment_11_8(['greet empfangen willkommen', 'silicon silicon', 'welcome willkommen']) ... OK
Call assignment_11_8(['mef mqax xkr', 'wemu mqax']) ... OK
Call assignment_11_8(['gfk whao', 'icfx nil whao', 'utzs brh']) ... OK
Call assignment_11_8(['casy moxd oosz pfh wfp', 'cgsy pfh', 'dlx fma lyq', 'vja fma lyq moxd oosz pfh wfp xrwv', 'zfq fma lyq tnqq xrwv']) ... OK
Call assignment_11_8(['drxpeycnkp fgoezltv fgqedltc fgqezltc fsqezltc tgqehltc', 'oaxpeycnkp fgqedltc pgqezltu tgqehltc', 'yaxpefcnkr fgqedltc fgqezljc fgqezlqc fgqezltc tgqehltc', 'yaxpeycnkp fgqedltc fgqezlqc fsqezltc', 'yaxpeycnks fgqedltc fsqezltc pgqezltu', 'yaxteyckp fgqedltc fgqezljc fgqezlqc fsqezltc pgqezltu']) ... OK
Call assignment_11_8(['fnpa kbyip', 'kbyip kbyip ojawe ysye', 'ojawe fnpa kbyip', 'uyjge kbyip', 'ysye ojawe']) ... OK
Call assignment_11
```