# 7. Lists

This Jupyter Notebook contains the [problems, or assignments, for Lesson 7](https://snakify.org/en/lessons/lists/problems/), gratefully imported from the online course on [Snakify](https://snakify.org). If you get stuck, you can either ask the teacher and the teaching assistents, or skim through the [theory](https://snakify.org/en/lessons/lists/) and [explanation of the steps](https://snakify.org/en/lessons/lists/steps/1/) on Snakify.

Some of these assignments can be solved with "comprehensions" -- which, as mentioned during the lecture, are not compulsory material. So you can try those if you have looked at them, but you can always fall back on using `for` loops.

```python
# Please run the contents of this cell.
# Select this cell and press Ctrl-Enter (or Command-Enter on mac).

from assertions import *
from test_cells import *
```

Output:
```
Assertions imported OK
Test cells imported OK
```

<br />
<br />
<br />
<br />

# 7.1 Even indices

Given a list of integers, find all the list elements with an even index and return them in a new list, *i.e*:
```python
[A[0], A[2], A[4], ... ]
```

The input list is created by entering a single line with numbers separated by spaces in the test cell.  
The test cell will take care of converting the input to a list of **integers** for you, which is then passed to your function.

## Example input

```python
1 2 3 4 5
```

## Example return value

```python
[1, 3, 5]
```

<!-- grade_id: assignment-07-01-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 7.1

def assignment_7_1(numbers):
    return numbers[0::2]
```

```python
# Test for assignment 7.1

test_cell_7_1(assignment_7_1)
```

Output:
```
Please enter a series of integer numbers, separated by spaces (Press Enter for default values 1 2 3 4 5):  

The elements in the list with even indices in [1, 2, 3, 4, 5] are:
[1, 3, 5]
```

<!-- grade_id: assignment-07-01-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 7.1

check_assertion(assignment_7_1)
```

Output:
```
Call assignment_7_1([1, 2, 3, 4, 5]) ... OK
Call assignment_7_1([4, 5, 3, 4, 2, 3]) ... OK
Call assignment_7_1([9, 4, 5, 2, 3]) ... OK
Call assignment_7_1([6]) ... OK
Call assignment_7_1([7, 8]) ... OK
Call assignment_7_1([1, 2, 3]) ... OK
Call assignment_7_1([90, 45, 3, 43]) ... OK
Call assignment_7_1([40, 64, -80, -98, -68, 56, 85, 87, -68, -78]) ... OK
Call assignment_7_1([-1, -11, -40, 3, -47, -37, -18, -100, 6, -29, 14, -1, -89, -20, -36, 7, 5, 78, -5, 85]) ... OK

Well done! You seem to have solved assignment_7_1!
```

<br />
<br />
<br />
<br />

# 7.2 Even elements

Given a list of integers, find and return all elements that are an even number.  
In this case use a for-loop that iterates over the items in list, and not over its indices! That is, don't use `range()`!  

The input list is created by entering a single line with numbers separated by spaces in the test cell.  
The test cell will take care of converting the input to a list of **integers** for you, which is then passed to your function.

## Example input
```python
1 2 2 3 3 3 4
```

## Example return value
```python
[2, 2, 4]
```

<!-- grade_id: assignment-07-02-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 7.2

def assignment_7_2(numbers):
    even_numbers = []
    for number in numbers: 
        if number % 2 == 0: 
            even_numbers.append(number)
    return even_numbers
        
    
```

```python
# Test for assignment 7.2

test_cell_7_2(assignment_7_2)
```

Output:
```
Please enter a series of integer numbers, separated by spaces (Press Enter for default values 1 2 2 3 3 3 4):  

The even elements in the list [1, 2, 2, 3, 3, 3, 4] are:
[2, 2, 4]
```

<!-- grade_id: assignment-07-02-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 7.2

check_assertion(assignment_7_2)
```

Output:
```
Call assignment_7_2([1, 2, 2, 3, 3, 3, 4]) ... OK
Call assignment_7_2([1, 2, 3, 4, 5]) ... OK
Call assignment_7_2([2, 4, 6, 8]) ... OK
Call assignment_7_2([1, 1, 1, 1, 1, 1, 2]) ... OK
Call assignment_7_2([44, -51, 32, -60, -7, -67, 62, 91, 26, -14]) ... OK
Call assignment_7_2([32, -95, -78, -15, -71, -73, -83, 53, -91, -82, -71, -52, -60, -18, 12, 8, 98, -92, -83, 7]) ... OK
Call assignment_7_2([39, 55, -88, 56, -39, -92, -30, 54, 15, 94, 27, -87, -85, -53, -44, -87, 33, 38, -22, 56]) ... OK

Well done! You seem to have solved assignment_7_2!
```

<br />
<br />
<br />
<br />

# 7.3 Greater than previous

Given a list of integers, find and return all the elements that are greater than the previous element.  
If there are no value greater than their previous element, return an empty list.
The first item of the list shouldn't be considered because it doesn't have a predecessor.

The input list is created by entering a single line with numbers separated by spaces in the test cell.  
The test cell will take care of converting the input to a list of **integers** for you, which is then passed to your function.

## Example input
```python
1 5 2 4 3
```

## Example return value
```python
[5, 4]
```

<!-- grade_id: assignment-07-03-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 7.3

def assignment_7_3(numbers):
    result = []
    for index, number in enumerate(numbers[1:], start=1):
        if number > numbers[index - 1]:
            result.append(number)
    return result
               
    
```

```python
# Test for assignment 7.3

test_cell_7_3(assignment_7_3)
```

Output:
```
Please enter a series of integer numbers, separated by spaces (Press Enter for default values 1 5 2 4 3):  

The elements in the list [1, 5, 2, 4, 3] that are bigger than their predecessor are:
[5, 4]
```

<!-- grade_id: assignment-07-03-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 7.3

check_assertion(assignment_7_3)
```

Output:
```
Call assignment_7_3([1, 5, 2, 4, 3]) ... OK
Call assignment_7_3([1, 2, 3, 4, 5]) ... OK
Call assignment_7_3([5, 4, 3, 2, 1]) ... OK
Call assignment_7_3([1, 5, 1, 5, 1]) ... OK
Call assignment_7_3([5, 1, 5, 1, 5]) ... OK
Call assignment_7_3([5, 5, 5, 5, 5]) ... OK
Call assignment_7_3([1, 1, 1, 5, 1]) ... OK
Call assignment_7_3([345354]) ... OK
Call assignment_7_3([1, 465]) ... OK
Call assignment_7_3([4, -54643]) ... OK
Call assignment_7_3([2147483647, 0, 1, 2, 3]) ... OK
Call assignment_7_3([1, 2, 3, 4, -2147483648]) ... OK
Call assignment_7_3([-9, 29, -100, 64, 26, 73, -96, 28, -92, 11, -14, -86, -54, -67]) ... OK

Well done! You seem to have solved assignment_7_3!
```

<br />
<br />
<br />
<br />

# 7.4 Neighbors of the same sign

Given a list of integers, find and return the first adjacent elements which have the same sign.  
Return the pair as a `tuple`. If there is no such pair, return `None` (default when nothing is returned from a function).

The input list is created by entering a single line with numbers separated by spaces in the test cell.  
The test cell will take care of converting the input to a list of **integers** for you, which is then passed to your function.

You may assume that none of the input numbers are zero.

## Example input
```python
-1 2 3 -1 -2
```

## Example return value
```python
(2, 3)
```

<!-- grade_id: assignment-07-04-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 7.4

def assignment_7_4(numbers):
    result = [numbers[0]]
    for number in numbers[1:]: 
        result.append(number)
        if result[0] * result[1] > 0:
            return tuple(result)
        else:
            del result[0]

```

```python
# Test for assignment 7.4

test_cell_7_4(assignment_7_4)
```

Output:
```
Please enter a series of integer numbers, separated by spaces (Press Enter for default values -1 2 3 -1 -2):  

The first pair with the same sign in the list [-1, 2, 3, -1, -2] (if any) is:
(2, 3)
```

<!-- grade_id: assignment-07-04-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 7.4

check_assertion(assignment_7_4)
```

Output:
```
Call assignment_7_4([-1, 2, 3, -1, -2]) ... OK
Call assignment_7_4([1, -3, 4, -2, 1]) ... OK
Call assignment_7_4([1, 2, -3, -4, -5]) ... OK
Call assignment_7_4([1, -2, 3, -4, -5]) ... OK
Call assignment_7_4([1, -2, 3, -4, 5, 6]) ... OK
Call assignment_7_4([-1, 2, -3, 4, -5, 6]) ... OK
Call assignment_7_4([1, -2, 3, -4, 5, -6]) ... OK
Call assignment_7_4([1, -2, 3, -4]) ... OK
Call assignment_7_4([-1, 2, -3, 4]) ... OK
Call assignment_7_4([1, 2, 3, 4]) ... OK
Call assignment_7_4([-1, -2, -3, -4]) ... OK
Call assignment_7_4([1, -1, 1, -1, -1, 1, -1, 1, -1, 1]) ... OK
Call assignment_7_4([1, -1, 1, -1, 1, 1, -1, 1, -1, 1]) ... OK
Call assignment_7_4([1, -1, 1, -1, 1]) ... OK
Call assignment_7_4([-1, 1, -1, 1, -1]) ... OK

Well done! You seem to have solved assignment_7_4!
```

<br />
<br />
<br />
<br />

# 7.5 Greater than neighbors

Given a list of integers, determine and return how many of the elements are greater than both of their neighbours (i.e. greater than its predecessor as well as its successor).

The first and the last items of the list shouldn't be considered because they don't have two neighbours.

The input list is created by entering a single line with numbers separated by spaces in the test cell.  
The test cell will take care of converting the input to a list of **integers** for you, which is then passed to your function.

## Example input
```python
1 2 5 4 3
```

## Example return value
```python
1
```

<!-- grade_id: assignment-07-05-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 7.5

def assignment_7_5(numbers):
    result = []
    for pos, number in enumerate(numbers[1:len(numbers)-1], start = 1): 
        if number > max(numbers[pos - 1], numbers[pos + 1]):
            result.append(number)
    return len(result)
```

```python
## Test for assignment 7.5

test_cell_7_5(assignment_7_5)
```

Output:
```
Please enter a series of integer numbers, separated by spaces (Press Enter for default values 1 2 5 4 3):  

The number of elements in [1, 2, 5, 4, 3] that are greater than both neighbors is:
1
```

<!-- grade_id: assignment-07-05-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 7.5

check_assertion(assignment_7_5)
```

Output:
```
Call assignment_7_5([1, 2, 3, 4, 5]) ... OK
Call assignment_7_5([5, 4, 3, 2, 1]) ... OK
Call assignment_7_5([1, 5, 1, 5, 1]) ... OK
Call assignment_7_5([5, 1, 5, 1, 5]) ... OK
Call assignment_7_5([5, 5, 5, 5, 5]) ... OK
Call assignment_7_5([1, 1, 1, 5, 1]) ... OK
Call assignment_7_5([345354]) ... OK
Call assignment_7_5([1, 465]) ... OK
Call assignment_7_5([4, -54643]) ... OK
Call assignment_7_5([2147483647, 0, 1, 2, 3]) ... OK
Call assignment_7_5([1, 2, 3, 4, -2147483648]) ... OK
Call assignment_7_5([-9, 29, -100, 64, 26, 73, -96, 28, -92, 11, -14, -86, -54, -67]) ... OK
Call assignment_7_5([2147483647, 0, 1, 0, 2147483647]) ... OK

Well done! You seem to have solved assignment_7_5!
```

<br />
<br />
<br />
<br />

# 7.6 The largest element

Given a list of integers. Determine the element in the list with the largest value.  
Return a tuple with the value of the largest element and the corresponding index.  
If the highest element is not unique, return the index of the first instance.

The input list is created by entering a single line with numbers separated by spaces in the test cell.  
The test cell will take care of converting the input to a list of **integers** for you, which is then passed to your function.

## Example input
```python
1 2 3 2 1
```

## Example return value
```python
(3, 2)
```

<!-- grade_id: assignment-07-06-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 7.6

def assignment_7_6(numbers):
    return (max(numbers), numbers.index(max(numbers)))
```

```python
# Test for assignment 7.6

test_cell_7_6(assignment_7_6)
```

Output:
```
Please enter a series of integer numbers, separated by spaces (Press Enter for default values 1 2 3 2 1):  

The largest element and its index in [1, 2, 3, 2, 1] are:
(3, 2)
```

<!-- grade_id: assignment-07-06-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 7.6

check_assertion(assignment_7_6)
```

Output:
```
Call assignment_7_6([1, 2, 3, 2, 1]) ... OK
Call assignment_7_6([1, 2, 3]) ... OK
Call assignment_7_6([1, 3, 2]) ... OK
Call assignment_7_6([2, 1, 3]) ... OK
Call assignment_7_6([2, 3, 1]) ... OK
Call assignment_7_6([3, 2, 1]) ... OK
Call assignment_7_6([3, 1, 2]) ... OK
Call assignment_7_6([1, 2, 3, 4, 5]) ... OK
Call assignment_7_6([5, 4, 3, 2, 1]) ... OK
Call assignment_7_6([-1, -2, -3, -4, -5]) ... OK
Call assignment_7_6([-5, -4, -3, -2, -1]) ... OK
Call assignment_7_6([2147483647, 2147483646, 2147483645]) ... OK
Call assignment_7_6([-2147483648, -2147483647, -2147483646]) ... OK
Call assignment_7_6([-2147483648, -2147483648, -2147483648]) ... OK
Call assignment_7_6([-2147483646, -2147483647, -2147483648]) ... OK
Call assignment_7_6([-100, -100, -100, -100, -100]) ... OK
Call assignment_7_6([-2147483648, 0, 2147483647]) ... OK
Call assignment_7_6([2147483647, 0, -2147483648]) ... OK
Call assignment_7_6([-79, 92, 70, 72, 28, 45, 37, -86, -12, 0, -27, -14, -69, 63, -68, 74, 45, 49, 9
```

<br />
<br />
<br />
<br />

# 7.7 The number of distinct elements

Given a list of integers with all of its elements sorted in ascending order, determine and return how many distinct elements it contains.

The input list is created by entering a single line with numbers separated by spaces in the test cell.  
The test cell will take care of converting the input to a list of **integers** for you, which is then passed to your function.

## Example input
```python
1 2 2 3 3 3
```

## Example return value
```python
3
```

<!-- grade_id: assignment-07-07-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 7.7

def assignment_7_7(numbers):
    result = []
    for number in numbers:
        if number not in result:
            result.append(number)
    return len(result) 
```

```python
# Test for assignment 7.7
    
test_cell_7_7(assignment_7_7)
```

Output:
```
Please enter a series of integer numbers, separated by spaces (Press Enter for default values 1 2 2 3 3 3):  

The quantity of distinct elements in [1, 2, 2, 3, 3, 3] is:
3
```

<!-- grade_id: assignment-07-07-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 7.7

check_assertion(assignment_7_7)
```

Output:
```
Call assignment_7_7([1, 2, 2, 3, 3, 3]) ... OK
Call assignment_7_7([1, 2, 3, 4, 5]) ... OK
Call assignment_7_7([1, 1, 1, 1, 1]) ... OK
Call assignment_7_7([1, 1, 2, 2, 2, 3, 4, 5, 6, 7]) ... OK
Call assignment_7_7([-100, -99, -98, -97, -96]) ... OK
Call assignment_7_7([-100, -100, -98, -50, 1]) ... OK
Call assignment_7_7([-64, -45, -38, -16, -14, -7, 2, 53, 70, 80, 80]) ... OK

Well done! You seem to have solved assignment_7_7!
```

<br />
<br />
<br />
<br />

# 7.8 Swap neighbors

Given a list of integers, swap adjacent items in pairs (so swap `A[0]` with `A[1]`, swap `A[2]` with `A[3]`, etc).  
Return the resulting list. If a list has an odd number of elements, leave the last element in its place.

The input list is created by entering a single line with numbers separated by spaces in the test cell.  
The test cell will take care of converting the input to a list of **integers** for you, which is then passed to your function.

## Example input
```python
1 2 3 4 5
```

## Example return value
```python
[2, 1, 4, 3, 5]
```

<!-- grade_id: assignment-07-08-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 7.8

def assignment_7_8(numbers):
    result = []
    
    if len(numbers) % 2 == 0:
        for index in range(0, len(numbers), 2):
            result += [numbers[index + 1], numbers[index]]
    
    if len(numbers) % 2 == 1: 
        for index in range(0, len(numbers)-1, 2):
            result += [numbers[index + 1], numbers[index]]
        result.append(numbers[-1])
                    
    return result 
        
            
            
            
            
```

```python
# Test for assignment 7.8

test_cell_7_8(assignment_7_8)
```

Output:
```
Please enter a series of integer numbers, separated by spaces (Press Enter for default values 1 2 3 4 5):  

The result after swapping adjacent numbers in [1, 2, 3, 4, 5] is:
[2, 1, 4, 3, 5]
```

<!-- grade_id: assignment-07-08-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 7.8

check_assertion(assignment_7_8)
```

Output:
```
Call assignment_7_8([1, 2, 3, 4, 5]) ... OK
Call assignment_7_8([4, 5, 3, 4, 2, 3]) ... OK
Call assignment_7_8([9, 4, 5, 2, 3]) ... OK
Call assignment_7_8([6]) ... OK
Call assignment_7_8([7, 8]) ... OK
Call assignment_7_8([1, 2, 3]) ... OK
Call assignment_7_8([90, 45, 3, 43]) ... OK
Call assignment_7_8([40, 64, -80, -98, -68, 56, 85, 87, -68, -78]) ... OK
Call assignment_7_8([-1, -11, -40, 3, -47, -37, -18, -100, 6, -29, 14, -1, -89, -20, -36, 7, 5, 78, -5, 85]) ... OK
Call assignment_7_8([-46, 35, 72, -79, 38, -32, 1, 23, -84, -84, 50, -46, -57, -44, 72, 33, 91, -6, 88, 91, -12, -76, -78, 95, 63, 45, 54, 13, 58, -19, 30, 48, -98, 37, 97]) ... OK

Well done! You seem to have solved assignment_7_8!
```

<br />
<br />
<br />
<br />

# 7.9 Swap min and max

Given a list of distinct integers, swap the smallest and largest elements of this list, leaving all of the other elements in place.

Return the resulting list.

The input list is created by entering a single line with numbers separated by spaces in the test cell.  
The test cell will take care of converting the input to a list of **integers** for you, which is then passed to your function.

## Example input
```python
3 4 5 2 1
```

## Example return value
```python
[3, 4, 1, 2, 5]
```

<!-- grade_id: assignment-07-09-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 7.9

def assignment_7_9(numbers):
    max_num = max(numbers)
    min_num = min(numbers)
    
    index_max = numbers.index(max_num)
    index_min = numbers.index(min_num) 
    
    numbers[index_max] = min_num
    numbers[index_min] = max_num 
    
    return numbers
```

```python
# Test for assignment 7.9

test_cell_7_9(assignment_7_9)
```

Output:
```
Please enter a series of integer numbers, separated by spaces (Press Enter for default values 3 4 5 2 1):  

The result of swapping the minimum and maximum values in [3, 4, 5, 2, 1] is:
[3, 4, 1, 2, 5]
```

<!-- grade_id: assignment-07-09-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 7.9

check_assertion(assignment_7_9)
```

Output:
```
Call assignment_7_9([3, 4, 5, 2, 1]) ... OK
Call assignment_7_9([1, 5, 4, 3, 2]) ... OK
Call assignment_7_9([-30000, 30000]) ... OK
Call assignment_7_9([1]) ... OK
Call assignment_7_9([2147483647, -2147483648]) ... OK
Call assignment_7_9([1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 17, 16, 15, 14]) ... OK
Call assignment_7_9([1, 2, 3, 4, 5, 6, 7, 8, 9, 10]) ... OK
Call assignment_7_9([10, 9, 8, 7, 6, 5, 4, 3, 2, 1]) ... OK

Well done! You seem to have solved assignment_7_9!
```

<br />
<br />
<br />
<br />

# 7.10 The number of equal pairs

Given a list of integers, count and return how many element pairs have the same value (are equal).  
Any two elements that are equal to each other should be counted exactly once.
Note that the elements that you inspect as a pair do not have to be adjacent!

The input list is created by entering a single line with numbers separated by spaces in the test cell.  
The test cell will take care of converting the input to a list of **integers** for you, which is then passed to your function.

**Hint**
You will need to use a `for` loop inside of another `for` loop. The outer loop should cover all indices `idx1` of the list, and the inner list should cover all indices `idx2` which are greater than the current `idx1`. In this way, you will visit every distinct pair `idx1, idx2` exactly once.


## Example input
```python
1 2 3 2 3
```

## Example return value
```python
2
```

<!-- grade_id: assignment-7-10-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 7.10

def assignment_7_10(numbers):
    pairs = 0
    counter = 0
    for number in numbers[:-1]: 
        for pair in numbers[counter+1:]: 
            if number == pair: 
                pairs += 1
        counter += 1
        
    return pairs 
                
```

```python
# Test for assignment 7.10

test_cell_7_10(assignment_7_10)
```

Output:
```
Please enter a series of integer numbers, separated by spaces (Press Enter for default values 1 2 3 2 3):  

The number of equal pairs in [1, 2, 3, 2, 3] is:
2
```

<!-- grade_id: assignment-7-10-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 7.10

check_assertion(assignment_7_10)
```

Output:
```
Call assignment_7_10([1, 2, 3, 2, 3]) ... OK
Call assignment_7_10([1, 1, 1, 1, 1]) ... OK
Call assignment_7_10([1, 2, 3]) ... OK
Call assignment_7_10([1, 1, 1]) ... OK
Call assignment_7_10([2, 3]) ... OK
Call assignment_7_10([2, 2]) ... OK
Call assignment_7_10([0, 0]) ... OK
Call assignment_7_10([0, 0, 6, 1, 1, 4, 2, 1, 2, 2]) ... OK
Call assignment_7_10([0, 1, 43, 10, 13, 33, 15, 8, 18, 21]) ... OK
Call assignment_7_10([3, 2, 1, 2, 2, 4, 3, 2, 5, 3, 2]) ... OK

Well done! You seem to have solved assignment_7_10!
```

<br />
<br />
<br />
<br />

# 7.11 Unique elements

Given a list of integers, find and return the elements that appear in the list only once.  
The elements in the returned list must be in the same order in which they occur in the original list.  
Return an empty list in case there are no elements that only appear once.

The input list is created by entering a single line with numbers separated by spaces in the test cell.  
The test cell will take care of converting the input to a list of **integers** for you, which is then passed to your function.

## Example input
```python
1 2 2 3 3 3
```

## Example return value
```python
[1]
```

<!-- grade_id: assignment-7-11-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 7.11

def assignment_7_11(numbers):
    result = []
    for number in numbers: 
        if number not in result:
            counter = 0 
            for occurrence in numbers: 
                if number == occurrence:
                    counter += 1
        if counter == 1:
            result.append(number) 
    
    return result 
        
```

```python
# Test for assignment 7.11

test_cell_7_11(assignment_7_11)
```

Output:
```
Please enter a series of integer numbers, separated by spaces (Press Enter for default values 1 2 2 3 3 3):  

The numbers that only occur once in [1, 2, 2, 3, 3, 3] are:
[1]
```

<!-- grade_id: assignment-7-11-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 7.11

check_assertion(assignment_7_11)
```

Output:
```
Call assignment_7_11([1, 2, 2, 3, 3, 3]) ... OK
Call assignment_7_11([4, 3, 5, 2, 5, 1, 3, 5]) ... OK
Call assignment_7_11([6, 9, 6, 23, 12, 19, 14, 26]) ... OK
Call assignment_7_11([21, 21, 16, 20, 28, 28, 7]) ... OK
Call assignment_7_11([21]) ... OK
Call assignment_7_11([5, 2, 29, 4, 28, 8]) ... OK
Call assignment_7_11([4, 2, 9, 3, 2, 9, 7, 8, 4]) ... OK

Well done! You seem to have solved assignment_7_11!
```

<br />
<br />
<br />
<br />

# 7.12 Queens

In chess it is known to be possible to place 8 queens on an 88 chess board such that none of them can attack another.  
Given a placement of 8 queens on the board, determine if there is a pair of queens that can attach each other on the next move.

Return `False` if no queen can attack another, otherwise return `True`.

In assignment 3.9, you have been working with the movement of the queen already.  
As a refresher, the picture from that assignment is shown here again:

<img src="figures/queen_move.svg" alt="Queen move" width="50%"/>

You probably can use (some of) the logic from that assignment here.

In the test cell, the coordinates for each queen is requested as input. These are passed as a list of 8 tuples to your function. Each tuple contains the row and column index for each queen.

You will notice that it is tougher to position 8 queens on a chess board where no queen is attacking another than when one or more are attacking each other. The example is one of those fewer positions.

## Example input
```python
0 6
1 3
2 1
3 7
4 5
5 0
6 2
7 4 
```

## Example return value
```python
False
```

<!-- grade_id: assignment-7-12-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 7.12

def assignment_7_12(coordinates):
    counter = 0
    is_attacking = False
    for ref in coordinates[:-1]: 
        counter += 1
        for i in coordinates[counter:]: 
            if (ref[0] == i[0]) or (ref[1] == i[1]) or (abs(ref[1]-i[1]) == abs(ref[0]-i[0])):
                is_attacking = True
                break
        if is_attacking:
            break
    return is_attacking
                
```

```python
# Test for assignment 7.12

test_cell_7_12(assignment_7_12)
```

Output:
```
Please enter a series of coordinates of for queen 1, separated by spaces (Press Enter for default values 0 6):  
Please enter a series of coordinates of for queen 2, separated by spaces (Press Enter for default values 1 3):  
Please enter a series of coordinates of for queen 3, separated by spaces (Press Enter for default values 2 1):  
Please enter a series of coordinates of for queen 4, separated by spaces (Press Enter for default values 3 7):  
Please enter a series of coordinates of for queen 5, separated by spaces (Press Enter for default values 4 5):  
Please enter a series of coordinates of for queen 6, separated by spaces (Press Enter for default values 5 0):  
Please enter a series of coordinates of for queen 7, separated by spaces (Press Enter for default values 6 2):  
Please enter a series of coordinates of for queen 8, separated by spaces (Press Enter for default values 7 4):  

With queens on [(0, 6), (1, 3), (2, 1), (3, 7), (4, 5), (5, 0), (6, 2), (7, 4)],
no queen is at
```

<!-- grade_id: assignment-7-12-assert, points: 1, grade: true, solution: false -->
```python
check_assertion(assignment_7_12)
```

Output:
```
Call assignment_7_12([(0, 6), (1, 3), (2, 1), (3, 7), (4, 5), (5, 0), (6, 2), (7, 4)]) ... OK
Call assignment_7_12([(0, 7), (1, 6), (2, 5), (3, 4), (4, 3), (5, 2), (6, 1), (7, 0)]) ... OK
Call assignment_7_12([(0, 6), (1, 1), (2, 3), (3, 0), (4, 7), (5, 5), (6, 2), (7, 4)]) ... OK
Call assignment_7_12([(0, 1), (1, 4), (2, 6), (3, 0), (4, 2), (5, 7), (6, 5), (7, 3)]) ... OK
Call assignment_7_12([(0, 2), (1, 6), (2, 1), (3, 7), (4, 4), (5, 0), (6, 3), (7, 5)]) ... OK
Call assignment_7_12([(0, 3), (1, 1), (2, 5), (3, 2), (4, 6), (5, 0), (6, 7), (7, 4)]) ... OK
Call assignment_7_12([(0, 3), (1, 1), (2, 6), (3, 0), (4, 5), (5, 7), (6, 4), (7, 2)]) ... OK
Call assignment_7_12([(0, 3), (1, 1), (2, 6), (3, 5), (4, 0), (5, 7), (6, 4), (7, 2)]) ... OK
Call assignment_7_12([(0, 3), (1, 1), (2, 4), (3, 2), (4, 0), (5, 6), (6, 3), (7, 5)]) ... OK
Call assignment_7_12([(0, 4), (1, 1), (2, 3), (3, 0), (3, 6), (4, 2), (5, 7), (6, 5)]) ... OK
Call assignment_7_12([(0, 2), (1, 0), (2, 5), (3, 7), (4, 5)
```

<br />
<br />
<br />
<br />

# 7.13 The bowling alley

In bowling, the player starts with 10 pins at the far end of a lane.  
The objective is to knock all the pins down with two balls.  
A schematic overview is shown in the figure below.

<img src="figures/bowling.jpg" alt="Bowling pins" width="50%"/>

In our representation, the bowling pins are represented by a list of booleans (`True` or `False`).
At the start of the bowling turn, all pins are standing.  
This is represented by a list with 10 `True` values.  
The four pins in the back (from left to right) are represented in the list with indices 0, 1, 2 and 3.  
Continuing in reading direction, the pin in front will be index 9. That is, the pins will be assigned positions as follows:

```python
0 1 2 3
 4 5 6
  7 8 
   9
```

The assignment function takes a list of two tuples as input.  
Each tuple contains the indices of the pins being knocked down with the ball.

All pins standing (represented by an asterix) at the start of the round:
```python
* * * *
 * * *
  * * 
   *
```

After a bowling **turn** with the list of two tuples:
```python
[[4, 7, 9], [8,]]
```
the following pins will still be standing (dots for knocked down pins):

```python
* * * *
 . * *
  . .
   .
```

The assignment function should return a list with 10 booleans indicating which pins are still standing (`True`)<br/>and which ones are knocked down (`False`). 

The test cell will request two lines with integers (0..9). Type them in separated by spaces.

## Example input
```python
4 7 9
8
```

## Example return value
```python
[True, True, True, True, False, True, True, False, False, False]
```

Represented as:
```python
* * * *
 . * *
  . .
   .
```

<!-- grade_id: assignment-7-13-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 7.13

def assignment_7_13(turn):
    result =[True for _ in range(10)]
    
    for i in turn: 
        for k in i: 
            result[k] = False
    
    return result
```

```python
# Test for assignment 7.13

test_cell_7_13(assignment_7_13)
```

Output:
```
Please enter a series of pins knocked down in throw 1, separated by spaces (Press Enter for default values 4 7 9):  
Please enter a series of pins knocked down in throw 2, separated by spaces (Press Enter for default values 8):  
Input as list of lists: [[4, 7, 9], [8]]


You returned [True, True, True, True, False, True, True, False, False, False], representing:

* * * * 
 . * * 
  . . 
   .
```

<!-- grade_id: assignment-7-13-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 7.13

check_assertion(assignment_7_13)
```

Output:
```
Call assignment_7_13([(4, 7, 9), (8,)]) ... OK
Call assignment_7_13([(0, 1), (2, 3)]) ... OK
Call assignment_7_13([(), ()]) ... OK
Call assignment_7_13([(0, 1, 2, 3, 4), (5, 6, 7, 8, 9)]) ... OK
Call assignment_7_13([(0, 2, 4), (6, 8)]) ... OK
Call assignment_7_13([(1, 3, 5), (7, 9)]) ... OK

Well done! You seem to have solved assignment_7_13!
```