# 5. Strings

This Jupyter Notebook contains the [problems, or assignments, for Lesson 5](https://snakify.org/en/lessons/strings_str/problems/), gratefully imported from the online course on [Snakify](https://snakify.org). If you get stuck, you can either ask the teacher and the teaching assistents, or skim through the [theory](https://snakify.org/en/lessons/strings_str/) and [explanation of the steps](https://snakify.org/en/lessons/strings_str/steps/1/) on Snakify.

```python
# Please run the contents of this cell.
# Select this cell and press Ctrl-Enter (or Command-Enter on mac).

from assertions import *
from test_cells import *
```

Output:
```
Assertions imported OK
Test cells imported OK
```

<br />
<br />
<br />
<br />

# 5.1 Slices (a)

(Please note that this assignment has been split into three parts)

You are given a string. Create a result list that contains the following three items:

* The first item should contain the third character of this string.
* The second item shoud contain the second to last character of this string.
* The third item should contain the first five characters of this string.

## Example input

```python
Abrakadabra
```

## Example return value

```python
['r', 'r', 'Abrak']
```

<!-- grade_id: assignment-05-01-a-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 5.1a

def assignment_5_1a(string):
    list = [string[2], string[::-1][1], string[:5]]
    return list
```

```python
# Test for assignment 5.1a

test_cell_5_1(assignment_5_1a)
```

Output:
```
Please enter a string (Press Enter for default "Abrakadabra"):  

For string 'Abrakadabra', this results in: 
['r', 'r', 'Abrak']
```

<!-- grade_id: assignment-05-01-a-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 5.1a

check_assertion(assignment_5_1a)
```

Output:
```
Call assignment_5_1a('Abrakadabra') ... OK
Call assignment_5_1a('Hello') ... OK
Call assignment_5_1a('qwertyuiop') ... OK
Call assignment_5_1a('asdfghjklzxcv') ... OK

Well done! You seem to have solved assignment_5_1a!
```

<br />
<br />
<br />
<br />

# 5.1 Slices (b)

(Please note that this assignment has been split into three parts)

You are given a string. Create a result list that contains the following three items:
* The first item should contain all but the last two characters of this string.
* The second item should contain all the characters of this string with even indices (remember indexing starts at 0, so the characters are displayed starting with the first).
* The third item should contain all the characters of this string with odd indices (i.e. starting with the second character in the string).

## Example input

```python
Abrakadabra
```

## Example return value

```python
['Abrakadab', 'Arkdba', 'baaar']
```

<!-- grade_id: assignment-05-01-b-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 5.1b

def assignment_5_1b(string):
    list = [string[:len(string)-2], string[::2], string[1::2]]
    return list
```

```python
# Test for assignment 5.1b

test_cell_5_1(assignment_5_1b)
```

Output:
```
Please enter a string (Press Enter for default "Abrakadabra"):  

For string 'Abrakadabra', this results in: 
['Abrakadab', 'Arkdba', 'baaar']
```

<!-- grade_id: assignment-05-01-b-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 5.1b

check_assertion(assignment_5_1b)
```

Output:
```
Call assignment_5_1b('Abrakadabra') ... OK
Call assignment_5_1b('Hello') ... OK
Call assignment_5_1b('qwertyuiop') ... OK
Call assignment_5_1b('asdfghjklzxcv') ... OK

Well done! You seem to have solved assignment_5_1b!
```

<br />
<br />
<br />
<br />

# 5.1 Slices (c)

(Please note that this assignment has been split into three parts)

You are given a string. Create a result list that contains the following three items:
* The first item should contain all the characters of the string in reverse order.
* The second item should contain every second character of the string in reverse order, starting from the last one.
* The third item should contain the length of the given string (**this is an integer**).

## Example input

```python
Abrakadabra
```

## Example return value

```python
['arbadakarbA', 'abdkrA', 11]
```

<!-- grade_id: assignment-05-01-c-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 5.1c

def assignment_5_1c(string):
    list = [string[::-1], string[::-1][::2], len(string)]
    return list

```

```python
# Test for assignment 5.1c

test_cell_5_1(assignment_5_1c)
```

Output:
```
Please enter a string (Press Enter for default "Abrakadabra"):  

For string 'Abrakadabra', this results in: 
['arbadakarbA', 'abdkrA', 11]
```

<!-- grade_id: assignment-05-01-c-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 5.1c

check_assertion(assignment_5_1c)
```

Output:
```
Call assignment_5_1c('Abrakadabra') ... OK
Call assignment_5_1c('Hello') ... OK
Call assignment_5_1c('qwertyuiop') ... OK
Call assignment_5_1c('asdfghjklzxcv') ... OK

Well done! You seem to have solved assignment_5_1c!
```

<br />
<br />
<br />
<br />

# 5.2 The number of words

Given a string consisting of words separated by (single) spaces. Determine how many words it has. To solve the problem, use the method `count`.

## Example input

```python
Hello world
```

## Example return value

```python
2
```

<!-- grade_id: assignment-05-02-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 5.2

def assignment_5_2(string):
    list = string.split()
    return len(list)
```

```python
# Test for assignment 5.2

test_cell_5_2(assignment_5_2)
```

Output:
```
Please enter a string (Press Enter for default "Hello world"):  

The number of words in 'Hello world' is:
2
```

<!-- grade_id: assignment-05-02-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 5.2

check_assertion(assignment_5_2)
```

Output:
```
Call assignment_5_2('Hello world') ... OK
Call assignment_5_2('Hello') ... OK
Call assignment_5_2('q w e') ... OK
Call assignment_5_2('In the hole in the ground there lived a hobbit') ... OK
Call assignment_5_2('One two three four five') ... OK

Well done! You seem to have solved assignment_5_2!
```

<br />
<br />
<br />
<br />

# 5.3 The two halves

Given a string. Cut it into two "equal" parts (If the length of the string is odd, place the middle character in the first string, so that the first string contains one more character than the second). Now `return` a new string with the first and second halves interchanged (second half first and the first half second).

Don't use the statement `if` in this task.

## Example input

```python
Hi
```

## Example return value

```python
iH
```

<!-- grade_id: assignment-05-03-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 5.3
def assignment_5_3(string):
    middle = len(string) % 2 + len(string) // 2
    return string[middle:] + string[:middle]
    
```

```python
# Test for assignment 5.3

test_cell_5_3(assignment_5_3)
```

Output:
```
Please enter a string (Press Enter for default "Hi"):  

The two halves of 'Hi' inverted are: 
iH
```

<!-- grade_id: assignment-05-03-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 5.3

check_assertion(assignment_5_3)
```

Output:
```
Call assignment_5_3('Hi') ... OK
Call assignment_5_3('Hello') ... OK
Call assignment_5_3('Qwerty') ... OK
Call assignment_5_3('Z') ... OK
Call assignment_5_3('asdfghj') ... OK
Call assignment_5_3('asdfghjzxc') ... OK

Well done! You seem to have solved assignment_5_3!
```

<br />
<br />
<br />
<br />

# 5.4 To swap the two words

Given a string consisting of exactly two words separated by a space. Return a new string with the first and second words swapped (the second word is printed first).

This task should not use loops and `if`.

## Example input

```python
Hello, world!
```

## Example return value

```python
world! Hello,
```

<!-- grade_id: assignment-05-04-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 5.4

def assignment_5_4(string):
    list = string.split()
    return list[1] + ' ' + list[0]
```

```python
# Test for assignment 5.4

test_cell_5_4(assignment_5_4)
```

Output:
```
Please enter a string of exactly two words (Press Enter for default "Hello, world!"):  
The two words in 'Hello, world!' swapped are: 
world! Hello,
```

<!-- grade_id: assignment-05-04-assert, points: 1, grade: true, solution: false -->
```python
# Code for assignment 5.4

check_assertion(assignment_5_4)
```

Output:
```
Call assignment_5_4('Hello, world!') ... OK
Call assignment_5_4('A B') ... OK
Call assignment_5_4('Q WERRTYUIOP') ... OK
Call assignment_5_4('QWERTYUIOP M') ... OK
Call assignment_5_4('sadfahsjkldfhasjkdfhaklsjdfhjkl asdlkfhasdjklfhaslkdfjhalskdfgalsdf') ... OK

Well done! You seem to have solved assignment_5_4!
```

<br />
<br />
<br />
<br />

# 5.5 The first and last occurrence

Given a string that may or may not contain a letter of interest. Return the index location of the first and last appearance of the letter `f`, as a tuple. If the letter `f` occurs only once, then return its index. If the letter `f` does not occur, then return nothing (`None`).

## Example input

```python
comfort
```

## Example return value

```python
3
```


<!-- grade_id: assignment-05-05-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 5.5

def assignment_5_5(string):
    i = 0
    for letter in string:  
        if letter in ["f"]: 
            i +=1
    if i == 0:
        return None
    elif i == 1: 
        return string.index("f")
    else: 
        return (string.index("f"), (len(string)-1) - string[::-1].index("f"))
                
    
    
            
```

```python
# Test for assignment 5.5

test_cell_5_5(assignment_5_5)
```

Output:
```
Please enter a string (Press Enter for default "comfort"):  

The first (and last) occurrence of 'f' in 'comfort' is:
3
```

<!-- grade_id: assignment-05-05-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 5.5

check_assertion(assignment_5_5)
```

Output:
```
Call assignment_5_5('comfort') ... OK
Call assignment_5_5('office') ... OK
Call assignment_5_5('hello') ... OK
Call assignment_5_5('fffffffffffffff') ... OK
Call assignment_5_5('aaaaaaaaaaaaaaaaaaaaafaaaaaaaaaaaaaaaaaaaaaaaa') ... OK
Call assignment_5_5('aaaaaaaaaaaaaaaaaaaaaaaaffaaaaaaaaaaaaaaaaaaaa') ... OK
Call assignment_5_5('afafafafafafafa') ... OK

Well done! You seem to have solved assignment_5_5!
```

<br />
<br />
<br />
<br />

# 5.6 The second occurrence

Given a string that may or may not contain the letter of interest. Return the index location of the **second** occurrence of the letter `f`. If the string contains the letter `f` only once, then return `-1`, and if the string does not contain the letter `f`, then return `-2`.

## Example input

```python
comfort
```

## Example return value

```python
-1
```

<!-- grade_id: assignment-05-06-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 5.6

def assignment_5_6(string):
    i=0
    for letter in string: 
        if letter in ["f"]: 
            i+=1
    if i == 1: 
        return -1
    elif i == 0:
        return -2
    
    letters = list(string)
    first_appearance = letters.index("f")
    letters[first_appearance] = ' '
    return letters.index("f") 
    
```

```python
# Test for assignment 5.6

test_cell_5_6(assignment_5_6)
```

Output:
```
Please enter a string (Press Enter for default "comfort"):  

The string 'comfort' only contains a single 'f'.
```

<!-- grade_id: assignment-05-06-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 5.6

check_assertion(assignment_5_6)
```

Output:
```
Call assignment_5_6('comfort') ... OK
Call assignment_5_6('coffee') ... OK
Call assignment_5_6('qwerty') ... OK
Call assignment_5_6('f') ... OK
Call assignment_5_6('oooooooooooooooooof') ... OK
Call assignment_5_6('ff') ... OK
Call assignment_5_6('ffffffffffffffff') ... OK
Call assignment_5_6('foooooooooooooof') ... OK
Call assignment_5_6('oooooooooooooff') ... OK
Call assignment_5_6('ofofofofofofofofo') ... OK

Well done! You seem to have solved assignment_5_6!
```

<br />
<br />
<br />
<br />

# 5.7 Remove the fragment

Given a string in which the letter `h` occurs at least twice. Remove from that string the first and the last occurrence of the letter `h`, as well as all of the characters between them. Return the resulting string.

## Example input

```python
In the hole in the ground there lived a hobbit
```

## Example return value

```python
In tobbit
```

<!-- grade_id: assignment-05-07-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 5.7

def assignment_5_7(string):
    first_occ = string.index("h")
    last_occ = (len(string) - 1) - string[::-1].index("h")
    return string.replace(string[first_occ:last_occ+1], "")
    
```

```python
# Test for assignment 5.7

test_cell_5_7_8_11(assignment_5_7)
```

Output:
```
Please enter a string with at least two h's (Press Enter for default "In the hole in the ground there lived a hobbit"):  

The resulting string is: 
In tobbit
```

<!-- grade_id: assignment-05-07-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 5.7

check_assertion(assignment_5_7)
```

Output:
```
Call assignment_5_7('In the hole in the ground there lived a hobbit') ... OK
Call assignment_5_7('qwertyhasdfghzxcvb') ... OK
Call assignment_5_7('asdfghhzxcvb') ... OK
Call assignment_5_7('ahahahahaha') ... OK
Call assignment_5_7('haaaaaaaaaaaaaaaaaah') ... OK
Call assignment_5_7('hh') ... OK
Call assignment_5_7('ahaha') ... OK

Well done! You seem to have solved assignment_5_7!
```

<br />
<br />
<br />
<br />

# 5.8 Reverse the fragment

Given a string in which the letter `h` occurs at least two times, reverse the sequence of characters enclosed between the first and last appearances. Return the resulting string.

## Example input

```python
In the hole in the ground there lived a hobbit
```

## Example return value

```python
In th a devil ereht dnuorg eht ni eloh ehobbit
```

<!-- grade_id: assignment-05-08-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 5.8

def assignment_5_8(string):
    first_occ = string.index("h")
    last_occ = (len(string) - 1) - string[::-1].index("h")
    substring = string[first_occ:last_occ+1] 
    inverted_substring = substring[::-1]
    return string[:first_occ] + inverted_substring + string[last_occ+1:]
```

```python
# Test for assignment 5.8

test_cell_5_7_8_11(assignment_5_8)
```

Output:
```
Please enter a string with at least two h's (Press Enter for default "In the hole in the ground there lived a hobbit"):  

The resulting string is: 
In th a devil ereht dnuorg eht ni eloh ehobbit
```

<!-- grade_id: assignment-05-08-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 5.8

check_assertion(assignment_5_8)
```

Output:
```
Call assignment_5_8('In the hole in the ground there lived a hobbit') ... OK
Call assignment_5_8('qwertyhasdfghzxcvb') ... OK
Call assignment_5_8('asdfghhzxcvb') ... OK
Call assignment_5_8('ahbhchdheha') ... OK
Call assignment_5_8('habcdefgijklmnh') ... OK
Call assignment_5_8('hh') ... OK
Call assignment_5_8('hah') ... OK
Call assignment_5_8('habh') ... OK
Call assignment_5_8('ahcvbhs') ... OK

Well done! You seem to have solved assignment_5_8!
```

<br />
<br />
<br />
<br />

# 5.9 Replace the substring

Given a string, replace in this string all the numbers `1` by the word `one`. Return the resulting string.

## Example input

```python
1+1=2
```

## Example return value

```python
one+one=2
```

<!-- grade_id: assignment-05-09-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 5.9

def assignment_5_9(string):
    return string.replace("1", "one")
```

```python
# Test for assignment 5.9

test_cell_5_9(assignment_5_9)
```

Output:
```
Please enter a string (Press Enter for default "1+1=2"):  

The resulting string with 1s replace by the word 'one' (if any): 
one+one=2
```

<!-- grade_id: assignment-05-09-assertion, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 5.9

check_assertion(assignment_5_9)
```

Output:
```
Call assignment_5_9('1+1=2') ... OK
Call assignment_5_9('Hello, 2345678990') ... OK
Call assignment_5_9('1') ... OK
Call assignment_5_9('1111111111111111111111111111111111') ... OK
Call assignment_5_9('1213141516171819101') ... OK

Well done! You seem to have solved assignment_5_9!
```

<br />
<br />
<br />
<br />

# 5.10 Delete a character

Given a string, remove from this string all the characters `@`. Return the resulting string.

## Example input

```python
Bilbo.Baggins@bagend.hobbiton.shire.me
```

## Example return value

```python
Bilbo.Bagginsbagend.hobbiton.shire.me
```

<!-- grade_id: assignment-5-10-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 5.10

def assignment_5_10(string):
    return string.replace("@", '')
```

```python
# Test for assignment 5.10

test_cell_5_10(assignment_5_10)
```

Output:
```
Please enter a string (Press Enter for default "Bilbo.Baggins@bagend.hobbiton.shire.me"):  

The resulting string without @s: 
Bilbo.Bagginsbagend.hobbiton.shire.me
```

<!-- grade_id: assignment-5-10-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 5.10

check_assertion(assignment_5_10)
```

Output:
```
Call assignment_5_10('Bilbo.Baggins@bagend.hobbiton.shire.me') ... OK
Call assignment_5_10('dfa;sdkfj;ajva;bvna'sdasdfasdglJLHJKFHLDKJFh') ... OK
Call assignment_5_10('@') ... OK
Call assignment_5_10('@W@E@E@E@R') ... OK
Call assignment_5_10('@@@@@@@@@@') ... OK
Call assignment_5_10('H@K@M@J@M@L') ... OK
Call assignment_5_10('@@@@@F@@@@KL@@@@@J@J@J@JJJH@JJKK@JJJ@JJ') ... OK

Well done! You seem to have solved assignment_5_10!
```

<br />
<br />
<br />
<br />

# 5.11 Replace within the fragment

Given a string, replace every occurrence of the letter `h` by the letter `H`, **except** for the first and the last occurrence. Return the resulting string.

## Example input

```python
In the hole in the ground there lived a hobbit
```

## Example return value

```python
In the Hole in tHe ground tHere lived a hobbit
```

<!-- grade_id: assignment-5-11-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 5.11

def assignment_5_11(string):
    first_occ = string.index("h")
    last_occ = (len(string) - 1) - string[::-1].index("h")
    for i in range(first_occ + 1, last_occ): 
        if string[i] == "h": 
            string = string[:i] + "H" + string[i+1:]
    return string
```

```python
# Test for assignment 5.11

test_cell_5_7_8_11(assignment_5_11)
```

Output:
```
Please enter a string with at least two h's (Press Enter for default "In the hole in the ground there lived a hobbit"):  

The resulting string is: 
In the Hole in tHe ground tHere lived a hobbit
```

<!-- grade_id: assignment-5-11-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 5.11

check_assertion(assignment_5_11)
```

Output:
```
Call assignment_5_11('In the hole in the ground there lived a hobbit') ... OK
Call assignment_5_11('qwertyhahsdhfghzxcvb') ... OK
Call assignment_5_11('asdfghhzxcvb') ... OK
Call assignment_5_11('ahbhchdheha') ... OK
Call assignment_5_11('habcdefgijklmnh') ... OK
Call assignment_5_11('hh') ... OK
Call assignment_5_11('hah') ... OK
Call assignment_5_11('hhh') ... OK
Call assignment_5_11('hhhhhhhhhhhhhhhhhhhhhhhhhhhhhhhhhhhhhhhhh') ... OK
Call assignment_5_11('aHbhchdheha') ... OK

Well done! You seem to have solved assignment_5_11!
```

<br />
<br />
<br />
<br />

# 5.12 Delete every third character

Given a string, delete from it all the characters whose indices are divisible by 3. Return the resulting string. It is is possible to solve this without an `if` statement, by using the `step` argument to the `range` function.

## Example input

```python
Python
```

## Example return value

```python
yton
```

<!-- grade_id: assignment-5-12-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 5.12

def assignment_5_12(string):
    letters = list(string)
    for i in range(0, len(letters), 3): 
        letters[i] = ''
    string = ''.join(letters)
    return string 
```

```python
# Test for assignment 5.12

test_cell_5_12(assignment_5_12)
```

Output:
```
Please enter a string (Press Enter for default "Python"):  

The resulting string without characters which indices are divisible by 3:
yton
```

<!-- grade_id: assignment-5-12-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 5.12

check_assertion(assignment_5_12)
```

Output:
```
Call assignment_5_12('Python') ... OK
Call assignment_5_12('Hello') ... OK
Call assignment_5_12('qwer') ... OK
Call assignment_5_12('a') ... OK
Call assignment_5_12('ab') ... OK
Call assignment_5_12('abc') ... OK
Call assignment_5_12('abcdefg') ... OK
Call assignment_5_12('abcdefgh') ... OK
Call assignment_5_12('abcdefghi') ... OK
Call assignment_5_12('abcdefghij') ... OK
Call assignment_5_12('abcdefghijk') ... OK
Call assignment_5_12('abcdefghijkl') ... OK
Call assignment_5_12('qwertyuiopasdfghjklzxcvbnm') ... OK

Well done! You seem to have solved assignment_5_12!
```