# 6. While loop

This Jupyter Notebook contains the [problems, or assignments, for Lesson 6](https://snakify.org/en/lessons/while_loop/problems/), gratefully imported from the online course on [Snakify](https://snakify.org). If you get stuck, you can either ask the teacher and the teaching assistents, or skim through the [theory](https://snakify.org/en/lessons/while_loop/) and [explanation of the steps](https://snakify.org/en/lessons/while_loop/steps/1/) on Snakify.

```python
# Please run the contents of this cell.
# Select this cell and press Ctrl-Enter (or Command-Enter on mac).

from assertions import *
from test_cells import *
```

Output:
```
Assertions imported OK
Test cells imported OK
```

<br />
<br />
<br />
<br />

# 6.1 List of squares

For a given integer `n`, return a list with all the squares of integer numbers where the square is less than or equal to `n`, in ascending order.

## Example input

```python
50
```

## Example return value

```python
[1, 4, 9, 16, 25, 36, 49]
```

<!-- grade_id: assignment-06-01-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 6.1

def assignment_6_1(n):
    number = 1
    squares = []
    while number ** 2 <= n: 
        squares.append(number**2)
        number +=1
    
    return squares 
        
```

```python
# Test for assignment 6.1

test_cell_6_1(assignment_6_1)
```

Output:
```
Please enter integer value N (Press Enter for default value 50):  

The resulting list with all squares up to n=50 (inclusive) is:
[1, 4, 9, 16, 25, 36, 49]
```

<!-- grade_id: assignment-06-01-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 6.1

check_assertion(assignment_6_1)
```

Output:
```
Call assignment_6_1(50) ... OK
Call assignment_6_1(10) ... OK
Call assignment_6_1(9) ... OK
Call assignment_6_1(4) ... OK
Call assignment_6_1(1) ... OK
Call assignment_6_1(100) ... OK
Call assignment_6_1(99) ... OK

Well done! You seem to have solved assignment_6_1!
```

<br />
<br />
<br />
<br />

# 6.2 Least divisor

Given an integer `n`, which you may assume will be greater than `1`. Return its smallest integer divisor greater than `1`.

## Example input

```python
15
```

## Example return value

```python
3
```

<!-- grade_id: assignment-06-02-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 6.2

def assignment_6_2(n):
    div = 2
    while n % div != 0: 
        div += 1
        continue
    else: 
        return div 
        
```

```python
# Test for assignment 6.2

test_cell_6_2(assignment_6_2)
```

Output:
```
Please enter an integer not smaller than 2 (n) (Press Enter for default value 15):  

The least divisor for n=15 is:
3
```

<!-- grade_id: assignment-06-02-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 6.2

check_assertion(assignment_6_2)
```

Output:
```
Call assignment_6_2(15) ... OK
Call assignment_6_2(2) ... OK
Call assignment_6_2(3) ... OK
Call assignment_6_2(4) ... OK
Call assignment_6_2(5) ... OK
Call assignment_6_2(6) ... OK
Call assignment_6_2(7) ... OK
Call assignment_6_2(8) ... OK
Call assignment_6_2(35) ... OK
Call assignment_6_2(55) ... OK
Call assignment_6_2(179) ... OK

Well done! You seem to have solved assignment_6_2!
```

<br />
<br />
<br />
<br />

# 6.3 The power of two

For a given integer $n$, find the greatest integer $x$ where ${2}^{x}$ is less than or equal to $n$. Return the exponent value $x$ and the result of the expression ${2}^{x}$ as a tuple.

## Example input

```python
50
```

## Example return value

```python
(5, 32)
```

<!-- grade_id: assignment-06-03-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 6.3

def assignment_6_3(n):
    exp = 1
    while 2**exp <= n: 
        exp += 1
    else:
        return (exp-1, 2**(exp-1))
```

```python
# Test for assignment 6.3

test_cell_6_3(assignment_6_3)
```

Output:
```
Please enter an integer number (n) (Press Enter for default value 50):  

The greatest integer x for n=50 and the result of the expression are:
(5, 32)
```

<!-- grade_id: assignment-06-03-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 6.3

check_assertion(assignment_6_3)
```

Output:
```
Call assignment_6_3(50) ... OK
Call assignment_6_3(10) ... OK
Call assignment_6_3(8) ... OK
Call assignment_6_3(7) ... OK
Call assignment_6_3(1) ... OK
Call assignment_6_3(2) ... OK
Call assignment_6_3(3) ... OK
Call assignment_6_3(4) ... OK
Call assignment_6_3(5) ... OK
Call assignment_6_3(100) ... OK
Call assignment_6_3(1025) ... OK
Call assignment_6_3(15431543) ... OK

Well done! You seem to have solved assignment_6_3!
```

<br />
<br />
<br />
<br />

# 6.4 Morning jog
As a future athlete you just started your practice for an upcoming event. Given that on the first day you run `miles_day_one` miles, and by the event you must be able to run `miles_total` miles.

Calculate the number of days required for you to finally reach the required distance for the event, if you increases your distance each day by 10% from the previous day.

Return one integer representing the number of days to reach the required distance.

## Example input

```python
10
20
```

## Example return value

```python
9
```

<!-- grade_id: assignment-06-04-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 6.4

def assignment_6_4(miles_day_one, miles_total):
    days = 1
    next_day = miles_day_one
    while next_day < miles_total:
        next_day = next_day * 1.1
        days += 1 
    return days
```

```python
# Test for assignment 6.4

test_cell_6_4(assignment_6_4)
```

Output:
```
Please enter the number of miles run on the first day (Press Enter for default value 10):  
Please enter the number of miles needed to run during the event (Press Enter for default value 20):  

Going from 10 miles to 20 miles per day with an increase of 10% per day will take 9 days.
```

<!-- grade_id: assignment-06-04-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 6.4

check_assertion(assignment_6_4)
```

Output:
```
Call assignment_6_4(10, 20) ... OK
Call assignment_6_4(10, 30) ... OK
Call assignment_6_4(10, 100) ... OK
Call assignment_6_4(10, 10) ... OK
Call assignment_6_4(100, 101) ... OK
Call assignment_6_4(100, 99) ... OK
Call assignment_6_4(1, 1000) ... OK
Call assignment_6_4(10, 11) ... OK
Call assignment_6_4(100, 121) ... OK

Well done! You seem to have solved assignment_6_4!
```

<br />
<br />
<br />
<br />

## 6.5 The length of a sequence

Given a sequence of non-negative integers, which are entered by the user as follows:
```python
numbers = []
while True:
    number = int(input('Enter a number (0 to stop): '))
    if number == 0:
        break
    numbers.append(number)
    
### Your code below here
...

```
You can use this in your implementation: it will ask the user for numbers until she enters a `0` to terminate the sequence.

Determine the length of the sequence obtained from the user. Return the length of the sequence (not counting the final `0`).

## Example input

```python
1
7
9
0
```

## Example return value

```python
3
```

<!-- grade_id: assignment-06-05-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 6.5

def assignment_6_5():
    numbers = []
    while True:
        number = int(input('Enter a number (0 to stop): '))
        if number == 0:
            break
        numbers.append(number)
    
    return len(numbers)
```

```python
# Test for assignment 6.5

test_cell_6_5(assignment_6_5)
```

Output:
```
Enter a number (0 to stop):  6
Enter a number (0 to stop):  7
Enter a number (0 to stop):  8
Enter a number (0 to stop):  0

The length of sequence before the 0 is: 3
```

<!-- grade_id: assignment-06-05-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 6.5

check_assertion(assignment_6_5)
```

Output:
```
Call assignment_6_5()
  with inputs ('1', '7', '9', '0', '5') ... OK
Call assignment_6_5()
  with inputs ('1', '2', '3', '4', '5', '6', '7', '0', '1', '2', '3') ... OK
Call assignment_6_5()
  with inputs ('100', '0') ... OK
Call assignment_6_5()
  with inputs ('0',) ... OK
Call assignment_6_5()
  with inputs ('5', '2', '8', '0', '1', '4', '4', '3', '2', '0', '5', '0', '0') ... OK

Well done! You seem to have solved assignment_6_5!
```

<br />
<br />
<br />
<br />

# 6.6 The sum of the sequence

Determine the sum of all elements in the sequence, ending with the number `0`. Like in the previous assignment, make use of a construction with:

```python
number = int(input('Enter a number (0 to stop): '))
```
together with `while`. Return the sum.

## Example input

```python
1
7
9
0
```

## Example return value

```python
17
```

<!-- grade_id: assignment-06-06-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 6.6

def assignment_6_6():
    numbers = []
    while True:
        number = int(input('Enter a number (0 to stop): '))
        if number == 0:
            break
        numbers.append(number)
    
    return sum(numbers)
```

```python
# Test for assignment 6.6

test_cell_6_6(assignment_6_6)
```

Output:
```
Enter a number (0 to stop):  5 
Enter a number (0 to stop):  6
Enter a number (0 to stop):  0

The sum of the numbers in the sequence before the 0 is: 11
```

<!-- grade_id: assignment-06-06-assert, points: 1, grade: true, solution: false -->
```python
# Code for assignment 6.6

check_assertion(assignment_6_6)
```

Output:
```
Call assignment_6_6()
  with inputs ('1', '7', '9', '0') ... OK
Call assignment_6_6()
  with inputs ('1', '1', '1', '1', '1', '1', '1', '1', '1', '0') ... OK
Call assignment_6_6()
  with inputs ('34', '2345', '2345', '2345', '2345', '345', '3', '345', '3', '345', '1', '3', '424', '5', '453', '0') ... OK
Call assignment_6_6()
  with inputs ('7883', '0') ... OK
Call assignment_6_6()
  with inputs ('0',) ... OK
Call assignment_6_6()
  with inputs ('1', '2', '3', '4', '5', '6', '7', '0') ... OK
Call assignment_6_6()
  with inputs ('1', '2', '3', '4', '0') ... OK
Call assignment_6_6()
  with inputs ('0',) ... OK
Call assignment_6_6()
  with inputs ('1', '0') ... OK

Well done! You seem to have solved assignment_6_6!
```

<br />
<br />
<br />
<br />

# 6.7 The average of the sequence

Determine the average of all elements of the sequence ending with the number `0`. Like in the previous assignment, make use of a construction with:
```python
number = int(input('Enter a number (0 to stop): '))
```
together with `while`. Return the average as a floating point number.

## Example input

```python
1
7
9
0
```

## Example return value

```python
5.666666666666667
```

<!-- grade_id: assignment-06-07-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 6.7

def assignment_6_7():
    numbers = []
    while True:
        number = int(input('Enter a number (0 to stop): '))
        if number == 0:
            break
        numbers.append(number)
    
    return sum(numbers) / len(numbers)
```

```python
# Test for assignment 6.7

test_cell_6_7(assignment_6_7)
```

Output:
```
Enter a number (0 to stop):  5
Enter a number (0 to stop):   4
Enter a number (0 to stop):  3
Enter a number (0 to stop):  2
Enter a number (0 to stop):  65
Enter a number (0 to stop):  0

The average of the numbers in the sequence before the 0 is: 15.8
```

<!-- grade_id: assignment-06-07-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 6.7

check_assertion(assignment_6_7)
```

Output:
```
Call assignment_6_7()
  with inputs ('1', '7', '9', '0') ... OK
Call assignment_6_7()
  with inputs ('1', '1', '1', '1', '1', '1', '1', '1', '1', '0') ... OK
Call assignment_6_7()
  with inputs ('34', '2345', '2345', '2345', '2345', '345', '3', '345', '3', '345', '1', '3', '424', '5', '453', '0') ... OK
Call assignment_6_7()
  with inputs ('7883', '0') ... OK
Call assignment_6_7()
  with inputs ('1', '2', '3', '-6', '0') ... OK
Call assignment_6_7()
  with inputs ('1', '2', '3', '4', '5', '6', '7', '0') ... OK
Call assignment_6_7()
  with inputs ('1', '2', '3', '4', '0', '5', '6', '7', '0', '8', '9', '0', '0') ... OK
Call assignment_6_7()
  with inputs ('5', '3', '6', '7', '4', '2', '4', '5', '6', '7', '4', '3', '4', '6', '7', '4', '5', '7', '8', '4', '4', '3', '6', '3', '7', '5', '5', '4', '3', '5', '0') ... OK
Call assignment_6_7()
  with inputs ('1', '0', '2', '0', '3', '0') ... OK

Well done! You seem to have solved assignment_6_7!
```

<br />
<br />
<br />
<br />

# 6.8 The maximum of the sequence

A sequence consists of non-negative integer numbers and ends with the number `0`. Determine the largest element of the sequence. Like in the previous assignment, make use of a construction with:
```python
number = int(input('Enter a number (0 to stop): '))
```
together with `while`. Return the largest element.

## Example input

```python
1
7
9
0
```

## Example return value

```python
9
```

<!-- grade_id: assignment-06-08-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 6.8

def assignment_6_8():
    numbers = []
    while True:
        number = int(input('Enter a number (0 to stop): '))
        if number == 0:
            break
        numbers.append(number)
    return max(numbers)
```

```python
### Test for assignment 6.8

test_cell_6_8(assignment_6_8)
```

Output:
```
Enter a number (0 to stop):  4
Enter a number (0 to stop):  0

The maximum of the numbers in the sequence before the 0 is: 4
```

<!-- grade_id: assignment-06-08-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 6.8

check_assertion(assignment_6_8)
```

Output:
```
Call assignment_6_8()
  with inputs ('1', '7', '9', '0') ... OK
Call assignment_6_8()
  with inputs ('1', '2', '3', '2', '1', '0') ... OK
Call assignment_6_8()
  with inputs ('1', '2', '3', '0') ... OK
Call assignment_6_8()
  with inputs ('1', '3', '2', '0') ... OK
Call assignment_6_8()
  with inputs ('2', '1', '3', '0') ... OK
Call assignment_6_8()
  with inputs ('2', '3', '1', '0') ... OK
Call assignment_6_8()
  with inputs ('3', '2', '1', '0') ... OK
Call assignment_6_8()
  with inputs ('3', '1', '2', '0') ... OK
Call assignment_6_8()
  with inputs ('1', '2', '3', '4', '5', '0') ... OK
Call assignment_6_8()
  with inputs ('5', '4', '3', '2', '1', '0') ... OK
Call assignment_6_8()
  with inputs ('80', '94', '78', '67', '24', '4', '22', '73', '26', '38', '7', '90', '37', '56', '95', '37', '96', '77', '51', '43', '25', '91', '71', '89', '99', '95', '29', '87', '64', '10', '0') ... OK

Well done! You seem to have solved assignment_6_8!
```

<br />
<br />
<br />
<br />

# 6.9 The index of the maximum of a sequence

A sequence consists of non-negative integer numbers and ends with the number `0`. Determine the index of the largest element of the sequence. If the highest element is not unique, return the index of the first of them. Like in the previous assignment, make use of a construction with:
```python
number = int(input('Enter a number (0 to stop): '))
```
together with `while`. Return the index of the first occurence of the largest element. Remember that indices are 0-based!

## Example input

```python
1
7
9
0
```

## Example return value

```python
2
```

<!-- grade_id: assignment-06-09-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 6.9

def assignment_6_9():
    numbers = []
    while True:
        number = int(input('Enter a number (0 to stop): '))
        if number == 0:
            break
        numbers.append(number)
        
    maximum = max(numbers) 
    return numbers.index(maximum) 
```

```python
# Test for assignment 6.9

test_cell_6_9(assignment_6_9)
```

Output:
```
Enter a number (0 to stop):  2345
Enter a number (0 to stop):  324
Enter a number (0 to stop):  3
Enter a number (0 to stop):  4
Enter a number (0 to stop):  0

The index of the first occurence of the maximum in the numbers in the sequence before the 0 is: 0
```

<!-- grade_id: assignment-06-09-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 6.9

check_assertion(assignment_6_9)
```

Output:
```
Call assignment_6_9()
  with inputs ('1', '7', '9', '0') ... OK
Call assignment_6_9()
  with inputs ('1', '2', '3', '2', '1', '0') ... OK
Call assignment_6_9()
  with inputs ('1', '2', '3', '0') ... OK
Call assignment_6_9()
  with inputs ('1', '3', '2', '0') ... OK
Call assignment_6_9()
  with inputs ('2', '1', '3', '0') ... OK
Call assignment_6_9()
  with inputs ('2', '3', '1', '0') ... OK
Call assignment_6_9()
  with inputs ('3', '2', '1', '0') ... OK
Call assignment_6_9()
  with inputs ('3', '1', '2', '0') ... OK
Call assignment_6_9()
  with inputs ('1', '2', '3', '4', '5', '0') ... OK
Call assignment_6_9()
  with inputs ('5', '4', '3', '2', '1', '0') ... OK
Call assignment_6_9()
  with inputs ('3', '5', '1', '2', '5', '0') ... OK
Call assignment_6_9()
  with inputs ('80', '94', '78', '67', '24', '4', '22', '73', '26', '38', '7', '90', '37', '56', '95', '37', '96', '77', '51', '43', '25', '91', '71', '89', '99', '95', '29', '87', '64', '10', '0') ... OK

Well done! You seem to have 
```

<br />
<br />
<br />
<br />

# 6.10 The number of even elements of the sequence

Determine the number of even elements in the sequence ending with the number `0`. Like in the previous assignment, make use of a construction with:
```python
number = int(input('Enter a number (0 to stop): '))
```
together with `while`. Return the number of even elements.

## Example input

```python
2
1
4
0
```

## Example return value

```python
2
```

<!-- grade_id: assignment-6-10-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 6.10

def assignment_6_10():
    numbers = []
    while True:
        number = int(input('Enter a number (0 to stop): '))
        if number == 0:
            break
        numbers.append(number)
    even = 0 
    for i in numbers: 
        if i % 2 == 0:
            even += 1 
    return even 
```

```python
# Test for assignment 6.10

test_cell_6_10(assignment_6_10)
```

Output:
```
Enter a number (0 to stop):  23
Enter a number (0 to stop):  453
Enter a number (0 to stop):  754
Enter a number (0 to stop):  23
Enter a number (0 to stop):  0

The number of even elements in the sequence ending with the number 0 is: 1
```

<!-- grade_id: assignment-6-10-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 6.10

check_assertion(assignment_6_10)
```

Output:
```
Call assignment_6_10()
  with inputs ('2', '1', '4', '0') ... OK
Call assignment_6_10()
  with inputs ('1', '3', '5', '7', '9', '1', '3', '5', '7', '9', '0') ... OK
Call assignment_6_10()
  with inputs ('2', '4', '0') ... OK
Call assignment_6_10()
  with inputs ('7883', '0') ... OK
Call assignment_6_10()
  with inputs ('0',) ... OK
Call assignment_6_10()
  with inputs ('1', '2', '3', '4', '5', '6', '7', '0') ... OK
Call assignment_6_10()
  with inputs ('1', '2', '3', '4', '0', '5', '6', '7', '0', '8', '9', '0', '0') ... OK
Call assignment_6_10()
  with inputs ('5', '3', '6', '7', '4', '2', '4', '5', '6', '7', '4', '3', '4', '6', '7', '4', '5', '7', '8', '4', '4', '3', '6', '3', '7', '5', '5', '4', '3', '5', '0') ... OK
Call assignment_6_10()
  with inputs ('1', '0', '2', '0', '3', '0') ... OK
Call assignment_6_10()
  with inputs ('88', '69', '61', '66', '16', '39', '56', '97', '57', '84', '13', '84', '63', '26', '68', '50', '34', '68', '67', '33', '88', '89', '8', '35', '77', '22', '57
```

<br />
<br />
<br />
<br />

# 6.11 The number of elements that are greater than the previous one

A sequence consists of non-negative integer numbers and ends with the number `0`. Determine how many elements of this sequence are greater than their immediate predecessor. Like in the previous assignment, make use of a construction with:
```python
number = int(input('Enter a number (0 to stop): '))
```
together with `while`. Return the number of elements that are greater than the previous one.

## Example input

```python
1
7
9
0
```

## Example return value

```python
2
```

<!-- grade_id: assignment-6-11-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 6.11

def assignment_6_11():
    numbers = []
    while True:
        number = int(input('Enter a number (0 to stop): '))
        if number == 0:
            break
        numbers.append(number)
    counter = 0
    for i in range(len(numbers)-1): 
        if numbers[i] < numbers[i+1]: 
            counter += 1
    return counter
```

```python
# Test for assignment 6.11

test_cell_6_11(assignment_6_11)
```

<!-- grade_id: assignment-6-11-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 6.11

check_assertion(assignment_6_11)
```

Output:
```
Call assignment_6_11()
  with inputs ('1', '7', '9', '0') ... OK
Call assignment_6_11()
  with inputs ('1', '5', '2', '4', '3', '0') ... OK
Call assignment_6_11()
  with inputs ('1', '2', '3', '4', '5', '0') ... OK
Call assignment_6_11()
  with inputs ('5', '4', '3', '2', '1', '0') ... OK
Call assignment_6_11()
  with inputs ('1', '5', '1', '5', '1', '0') ... OK
Call assignment_6_11()
  with inputs ('5', '1', '5', '1', '5', '0') ... OK
Call assignment_6_11()
  with inputs ('5', '5', '5', '5', '5', '0') ... OK
Call assignment_6_11()
  with inputs ('1', '1', '1', '5', '1', '0') ... OK
Call assignment_6_11()
  with inputs ('345354', '0') ... OK
Call assignment_6_11()
  with inputs ('1', '465', '0') ... OK
Call assignment_6_11()
  with inputs ('4', '1', '0') ... OK
Call assignment_6_11()
  with inputs ('1', '2', '3', '0') ... OK
Call assignment_6_11()
  with inputs ('1', '2', '3', '4', '1', '0') ... OK

Well done! You seem to have solved assignment_6_11!
```

<br />
<br />
<br />
<br />

# 6.12 The second maximum

The sequence consists of distinct non-negative integer numbers and ends with the number `0`. Determine the value of the second largest element in this sequence. Please make sure that the sequences you enter will always have at least two elements. Like in the previous assignment, make use of a construction with:
```python
number = int(input('Enter a number (0 to stop): '))
```
together with `while`. Return the second maximum.

## Example input

```python
1
7
9
0
```

## Example return value

```python
7
```

<!-- grade_id: assignment-6-12-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 6.12

def assignment_6_12():
    numbers = []
    while True:
        number = int(input('Enter a number (0 to stop): '))
        if number == 0:
            break
        numbers.append(number)
    numbers.remove(max(numbers))
    return max(numbers)
```

```python
# Test for assignment 6.12

test_cell_6_12(assignment_6_12)
```

Output:
```
Enter a number (0 to stop):  4
Enter a number (0 to stop):  65
Enter a number (0 to stop):  4
Enter a number (0 to stop):  3
Enter a number (0 to stop):  7
Enter a number (0 to stop):  0

The second largest item in the sequence is : 7
```

<!-- grade_id: assignment-6-12-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 6.12

check_assertion(assignment_6_12)
```

Output:
```
Call assignment_6_12()
  with inputs ('1', '7', '9', '0') ... OK
Call assignment_6_12()
  with inputs ('2', '1', '0') ... OK
Call assignment_6_12()
  with inputs ('1', '2', '0') ... OK
Call assignment_6_12()
  with inputs ('1', '2', '3', '0') ... OK
Call assignment_6_12()
  with inputs ('1', '3', '2', '0') ... OK
Call assignment_6_12()
  with inputs ('2', '3', '1', '0') ... OK
Call assignment_6_12()
  with inputs ('2', '1', '3', '0') ... OK
Call assignment_6_12()
  with inputs ('3', '1', '2', '0') ... OK
Call assignment_6_12()
  with inputs ('3', '2', '1', '0') ... OK
Call assignment_6_12()
  with inputs ('1', '4', '2', '3', '0') ... OK
Call assignment_6_12()
  with inputs ('1', '4', '3', '2', '0') ... OK
Call assignment_6_12()
  with inputs ('4', '1', '2', '3', '0') ... OK
Call assignment_6_12()
  with inputs ('4', '1', '3', '2', '0') ... OK
Call assignment_6_12()
  with inputs ('10', '1', '2', '3', '4', '5', '6', '7', '8', '9', '0') ... OK
Call assignment_6_12()
  with inputs ('1', '
```

<br />
<br />
<br />
<br />

# 6.13 The number of elements equal to the maximum

A sequence consists of integer numbers and ends with the number `0`. Determine how many elements of this sequence are equal to its largest element. Like in the previous assignment, make use of a construction with:
```python
number = int(input('Enter a number (0 to stop): '))
```
together with `while`. Return the number of elements equal to the maximum.

## Example input

```python
1
7
9
0
```

## Example output

```python
1
```

<!-- grade_id: assignment-6-13-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 6.13

def assignment_6_13():
    numbers = []
    while True:
        number = int(input('Enter a number (0 to stop): '))
        if number == 0:
            break
        numbers.append(number)
    counter = 0
    maximum = max(numbers) 
    for i in numbers:
        if i == maximum: 
            counter += 1
    return counter
```

```python
# Test for assignment 6.13

test_cell_6_13(assignment_6_13)
```

<!-- grade_id: assignment-6-13-assert, points: 1, grade: true, solution: false -->
```python
# Code for assignment 6.13

check_assertion(assignment_6_13)
```

Output:
```
Call assignment_6_13()
  with inputs ('1', '7', '9', '0') ... OK
Call assignment_6_13()
  with inputs ('1', '3', '3', '1', '0') ... OK
Call assignment_6_13()
  with inputs ('1', '2', '3', '4', '5', '0') ... OK
Call assignment_6_13()
  with inputs ('5', '4', '3', '2', '1', '0') ... OK
Call assignment_6_13()
  with inputs ('1', '2', '3', '2', '3', '3', '3', '0') ... OK
Call assignment_6_13()
  with inputs ('1', '2', '3', '2', '3', '3', '3', '2', '0') ... OK
Call assignment_6_13()
  with inputs ('1', '1', '1', '1', '1', '1', '0') ... OK
Call assignment_6_13()
  with inputs ('179', '0') ... OK
Call assignment_6_13()
  with inputs ('5', '4', '3', '2', '1', '5', '0') ... OK
Call assignment_6_13()
  with inputs ('10', '39', '77', '33', '86', '79', '19', '62', '27', '79', '0') ... OK
Call assignment_6_13()
  with inputs ('74', '56', '70', '43', '28', '20', '97', '95', '58', '6', '56', '41', '96', '82', '44', '82', '42', '65', '14', '3', '0') ... OK
Call assignment_6_13()
  with inputs ('78', '
```

<br />
<br />
<br />
<br />

# 6.14 Fibonacci numbers

The Fibonacci sequence is defined as follows:

$F_{0} = 0$

$F_{1} = 1$

$F_{n} = F_{n1} + F_{n2}$

Given a non-negative integer $n$, return the $n$th Fibonacci number $F_{n}$.

This problem can also be solved with a `for` loop.

## Example input

```python
6
```
## Example return value

```python
8
```

<!-- grade_id: assignment-6-14-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 6.14

def assignment_6_14(n):
    if n == 0:
        return 0
    
    counter = 2
    prev = 0
    current = 1
    
    while counter <= n:
        next_num = current + prev
        prev = current
        current = next_num
        
        counter +=1
    return current
```

```python
# Test for assignment 6.14

test_cell_6_14(assignment_6_14)
```

Output:
```
Please enter the index for the Fibonacci number (n) (Press Enter for default value 6):  

The 6th Fibonacci number is: 
8
```

<!-- grade_id: assignment-6-14-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 6.14

check_assertion(assignment_6_14)
```

Output:
```
Call assignment_6_14(6) ... OK
Call assignment_6_14(0) ... OK
Call assignment_6_14(2) ... OK
Call assignment_6_14(3) ... OK
Call assignment_6_14(4) ... OK
Call assignment_6_14(1) ... OK
Call assignment_6_14(5) ... OK
Call assignment_6_14(7) ... OK
Call assignment_6_14(8) ... OK
Call assignment_6_14(9) ... OK
Call assignment_6_14(10) ... OK
Call assignment_6_14(11) ... OK
Call assignment_6_14(12) ... OK
Call assignment_6_14(13) ... OK
Call assignment_6_14(14) ... OK
Call assignment_6_14(15) ... OK
Call assignment_6_14(16) ... OK
Call assignment_6_14(17) ... OK
Call assignment_6_14(18) ... OK
Call assignment_6_14(19) ... OK
Call assignment_6_14(20) ... OK

Well done! You seem to have solved assignment_6_14!
```

<br />
<br />
<br />
<br />

# 6.15 The index of a Fibonacci number

The Fibonacci sequence is defined as in the previous assignment.

Given an integer $x$, determine its index among the Fibonacci numbers, that is, return the number $n$ such that $F_{n} = x$. If $x$ is not a Fibonacci number, return `None`.

**Please note** that the assertion cell will check your solution with very large numbers. One way many students tried to
solve this in the past was by creating a list with the Fibonacci numbers and then determining the index of argument
value . This will create a very large list that might not even fit in memory. It is also not necessary.

A suggestion to solve this efficiently is to create a loop in which the current Fibonacci number for the iterator is
calculated, then it is compared to $x$. If it is equal to $x$, your function should return the current value of the
iterator, if it is larger than $x$, your function should return ```None```, if it is still smaller, just continue with
the next value for the iterator.

## Example input

```python
8
```

## Example return value

```python
6
```


<!-- grade_id: assignment-6-15-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 6.15

def assignment_6_15(x):    
    counter = 2
    prev = 0
    current = 1
    
    while True:
        next_num = current + prev
        prev = current
        current = next_num
        if current == x: 
            return counter
        elif current > x:
            return None
        
        counter +=1
    return counter
```

```python
# Test for assignment 6.15

test_cell_6_15(assignment_6_15)
```

Output:
```
Please enter an integer number (x) (Press Enter for default value 8):  

The number 8 is the 6th number in the Fibonacci series.
```

<!-- grade_id: assignment-6-15-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 6.15

check_assertion(assignment_6_15)
```

Output:
```
Call assignment_6_15(8) ... OK
Call assignment_6_15(10) ... OK
Call assignment_6_15(13) ... OK
Call assignment_6_15(55) ... OK
Call assignment_6_15(56) ... OK
Call assignment_6_15(57) ... OK
Call assignment_6_15(54) ... OK
Call assignment_6_15(2) ... OK
Call assignment_6_15(3) ... OK
Call assignment_6_15(4) ... OK
Call assignment_6_15(5) ... OK
Call assignment_6_15(6) ... OK
Call assignment_6_15(1134903170) ... OK
Call assignment_6_15(1134903171) ... OK
Call assignment_6_15(1134903169) ... OK

Well done! You seem to have solved assignment_6_15!
```

<br />
<br />
<br />
<br />

# 6.16 The maximum number of consecutive equal elements

Given a sequence of integer numbers ending with the number `0`. Determine the length of the widest fragment where all the elements are equal to each other. Like in earlier assignments above, make use of a construction with:
```python
number = int(input('Enter a number (0 to stop): '))
```
together with `while`. Return the maximum number of consecutive equal elements.

## Example input

```python
1
7
7
9
1
0
```

## Example return value

```python
2
```


<!-- grade_id: assignment-6-16-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 6.16

def assignment_6_16():
    numbers = []
    while True:
        number = int(input('Enter a number (0 to stop): '))
        if number == 0:
            break
        numbers.append(number)
        
    counter = 1
    consecutive = []
    for i in range(len(numbers)-1): 
        if numbers[i+1] == numbers[i]:
            counter += 1
            continue
        if numbers[i+1] != numbers[i]:
            consecutive.append(counter)
            counter = 1

    consecutive.append(counter)
    return max(consecutive)
        
    
    
```

```python
# Test for assignment 6.16

test_cell_6_16(assignment_6_16)
```

Output:
```
Enter a number (0 to stop):  5
Enter a number (0 to stop):  6
Enter a number (0 to stop):  7
Enter a number (0 to stop):  5
Enter a number (0 to stop):  5
Enter a number (0 to stop):  5
Enter a number (0 to stop):  3
Enter a number (0 to stop):  3
Enter a number (0 to stop):  1
Enter a number (0 to stop):  0
[1, 1, 1, 3, 2, 1]

The number of consecutive equal numbers is: 3
```

<!-- grade_id: assignment-6-16-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 6.16

check_assertion(assignment_6_16)
```

Output:
```
Call assignment_6_16()
  with inputs ('1', '7', '7', '9', '1', '0') ... OK
Call assignment_6_16()
  with inputs ('1', '2', '3', '4', '5', '6', '7', '8', '9', '10', '11', '0') ... OK
Call assignment_6_16()
  with inputs ('4', '4', '4', '4', '4', '4', '4', '4', '4', '4', '4', '4', '4', '4', '4', '0') ... OK
Call assignment_6_16()
  with inputs ('6', '0') ... OK
Call assignment_6_16()
  with inputs ('1', '2', '3', '2', '3', '3', '2', '4', '4', '4', '4', '2', '3', '3', '3', '2', '2', '3', '0') ... OK
Call assignment_6_16()
  with inputs ('1', '2', '2', '1', '2', '1', '2', '1', '2', '0') ... OK
Call assignment_6_16()
  with inputs ('1', '2', '2', '0') ... OK
Call assignment_6_16()
  with inputs ('2', '2', '1', '0') ... OK
Call assignment_6_16()
  with inputs ('1', '2', '3', '3', '2', '2', '4', '4', '4', '4', '4', '5', '5', '5', '6', '3', '3', '3', '3', '3', '3', '5', '5', '5', '5', '3', '3', '3', '4', '5', '5', '5', '0') ... OK
Call assignment_6_16()
  with inputs ('1', '2', '1', '1', '1', 
```