# 3. Conditions: if, then, else

This Jupyter Notebook contains the [problems, or assignments, for Lesson 3](https://snakify.org/en/lessons/if_then_else_conditions/problems/), gratefully imported from the online course on [Snakify](https://snakify.org). If you get stuck, you can either ask the teacher and the teaching assistents, or skim through the [theory](https://snakify.org/en/lessons/if_then_else_conditions/) and [explanation of the steps](https://snakify.org/en/lessons/if_then_else_conditions/steps/1/) on Snakify.

```python
# Please run the contents of this cell.
# Select this cell and press Ctrl-Enter (or Command-Enter on mac).

from assertions import *
from test_cells import *
```

Output:
```
Assertions imported OK
Test cells imported OK
```

<br />
<br />
<br />
<br />

# 3.1 Minimum of two numbers

Given two integers `a` and `b`, return the lesser value.

## Example input

```python
3
7
```

## Example return value

```python
3
```

<!-- grade_id: assignment-03-01-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 3.1

def assignment_3_1(a, b):
    if a < b: 
        return a 
    else: 
        return b 
#Did not take into consideration the case when given a = b
```

```python
# Test for assignment 3.1

test_cell_3_1(assignment_3_1)
```

Output:
```
Please enter integer value a (Press Enter for default value 3):  
Please enter integer value b (Press Enter for default value 7):  

The minimum of 3 and 7 is:
3
```

<!-- grade_id: assignment-03-01-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 3.1

check_assertion(assignment_3_1)
```

Output:
```
Call assignment_3_1(3, 7) ... OK
Call assignment_3_1(2, 2) ... OK
Call assignment_3_1(15, 8) ... OK
Call assignment_3_1(-15, -8) ... OK

Well done! You seem to have solved assignment_3_1!
```

<br />
<br />
<br />
<br />

# 3.2 Sign function

For the given integer `number` return `1` if it's positive, `-1` if it's negative, or `0` if it's equal to zero.
Try to use the cascade `if-elif-else` for it.

## Example input

```python
179
```

## Example return value

```python
1
```

<!-- grade_id: assignment-03-02-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 3.2

def assignment_3_2(number):
    if number > 0:
        return 1 
    elif number < 0:
        return -1 
    else: 
        return 0 
```

```python
# Test for assignment 3.2

test_cell_3_2(assignment_3_2)
```

Output:
```
Please enter an integer value (Press Enter for default value 179):  

For number 179 the returned value (-1 or 1) is: 
1
```

<!-- grade_id: assignment-03-02-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 3.2

check_assertion(assignment_3_2)
```

Output:
```
Call assignment_3_2(179) ... OK
Call assignment_3_2(-42) ... OK
Call assignment_3_2(0) ... OK
Call assignment_3_2(17) ... OK
Call assignment_3_2(-100) ... OK

Well done! You seem to have solved assignment_3_2!
```

<br />
<br />
<br />
<br />

# 3.3 Minimum of three numbers

Given three integers `a`, `b` and `c`, return the smallest value.

## Example input

```python
5
3
7
```

## Example return value

```python
3
```

<!-- grade_id: assignment-03-03-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 3.3

def assignment_3_3(a,b,c):
    list = [a,b,c] 
    return min(list) 
```

```python
# Test for assignment 3.3

test_cell_3_3(assignment_3_3)
```

Output:
```
Please enter integer value a (Press Enter for default value 5):  
Please enter integer value b (Press Enter for default value 3):  
Please enter integer value c (Press Enter for default value 7):  

The smallest number of 5, 3 and 7 is:
3
```

<!-- grade_id: assignment-03-03-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 3.3

check_assertion(assignment_3_3)
```

Output:
```
Call assignment_3_3(5, 3, 7) ... OK
Call assignment_3_3(10, 30, 4) ... OK
Call assignment_3_3(-5, -3, -3) ... OK
Call assignment_3_3(1, 10, 20) ... OK
Call assignment_3_3(74, 32, 11) ... OK
Call assignment_3_3(50, 80, 25) ... OK

Well done! You seem to have solved assignment_3_3!
```

<br />
<br />
<br />
<br />

# 3.4 Equal numbers

Given three integers `a`, `b` and `c`, determine how many of them are equal to each other. The function must return one of these numbers: 3 (if all are the same), 2 (if two of them are equal to each other and the third is different) or 0 (if all numbers are different).

## Example input

```python
10
5
10
```

## Example return value

```python
2
```

<!-- grade_id: assignment-03-04-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 3.4

def assignment_3_4(a, b, c):
    if a != b: 
        if c == b or c == a: 
            return 2
        else:
            return 0 
    elif a == b: 
        if c == b or c == a: 
            return 3 
        else: 
            return 2 
    
    
```

```python
# Test for assignment 3.4

test_cell_3_4(assignment_3_4)
```

Output:
```
Please enter integer value a (Press Enter for default value 10):  
Please enter integer value b (Press Enter for default value 5):  
Please enter integer value c (Press Enter for default value 10):  

Of 10, 5 and 10, the number of equal numbers is: 
2
```

<!-- grade_id: assignment-03-04-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 3.4

check_assertion(assignment_3_4)
```

Output:
```
Call assignment_3_4(10, 5, 10) ... OK
Call assignment_3_4(17, 17, -9) ... OK
Call assignment_3_4(4, -82, -82) ... OK
Call assignment_3_4(5, 2, 4) ... OK
Call assignment_3_4(-149, -146, -142) ... OK
Call assignment_3_4(666, 666, 666) ... OK

Well done! You seem to have solved assignment_3_4!
```

<br />
<br />
<br />
<br />

# 3.5 Car route

A car can cover distance of number of kilometers per day (`kms_per_day`). How many days will it take to cover a route of length `total_distance` kilometers? The function takes two integer arguments: `kms_per_day` and `total_distance`. Please note that your function should return an integer, representing a whole number of days.


## Example input

```python
700
750
```

## Example return value

```python
2
```

<!-- grade_id: assignment-03-05-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 3.5

def assignment_3_5(kms_per_day, total_distance):
    if total_distance % kms_per_day != 0: 
        return (total_distance // kms_per_day) + 1 
    else: 
        return total_distance // kms_per_day 
```

```python
# Test for assignment 3.5

test_cell_3_5(assignment_3_5)
```

Output:
```
Please enter kilometers per day (Press Enter for default value 700):  
Please enter total distance (Press Enter for default value 750):  

The number of days it takes a car to drive 750 at 700 kilometers per day is: 
2
```

<!-- grade_id: assignment-03-05-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 3.5

check_assertion(assignment_3_5)
```

Output:
```
Call assignment_3_5(700, 750) ... OK
Call assignment_3_5(700, 2100) ... OK
Call assignment_3_5(10, 15) ... OK
Call assignment_3_5(10, 16) ... OK
Call assignment_3_5(10, 19) ... OK
Call assignment_3_5(10, 70) ... OK
Call assignment_3_5(10, 81) ... OK
Call assignment_3_5(10, 1000) ... OK
Call assignment_3_5(10, 1001) ... OK
Call assignment_3_5(10, 999) ... OK
Call assignment_3_5(10, 1) ... OK
Call assignment_3_5(10, 9) ... OK
Call assignment_3_5(483, 9382) ... OK
Call assignment_3_5(123, 234234) ... OK

Well done! You seem to have solved assignment_3_5!
```

<br />
<br />
<br />
<br />

# 3.6 Rook move

In the next couple of assignments, you will encounter multiple chess exercises. The notation for squares in chess is `a...h` for the columns and `1...8` for the rows. Since Python is zero-based, we will be using `0...7` for both row and column indices. See the same position with a white rook in both chess notation (`d4`) and Python notation (`3,3`).

<table><tr>
<td style="text-align: center; vertical-align: middle;"><img src="figures/rook_move.svg" alt="Chess" width="150%"/> <h2>Chess</h2></td>
<td style="text-align: center; vertical-align: middle;"><h2><--></h2></td>

<td style="text-align: center; vertical-align: middle;"><img src="figures/rook_python.svg" alt="Python" width="150%"/> <h2>Python</h2></td>
</tr></table>


The rook in chess moves horizontally or vertically. Given two different squares of the chessboard, determine whether a rook can go from the first square to the second in one move.

The function has four arguments, numbers from 0 to 7, each specifying the column and row number, first two - for the first square, and then the last two - for the second square. The function should return `YES` if a rook can go from the first square to the second in one move, or `NO` otherwise.

## Example input

```python
3
3
4
4
```

## Example return value

```python
NO
```

# 3.6 Rook move

In the next couple of assignments, you will encounter multiple chess exercises. The notation for squares in chess is `a...h` for the columns and `1...8` for the rows. Since Python is zero-based, we will be using `0...7` for both row and column indices. See the same position with a white rook in both chess notation (`d4`) and Python notation (`3,3`).

<table><tr>
<td style="text-align: center; vertical-align: middle;"><img src="figures/rook_move.svg" alt="Chess" width="150%"/> <h2>Chess</h2></td>
<td style="text-align: center; vertical-align: middle;"><h2><--></h2></td>

<td style="text-align: center; vertical-align: middle;"><img src="figures/rook_python.svg" alt="Python" width="150%"/> <h2>Python</h2></td>
</tr></table>


The rook in chess moves horizontally or vertically. Given two different squares of the chessboard, determine whether a rook can go from the first square to the second in one move.

The function has four arguments, numbers from 0 to 7, each specifying the column and row number, first two - for the first square, and then the last two - for the second square. The function should return `YES` if a rook can go from the first square to the second in one move, or `NO` otherwise.

## Example input

```python
3
3
4
4
```

## Example return value

```python
NO
```

# 3.6 Rook move

In the next couple of assignments, you will encounter multiple chess exercises. The notation for squares in chess is `a...h` for the columns and `1...8` for the rows. Since Python is zero-based, we will be using `0...7` for both row and column indices. See the same position with a white rook in both chess notation (`d4`) and Python notation (`3,3`).

<table><tr>
<td style="text-align: center; vertical-align: middle;"><img src="figures/rook_move.svg" alt="Chess" width="150%"/> <h2>Chess</h2></td>
<td style="text-align: center; vertical-align: middle;"><h2><--></h2></td>

<td style="text-align: center; vertical-align: middle;"><img src="figures/rook_python.svg" alt="Python" width="150%"/> <h2>Python</h2></td>
</tr></table>


The rook in chess moves horizontally or vertically. Given two different squares of the chessboard, determine whether a rook can go from the first square to the second in one move.

The function has four arguments, numbers from 0 to 7, each specifying the column and row number, first two - for the first square, and then the last two - for the second square. The function should return `YES` if a rook can go from the first square to the second in one move, or `NO` otherwise.

## Example input

```python
3
3
4
4
```

## Example return value

```python
NO
```

# 3.6 Rook move

In the next couple of assignments, you will encounter multiple chess exercises. The notation for squares in chess is `a...h` for the columns and `1...8` for the rows. Since Python is zero-based, we will be using `0...7` for both row and column indices. See the same position with a white rook in both chess notation (`d4`) and Python notation (`3,3`).

<table><tr>
<td style="text-align: center; vertical-align: middle;"><img src="figures/rook_move.svg" alt="Chess" width="150%"/> <h2>Chess</h2></td>
<td style="text-align: center; vertical-align: middle;"><h2><--></h2></td>

<td style="text-align: center; vertical-align: middle;"><img src="figures/rook_python.svg" alt="Python" width="150%"/> <h2>Python</h2></td>
</tr></table>


The rook in chess moves horizontally or vertically. Given two different squares of the chessboard, determine whether a rook can go from the first square to the second in one move.

The function has four arguments, numbers from 0 to 7, each specifying the column and row number, first two - for the first square, and then the last two - for the second square. The function should return `YES` if a rook can go from the first square to the second in one move, or `NO` otherwise.

## Example input

```python
3
3
4
4
```

## Example return value

```python
NO
```

<!-- grade_id: assignment-03-06-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 3.6

def assignment_3_6(col1, row1, col2, row2):
    if col1 == col2 and row1 != row2:
        return "YES" 
    elif col1 != col2 and row1 == row2:
        return "YES" 
    else: 
        return "NO"
        
```

```python
# Test for assignment 3.6

test_cell_3_6(assignment_3_6)
```

Output:
```
Please enter column number of From field (col1) (Press Enter for default value 3):  
Please enter row number of From field (row1) (Press Enter for default value 3):  
Please enter column number of To field (col2) (Press Enter for default value 4):  
Please enter row number of To field (row2) (Press Enter for default value 4):  

The rook can move from (3,3) to (4,4)? 
NO
```

<!-- grade_id: assignment-03-06-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 3.6

check_assertion(assignment_3_6)
```

Output:
```
Call assignment_3_6(3, 3, 4, 4) ... OK
Call assignment_3_6(3, 3, 4, 3) ... OK
Call assignment_3_6(3, 3, 4, 2) ... OK
Call assignment_3_6(3, 3, 3, 4) ... OK
Call assignment_3_6(3, 3, 2, 4) ... OK
Call assignment_3_6(3, 3, 3, 2) ... OK
Call assignment_3_6(3, 3, 2, 3) ... OK
Call assignment_3_6(3, 3, 2, 2) ... OK
Call assignment_3_6(0, 0, 0, 7) ... OK
Call assignment_3_6(0, 0, 7, 7) ... OK
Call assignment_3_6(0, 0, 7, 0) ... OK
Call assignment_3_6(0, 7, 7, 7) ... OK
Call assignment_3_6(0, 7, 7, 0) ... OK
Call assignment_3_6(0, 7, 0, 0) ... OK
Call assignment_3_6(7, 7, 7, 0) ... OK
Call assignment_3_6(7, 7, 0, 0) ... OK
Call assignment_3_6(7, 7, 0, 7) ... OK
Call assignment_3_6(7, 0, 0, 0) ... OK
Call assignment_3_6(7, 0, 0, 7) ... OK
Call assignment_3_6(7, 0, 7, 7) ... OK

Well done! You seem to have solved assignment_3_6!
```

<br />
<br />
<br />
<br />

# 3.7 Chess board - same color

Given two squares of a chessboard. Finalize the assignment function in such a way that, if the squares entered have the same color, the word `YES` is being returned, and `NO` otherwise.

The function has four arguments, numbers from 0 to 7, each specifying the column and row numbers: the first two designate the column and row of the first square, the latter two the column and row of the second square.

<table><tr>
<td style="text-align: center; vertical-align: middle;"><img src="figures/chess_board.svg" alt="Squares a1 (0,0) and b6 (1,5)" width="100%"/> <h2>Squares a1 (0,0) and b6 (1,5)</h2></td>
</tr></table>

## Example input

```python
0
0
1
5
```

## Example return value

```python
YES
```

<!-- grade_id: assignment-03-07-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 3.7

def assignment_3_7(col1, row1, col2, row2):
    col1_even = col1 % 2 == 0 
    row1_even = row1 % 2 == 0 
    col2_even = col2 % 2 == 0 
    row2_even = row2 % 2 == 0 
    if row1_even == col1_even: 
        if row2_even == col2_even: 
            return "YES" 
        else:
            return "NO"
    elif row1_even != col1_even: 
        if row2_even != col2_even:
            return "YES" 
        else: 
            return "NO" 
```

```python
# Test for assignment 3.7

test_cell_3_7(assignment_3_7)
```

Output:
```
Please enter column number of From field (col1) (Press Enter for default value 0):  
Please enter row number of From field (row1) (Press Enter for default value 0):  
Please enter column number of To field (col2) (Press Enter for default value 1):  
Please enter row number of To field (row2) (Press Enter for default value 5):  

Do the squares (0,0) and (1,5) have the same color? 
YES
```

<!-- grade_id: assignment-03-07-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 3.7

check_assertion(assignment_3_7)
```

Output:
```
Call assignment_3_7(0, 0, 1, 5) ... OK
Call assignment_3_7(1, 1, 1, 4) ... OK
Call assignment_3_7(1, 1, 1, 3) ... OK
Call assignment_3_7(1, 2, 2, 1) ... OK
Call assignment_3_7(1, 2, 6, 7) ... OK
Call assignment_3_7(1, 2, 7, 7) ... OK
Call assignment_3_7(4, 6, 4, 6) ... OK
Call assignment_3_7(1, 5, 2, 0) ... OK
Call assignment_3_7(1, 2, 3, 4) ... OK
Call assignment_3_7(6, 1, 1, 2) ... OK
Call assignment_3_7(0, 5, 6, 1) ... OK

Well done! You seem to have solved assignment_3_7!
```

<br />
<br />
<br />
<br />

# 3.8 King move


The chess king moves horizontally, vertically or diagonally to any adjacent square. Given two different squares of the chessboard, determine whether a king can go from the first square to the second in one move.

The function has four arguments, numbers from 0 to 7, each specifying the column and row number, first two - for the first square, and then the last two - for the second sqaure. The function should return `YES` if the king can go from the first square to the second in one move, or `NO` otherwise.

<table><tr>
<td style="text-align: center; vertical-align: middle;"><img src="figures/king_move.svg" alt="Possible moves for king on e4 (4,3)" width="100%"/> <h2>Possible moves for king on e4 (4,3)</h2></td>
</tr></table>

## Example input

```python
4
3
5
4
```

## Example return value

```python
YES
```

<!-- grade_id: assignment-03-08-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 3.8

def assignment_3_8(col1, row1, col2, row2):
    col_diff = abs(col1 - col2) 
    row_diff = abs(row1 - row2) 
    if col_diff != 1 and col_diff != 0:
        return "NO"
    if row_diff == 1 or row_diff == 0: 
        return "YES" 
    else: 
        return "NO" 
        
```

```python
# Test for assignment 3.8

test_cell_3_8(assignment_3_8)
```

Output:
```
Please enter column number of From field (col1) (Press Enter for default value 4):  
Please enter row number of From field (row1) (Press Enter for default value 3):  
Please enter column number of To field (col2) (Press Enter for default value 5):  
Please enter row number of To field (row2) (Press Enter for default value 4):  

The king can move from (4,3) to (5,4)? 
YES
```

<!-- grade_id: assignment-03-08-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 3.8

check_assertion(assignment_3_8)
```

Output:
```
Call assignment_3_8(3, 3, 4, 4) ... OK
Call assignment_3_8(3, 3, 4, 3) ... OK
Call assignment_3_8(3, 3, 4, 2) ... OK
Call assignment_3_8(3, 3, 3, 4) ... OK
Call assignment_3_8(3, 3, 2, 4) ... OK
Call assignment_3_8(3, 3, 3, 2) ... OK
Call assignment_3_8(3, 3, 2, 3) ... OK
Call assignment_3_8(3, 3, 2, 2) ... OK
Call assignment_3_8(0, 0, 0, 7) ... OK
Call assignment_3_8(0, 0, 7, 7) ... OK
Call assignment_3_8(0, 0, 7, 0) ... OK
Call assignment_3_8(0, 7, 7, 7) ... OK
Call assignment_3_8(0, 7, 7, 0) ... OK
Call assignment_3_8(0, 7, 0, 0) ... OK
Call assignment_3_8(7, 7, 7, 0) ... OK
Call assignment_3_8(7, 7, 0, 0) ... OK
Call assignment_3_8(7, 7, 0, 7) ... OK
Call assignment_3_8(7, 0, 0, 0) ... OK
Call assignment_3_8(7, 0, 0, 7) ... OK
Call assignment_3_8(7, 0, 7, 7) ... OK
Call assignment_3_8(0, 0, 0, 1) ... OK
Call assignment_3_8(0, 0, 1, 1) ... OK
Call assignment_3_8(0, 0, 1, 0) ... OK
Call assignment_3_8(3, 3, 5, 5) ... OK
Call assignment_3_8(3, 3, 1, 1) ... OK
Call assignment_3_8(3, 3,
```

<br />
<br />
<br />
<br />

# 3.9 Bishop move

In chess, the bishop moves diagonally, any number of squares. Given two different squares of the chessboard, determine whether a bishop can go from the first to the second in one move.

The function has four arguments, numbers from 0 to 7, specifying the column and row numbers of the starting square and the column and row numbers of the ending square. The function should return `YES` if a Bishop can go from the first square to the second in one move, or `NO` otherwise.

<table><tr>
<td style="text-align: center; vertical-align: middle;"><img src="figures/bishop_move.svg" alt="Possible moves for bishop on e3 (4,2)" width="100%"/> <h2>Possible moves for bishop on e3 (4,2)</h2></td>
</tr></table>

## Example input

```python
4
2
1
5
```

## Example return value

```python
YES
```

<!-- grade_id: assignment-03-09-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 3.9

def assignment_3_9(col1, row1, col2, row2):
    if abs(col1 - col2) == abs(row1 - row2): 
        return "YES" 
    else: 
        return "NO"
```

```python
# Test for assignment 3.9

test_cell_3_9(assignment_3_9)
```

Output:
```
Please enter column number of From field (col1) (Press Enter for default value 4):  
Please enter row number of From field (row1) (Press Enter for default value 2):  
Please enter column number of To field (col2) (Press Enter for default value 1):  
Please enter row number of To field (row2) (Press Enter for default value 5):  

The bishop can move from (4,2) to (1,5)? 
YES
```

<!-- grade_id: assignment-03-09-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 3.9

check_assertion(assignment_3_9)
```

Output:
```
Call assignment_3_9(3, 3, 4, 4) ... OK
Call assignment_3_9(3, 3, 4, 3) ... OK
Call assignment_3_9(3, 3, 4, 2) ... OK
Call assignment_3_9(3, 3, 3, 4) ... OK
Call assignment_3_9(3, 3, 2, 4) ... OK
Call assignment_3_9(3, 3, 3, 2) ... OK
Call assignment_3_9(3, 3, 2, 3) ... OK
Call assignment_3_9(3, 3, 2, 2) ... OK
Call assignment_3_9(0, 0, 0, 7) ... OK
Call assignment_3_9(0, 0, 7, 7) ... OK
Call assignment_3_9(0, 0, 7, 0) ... OK
Call assignment_3_9(0, 7, 7, 7) ... OK
Call assignment_3_9(0, 7, 7, 0) ... OK
Call assignment_3_9(0, 7, 0, 0) ... OK
Call assignment_3_9(7, 7, 7, 0) ... OK
Call assignment_3_9(7, 7, 0, 0) ... OK
Call assignment_3_9(7, 7, 0, 7) ... OK
Call assignment_3_9(7, 0, 0, 0) ... OK
Call assignment_3_9(7, 0, 0, 7) ... OK
Call assignment_3_9(7, 0, 7, 7) ... OK
Call assignment_3_9(0, 0, 0, 1) ... OK
Call assignment_3_9(0, 0, 1, 1) ... OK
Call assignment_3_9(0, 0, 1, 0) ... OK
Call assignment_3_9(3, 3, 5, 5) ... OK
Call assignment_3_9(3, 3, 1, 1) ... OK
Call assignment_3_9(3, 3,
```

<br />
<br />
<br />
<br />

# 3.10 Queen move

Chess queens can move horizontally, vertically or diagonally. Given two different squares of the chessboard, determine whether a queen can go from the first square to the second in one move.

The function has four arguments, numbers from 0 to 7, each specifying the column and row number, first two - for the first square, and then the last two - for the second square. The function should return `YES` if a queen can go from the first square to the second in one move, or `NO` otherwise.

<table><tr>
<td style="text-align: center; vertical-align: middle;"><img src="figures/queen_move.svg" alt="Possible moves for queen on e5 (4,4)" width="100%"/> <h2>Possible moves for queen on e5 (4,4)</h2></td>
</tr></table>

## Example input

```python
4
4
6
6
```

## Example return value

```python
YES
```

<!-- grade_id: assignment-3-10-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 3.10

def assignment_3_10(col1, row1, col2, row2): 
        same_col = col1 == col2 and row1 != row2
        same_row = row1 == row2 and col1 != col2 
        bishop_mov = abs(col1 - col2) == abs(row1 - row2) 
        if same_col or same_row: 
            return "YES" 
        elif bishop_mov: 
            return "YES" 
        else:
            return "NO" 
```

```python
# Test for assignment 3.10

test_cell_3_10(assignment_3_10)
```

Output:
```
Please enter column number of From field (col1) (Press Enter for default value 4):  
Please enter row number of From field (row1) (Press Enter for default value 4):  
Please enter column number of To field (col2) (Press Enter for default value 6):  
Please enter row number of To field (row2) (Press Enter for default value 6):  

The queen can move from (4,4) to (6,6)? 
YES
```

<!-- grade_id: assignment-3-10-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 3.10

check_assertion(assignment_3_10)
```

Output:
```
Call assignment_3_10(0, 0, 1, 1) ... OK
Call assignment_3_10(0, 0, 1, 2) ... OK
Call assignment_3_10(4, 5, 2, 2) ... OK
Call assignment_3_10(2, 2, 0, 0) ... OK
Call assignment_3_10(5, 4, 1, 4) ... OK
Call assignment_3_10(6, 5, 4, 1) ... OK
Call assignment_3_10(1, 6, 5, 6) ... OK
Call assignment_3_10(1, 6, 3, 5) ... OK
Call assignment_3_10(6, 3, 1, 4) ... OK
Call assignment_3_10(6, 4, 0, 0) ... OK
Call assignment_3_10(1, 3, 4, 6) ... OK
Call assignment_3_10(2, 4, 6, 0) ... OK
Call assignment_3_10(4, 1, 4, 7) ... OK
Call assignment_3_10(0, 1, 2, 0) ... OK
Call assignment_3_10(1, 0, 0, 2) ... OK

Well done! You seem to have solved assignment_3_10!
```

<br />
<br />
<br />
<br />

# 3.11 Knight move

The moves of the knights in chess resemble the letter **L**. A knight can move two squares horizontally and one square vertically, or two squares vertically and one square horizontally. Given two different squares of the chessboard, determine whether a knight can go from the first square to the second in one move.

The function has four arguments, numbers from 0 to 7, each specifying the column and row number, first two - for the first square, and then the last two - for the second square. The function should return `YES` if a knight can go from the first square to the second in one move, or `NO` otherwise.

<table><tr>
<td style="text-align: center; vertical-align: middle;"><img src="figures/knight_move.svg" alt="Possible moves for knight on e4 (4,3)" width="100%"/> <h2>Possible moves for knight on e4 (4,3)</h2></td>
</tr></table>

## Example input

```python
4
3
4
7
```

## Example return value

```python
NO
```


<!-- grade_id: assignment-3-11-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 3.11

def assignment_3_11(col1, row1, col2, row2):
    col_diff = abs(col1 - col2) 
    row_diff = abs(row1 - row2) 
    if col_diff == 1 and row_diff == 2: 
        return "YES" 
    elif col_diff == 2 and row_diff ==1: 
        return "YES" 
    else: 
        return "NO" 

```

```python
# Test for assignment 3.11

test_cell_3_11(assignment_3_11)
```

Output:
```
Please enter column number of From field (col1) (Press Enter for default value 4):  
Please enter row number of From field (row1) (Press Enter for default value 3):  
Please enter column number of To field (col2) (Press Enter for default value 4):  
Please enter row number of To field (row2) (Press Enter for default value 7):  

The knight can move from (4,3) to (4,7)? 
NO
```

<!-- grade_id: assignment-3-11-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 3.11

check_assertion(assignment_3_11)
```

Output:
```
Call assignment_3_11(0, 0, 0, 3) ... OK
Call assignment_3_11(0, 0, 7, 7) ... OK
Call assignment_3_11(1, 3, 2, 1) ... OK
Call assignment_3_11(4, 1, 3, 3) ... OK
Call assignment_3_11(1, 7, 2, 6) ... OK
Call assignment_3_11(1, 7, 2, 4) ... OK
Call assignment_3_11(4, 4, 2, 6) ... OK
Call assignment_3_11(1, 3, 1, 4) ... OK
Call assignment_3_11(3, 6, 5, 5) ... OK
Call assignment_3_11(3, 4, 1, 3) ... OK
Call assignment_3_11(1, 2, 2, 1) ... OK
Call assignment_3_11(4, 0, 3, 2) ... OK
Call assignment_3_11(5, 1, 7, 2) ... OK

Well done! You seem to have solved assignment_3_11!
```

<br />
<br />
<br />
<br />

# 3.12 Chocolate bar

A chocolate bar has the form of a rectangle divided into `cols x rows` portions. The chocolate bar can be split into two rectangular parts by breaking it along a selected straight line on its pattern. 

Determine whether it is possible to split it so that one of the parts will have exactly `squares` squares.

The function uses three integer arguments: `cols`, `rows`, and `squares`. It should return `YES` or `NO`.

An example chocolate bar is shown in this figure, where `cols` = 4 and `rows` = 6:

<img src="figures/chocolate_bar.jpg" alt="Chocolate bar" width="50%"/>

The chocolate bar is going to be broken into two pieces *along one of the lines, either horizontally, or vertically*, resulting into *two rectangular* pieces.

Imagine that we choose `squares` = 8. If we break the chocolate bar on the second horizontal line from the top, we will end up in one piece of 4 x 2 squares and one piece of 4 x 4 squares. Indeed, the 4 x 2 matches `squares` = 8.

Now, if we look in a more general way, we will see that if we break the chocolate bar on a horizontal line, we will end up with rectangles of 4 x 1, 4 x 2, 4 x 3, <br>
4 x 4 or 4 x 5 squares. Meaning that if `squares` is divisible by 4, in this case, we will have to return `YES`. We could use the remainder operator `%` to check for this.

The same is true for a vertical split. We end up with 1 x 6, 2 x 6 and 3 x 6 squares. Meaning that if `squares` is divisible by 6, in this case, we will have to return `YES`.

## Example input

```python
4
6
8
```

## Example return value

```python
YES
```

<!-- grade_id: assignment-3-12-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 3.12

def assignment_3_12(cols, rows, squares):
    if squares % cols == 0:
        if squares <= cols*(rows-1): 
            return "YES" 
        else: 
            return "NO"
    elif squares % rows == 0: 
        if squares <= rows*(cols-1): 
            return "YES" 
        else:
            return "NO" 
    else: 
        return "NO"
        
```

```python
# Test for assignment 3.12

test_cell_3_12(assignment_3_12)
```

Output:
```
Please enter number of columns (cols) for the chocolate bar (Press Enter for default value 4):  
Please enter number of rows (rows) for the chocolate bar (Press Enter for default value 6):  
Please enter number of squares (squares) for the part (Press Enter for default value 8):  

Can one part of a 4x6 squares chocolate bar have 8 squares?
YES
```

<!-- grade_id: assignment-3-12-assert, points: 1, grade: true, solution: false -->
```python
# Assertion  for assignment 3.12

check_assertion(assignment_3_12)
```

Output:
```
Call assignment_3_12(4, 2, 6) ... OK
Call assignment_3_12(2, 10, 7) ... OK
Call assignment_3_12(5, 7, 1) ... OK
Call assignment_3_12(7, 4, 21) ... OK
Call assignment_3_12(5, 12, 100) ... OK
Call assignment_3_12(6, 6, 6) ... OK
Call assignment_3_12(6, 6, 35) ... OK
Call assignment_3_12(6, 6, 37) ... OK
Call assignment_3_12(7, 1, 99) ... OK
Call assignment_3_12(300, 100, 3000) ... OK
Call assignment_3_12(256, 124, 4069) ... OK
Call assignment_3_12(348, 41, 6183) ... OK
Call assignment_3_12(387, 13, 2709) ... OK
Call assignment_3_12(13, 387, 2709) ... OK
Call assignment_3_12(1, 1, 2) ... OK

Well done! You seem to have solved assignment_3_12!
```

<br />
<br />
<br />
<br />

# 3.13 Leap year

Given the year number. You need to check if this year is a leap year. If it is, return `LEAP`, otherwise return `COMMON`.

The rules in Gregorian calendar are as follows:
 * a year is a leap year if its number is exactly divisible by 4 but it is not exactly divisible by 100
 * a year is always a leap year if its number is exactly divisible by 400
 
**Warning.** The words `LEAP` and `COMMON` should be all caps.

## Example input

```python
2012
```

## Example return value

```python
LEAP
```

<!-- grade_id: assignment-3-13-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 3.13

def assignment_3_13(year):
    if year % 4 == 0 and year % 100 != 0:
        return "LEAP" 
    elif year % 400 == 0: 
        return "LEAP"  
    else: 
        return "COMMON" 
```

```python
# Test for assignment 3.13

test_cell_3_13(assignment_3_13)
```

Output:
```
Please enter the year to be checked (Press Enter for default value 2012):  

Is the year 2012 a LEAP or a COMMON year? 
LEAP
```

<!-- grade_id: assignment-3-13-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 3.13

check_assertion(assignment_3_13)
```

Output:
```
Call assignment_3_13(2012) ... OK
Call assignment_3_13(2011) ... OK
Call assignment_3_13(1492) ... OK
Call assignment_3_13(1861) ... OK
Call assignment_3_13(1600) ... OK
Call assignment_3_13(1700) ... OK
Call assignment_3_13(1800) ... OK
Call assignment_3_13(1900) ... OK
Call assignment_3_13(2000) ... OK

Well done! You seem to have solved assignment_3_13!
```