# 8. Functions and recursion

This Jupyter Notebook contains the [problems, or assignments, for Lesson 8](https://snakify.org/en/lessons/functions/problems/), gratefully imported from the online course on [Snakify](https://snakify.org). If you get stuck, you can either ask the teacher and the teaching assistents, or skim through the [theory](https://snakify.org/en/lessons/functions/) and [explanation of the steps](https://snakify.org/en/lessons/functions/steps/1/) on Snakify.

**NB** Since the focus of these assignments is on `functions`, you are requested to define the whole function yourself.

```python
# Please run the contents of this cell.
# Select this cell and press Ctrl-Enter (or Command-Enter on mac).

from assertions import *
from test_cells import *
```

Output:
```
Assertions imported OK
Test cells imported OK
```

<br />
<br />
<br />
<br />

# 8.1 The length of the segment

Given four real numbers representing cartesian coordinates: $(x_1, y_1),(x_2, y_2)$.  
Write a function `distance(x1, y1, x2, y2)` to compute the distance between the points $(x_1, y_1)$ and $(x_2, y_2)$. 

You can make use of the function that computes the square root from the `math` module.
```python
math.sqrt()
```

Make sure your function returns the distance.

The formula for distance between two points can be found at [Wolfram](http://mathworld.wolfram.com/Distance.html).

In the test cell, enter the four numbers separated by spaces.

## Example input
```python
0 0 1 1
```

## Example return code
```python
1.4142135623730951
```


<!-- grade_id: assignment-08-01-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 8.1
import math

# YOUR CODE HERE
def distance(x1, y1, x2, y2): 
    return math.sqrt((y1-y2)**2 + (x1-x2)**2)
```

```python
# Test for assignment 8.1

test_cell_8_1(distance)
```

Output:
```
Enter 4 numbers, separated by spaces:  
Please enter a series of 4 numbers, separated by spaces (Press Enter for default values 0 0 1 1):  

The distance between the points (0.0, 0.0) and (1.0, 1.0):
1.4142135623730951
```

<!-- grade_id: assignment-08-01-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 8.1

check_assertion(distance)
```

Output:
```
Call distance(0.0, 0.0, 1.0, 1.0) ... OK
Call distance(0.0, 0.0, 1.0, 0.0) ... OK
Call distance(3.0, -2.0, -1.0, 7.0) ... OK
Call distance(0.1, 0.1, 0.2, 0.2) ... OK
Call distance(-1.0, -1.0, -3.0, -5.0) ... OK

Well done! You seem to have solved distance!
```

<br />
<br />
<br />
<br />

# 8.2 Exponentiation

Given a real number $a$ and an integer $n$, compute $a^n$.

Write a function `power(a, n)` that calculates $a^n$ and return the result of the expression.

Don't use the `math.pow()` function from the standard library, nor the syntax `a**n`. Instead, use the definition of exponentiation to an integer power, which is to multipy the value $a$ by itself repeatedly.

Please keep in mind that $n$ can also be negative, for this case you should use the fact that $a^{-n} = \frac{1}{a^n}$. Therefore, you can compute the value just as if $n$ were positive, but you take the reciprocal before returning the result. Also keep in mind that $a^0 = 1$ for all values of $a$.

## Example input
```python
2.0
-3
```

## Example return value
```python
0.125
```

<!-- grade_id: assignment-08-02-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 8.2

# YOUR CODE HERE
def power(a, n): 
    if n == 0:
        return 1.0
    result = 1
    for i in range(abs(n)): 
        result *= a
    
    if n<0:
        return 1/result 
    else:
        return result
        
        
```

```python
# Test for assignment 8.2

test_cell_8_2_4(power)
```

Output:
```
Please enter value a (Press Enter for default value 2.0):  
Please enter integer value n (Press Enter for default value -3):  

The result of 2.0 to the power of -3 is:
0.125
```

<!-- grade_id: assignment-08-02-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 8.2

check_assertion(power)
```

Output:
```
Call power(2.0, -3) ... OK
Call power(2.0, 1) ... OK
Call power(2.0, 2) ... OK
Call power(2.0, 3) ... OK
Call power(2.0, 4) ... OK
Call power(2.0, 9) ... OK
Call power(2.0, 10) ... OK
Call power(2.0, 15) ... OK
Call power(2.0, 0) ... OK
Call power(3.0, 1) ... OK
Call power(3.0, 2) ... OK
Call power(3.0, 3) ... OK
Call power(3.0, 10) ... OK
Call power(3.0, 0) ... OK
Call power(1.1414, 2) ... OK
Call power(1.5, 10) ... OK
Call power(1.0, -1) ... OK
Call power(2.0, -1) ... OK
Call power(2.0, -2) ... OK
Call power(2.0, -3) ... OK
Call power(2.0, -4) ... OK
Call power(2.0, -8) ... OK
Call power(2.0, -9) ... OK
Call power(2.0, -10) ... OK
Call power(2.0, -15) ... OK
Call power(3.0, -1) ... OK
Call power(3.0, -2) ... OK
Call power(3.0, -3) ... OK
Call power(3.0, -4) ... OK
Call power(3.0, -5) ... OK
Call power(3.0, -10) ... OK
Call power(3.0, -6) ... OK
Call power(1.1414, -2) ... OK
Call power(1.5, -10) ... OK

Well done! You seem to have solved power!
```

<br />
<br />
<br />
<br />

# 8.3 Uppercase

Write a function `capitalize(lower_case_words)` that takes a string with lower case words<br/>and returns a line with the first letter of each word capitalized, *e.g.*
```python
print(capitalize('hello world'))
```
should print `Hello World`. Assume that the words are separated by single spaces.

**Hint** In Python there is a function `ord(character)`, which returns the character code for a character in the [the ASCII chart](https://en.wikipedia.org/wiki/ASCII).<br/>
The function `chr(code)` returns the character itself from the ASCII code. For example,
```python
ord('a') == 97
chr(97) == 'a'
```

**NB** Please note the difference in character value between small and capital letters. Use this in your function.

It should be noted that this exercise is meant for experimenting with small functions, *i.e.* **ord()** and **chr()**. This does, in this case, *not* represent the most efficient or elegant way of solving this problem, which would be to just use the builtin `title()` method of strings.

<!-- grade_id: assignment-08-03-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 8.3

def capitalize(string): 
    words = string.split()
    new_string = []
    for word in words: 
        new_word = chr(ord(word[0]) - 32) + word[1:]
        new_string.append(new_word) 
        
    new_string = " ".join(new_string)
    return(new_string) 
```

```python
# Test for assignment 8.3

test_cell_8_3(capitalize)
```

Output:
```
Please enter a string containing only lowercase letters (Press Enter for default "hello world"):  

The capitalized copy of "hello world" is: 
Hello World
```

<!-- grade_id: assignment-08-03-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 8.3

check_assertion(capitalize)
```

Output:
```
Call capitalize('harry potter') ... OK
Call capitalize('procrastination') ... OK
Call capitalize('we danced the mamushka at waterloo') ... OK
Call capitalize('de noche todos los gatos son pardos') ... OK
Call capitalize('all i see turns to brown') ... OK

Well done! You seem to have solved capitalize!
```

<br />
<br />
<br />
<br />

# 8.4 Exponentiation - recursive

For this assignment, we are going to compute $a^n$ once again, like in assignment 8.2.  
But, in this case, **without** making use of loops.

To achieve our goal, we will write a recursive function, *i.e.* a function that calls itself.

For example, if we consider multiplication instead of exponentiation, a recursive solution might look like this:
```python
def multiply_recursive(a, n):
    if n == 0:
        return 0.0
    return a + multiply_recursive(a, n - 1)
```
This function compute the multiplication of $a$ by a non-negative integer $n$ using the definition of repeated addition, i.e. it adds $n$ copies of $a$ together.

As you can see, on the last line, the function calls itself with $n-1$, unless $n = 0$ in which case it does not call itself and the chain of function calls ends.

For this assignment, create a function `power_recursive(a, n)`  
that takes a real number $a$ and a *positive* integer $n$ and computes $a^n$ recursively.

## Example input
```python
2.0
3
```

## Example output
```python
8.0
```

```python
# Trial cell for experimentation with recursive function

def multiply_recursive(a, n):
    if n == 0:
        return 0.0
    return a + multiply_recursive(a, n - 1)
    
print(multiply_recursive(3, 4))
```

Output:
```
12.0
```

<!-- grade_id: assignment-08-04-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 8.4

# YOUR CODE HERE
def power_recursive(a, n):
    if n == 0:
        return 1.0
    return a * power_recursive(a, n - 1)
```

```python
# Test for assignment 8.4

test_cell_8_2_4(power_recursive)
```

Output:
```
Please enter value a (Press Enter for default value 2.0):  
Please enter integer value n (Press Enter for default value 3):  

The result of 2.0 to the power of 3 is:
8.0
```

<!-- grade_id: assignment-08-04-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 8.4

check_assertion(power_recursive)
```

Output:
```
Call power_recursive(2.0, 3) ... OK
Call power_recursive(2.0, 2) ... OK
Call power_recursive(2.0, 1) ... OK
Call power_recursive(2.0, 4) ... OK
Call power_recursive(2.0, 5) ... OK
Call power_recursive(2.0, 6) ... OK
Call power_recursive(2.0, 7) ... OK
Call power_recursive(2.0, 8) ... OK
Call power_recursive(2.0, 9) ... OK
Call power_recursive(2.0, 10) ... OK
Call power_recursive(2.0, 15) ... OK
Call power_recursive(2.0, 0) ... OK
Call power_recursive(3.0, 1) ... OK
Call power_recursive(3.0, 2) ... OK
Call power_recursive(3.0, 3) ... OK
Call power_recursive(3.0, 4) ... OK
Call power_recursive(3.0, 5) ... OK
Call power_recursive(3.0, 10) ... OK
Call power_recursive(3.0, 0) ... OK
Call power_recursive(1.1414, 2) ... OK
Call power_recursive(1.5, 10) ... OK

Well done! You seem to have solved power_recursive!
```

<br />
<br />
<br />
<br />

# 8.5 Factorial - recursively

Create a function `factorial(n)` that takes $n$ and computes $n!$. 
`factorial(n)` should be a recursive function.

Remember that the factorial of a positive integer $n$ is defined to be the product of all integers from $1$ up to and including $n$.

## Example input
```python
10
```

## Example return value
```python
3628800
```

<!-- grade_id: assignment-08-05-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 8.5

def factorial(n):
    if n == 0:
        return 1
    return n * factorial(n - 1)
```

```python
# Test cell for assignment 8.5

test_cell_8_5(factorial)
```

Output:
```
Please enter integer value n (Press Enter for default value 10):  

The value 10! is:
3628800
```

<!-- grade_id: assignment-08-05-assert, points: 1, grade: true, solution: false -->
```python
check_assertion(factorial)
```

Output:
```
Call factorial(10) ... OK
Call factorial(9) ... OK
Call factorial(1) ... OK
Call factorial(2) ... OK
Call factorial(3) ... OK
Call factorial(4) ... OK
Call factorial(5) ... OK

Well done! You seem to have solved factorial!
```

<br />
<br />
<br />
<br />

# 8.6 Fibonacci - recursively

Given a non-negative integer $n$, return the $n$th Fibonacci number.  
Do this by writing a function `fib(n)` which takes the non-negative integer $n$ and returns the $n$th Fibonacci number.

You can look at assignment 6.14 for the definition of the Fibonacci sequence.

Don't use loops, try to use recursion instead.

## Example input
```python
6
```

## Example return value
```python
8
```

<!-- grade_id: assignment-08-06-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 8.6

# YOUR CODE HERE
def fib(n): 
    if n == 1 or n == 2:
        return 1
    
    return fib(n-1) + fib(n-2) 
```

```python
# Test cell for assignment 8.6

test_cell_8_6(fib)
```

Output:
```
Please enter integer value n (max. 35) (Press Enter for default value 6):  

The 6th Fibonacci number is:
8
```

<!-- grade_id: assignment-08-06-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 8.6

check_assertion(fib)
```

Output:
```
Call fib(6) ... OK
Call fib(1) ... OK
Call fib(2) ... OK
Call fib(3) ... OK
Call fib(4) ... OK
Call fib(5) ... OK
Call fib(7) ... OK
Call fib(8) ... OK

Well done! You seem to have solved fib!
```