# 9. Lists-of-lists (Two-dimensional arrays)

This Jupyter Notebook contains the [problems, or assignments, for Lesson 9](https://snakify.org/en/lessons/two_dimensional_lists_arrays/problems/), gratefully imported from the online course on [Snakify](https://snakify.org). If you get stuck, you can either ask the teacher and the teaching assistents, or skim through the [theory](https://snakify.org/en/lessons/two_dimensional_lists_arrays/) and [explanation of the steps](https://snakify.org/en/lessons/two_dimensional_lists_arrays/steps/1/) on Snakify.

```python
# Please run the contents of this cell.
# Select this cell and press Ctrl-Enter (or Command-Enter on mac).

from assertions import *
from test_cells import *
```

<br />
<br />
<br />
<br />

# 9.1 Maximum

Given a matrix represented as list of list (rows of columns), find the maximum element and return a tuple of indices (`row`, `col`) of its row and column.

If there exist multiple such elements in different rows, return the one having the smallest row index. If there multiple such elements occur on the same row, return the one having the smallest column index.

The function `assignment_9_1()` takes a single parameters `matrix`, which contains a list of lists. 

## Example input
```python
0 3 2 4
2 3 5 5
5 1 2 3
```

## Example return value
```python
(1, 2)
```

<!-- grade_id: assignment-09-01-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 9.1

def assignment_9_1(matrix):
    max_number = []
    for row in matrix: 
        max_number.append(max(row))
    max_number = max(max_number)
    
    for row in matrix: 
        if max_number in row: 
            max_number = (matrix.index(row), row.index(max_number))
            break
    
    return max_number
        
```

```python
# Test for assignment 9.1

test_cell_9_1(assignment_9_1)
```

Output:
```
Please enter a series of integer numbers, separated by spaces (Press Enter for default values 0 3 2 4):  

Working with default matrix: [[0, 3, 2, 4], [2, 3, 5, 5], [5, 1, 2, 3]]

The largest element in this matrix is at position (1, 2).
```

<!-- grade_id: assignment-09-01-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 9.1

check_assertion(assignment_9_1)
```

Output:
```
Call assignment_9_1([[0, 3, 2, 4], [2, 3, 5, 5], [5, 1, 2, 3]]) ... OK
Call assignment_9_1([[1]]) ... OK
Call assignment_9_1([[1, 2, 3, 4, 5], [6, 7, 8, 9, 10], [11, 12, 13, 14, 15]]) ... OK
Call assignment_9_1([[1], [2], [3], [2], [1], [2]]) ... OK
Call assignment_9_1([[1, 2, 3, 3, 2, 1]]) ... OK
Call assignment_9_1([[1, 4, 5, 2, 7], [3, 4, 8, 0, 8], [8, 8, 8, 8, 8]]) ... OK
Call assignment_9_1([[-3, -2], [-2, -1]]) ... OK
Call assignment_9_1([[-1000000003, -1000000002], [-1000000002, -1000000001]]) ... OK

Well done! You seem to have solved assignment_9_1!
```

<br />
<br />
<br />
<br />

# 9.2 Snowflake

Given an odd integer `n`, produce a two-dimensional array of size (`n` x `n`). Fill each element with a single character string of `'.'` . Then fill the middle row, the middle column and the diagnals with the single character string of `'*'` (an image of a snow flake). Return the two-dimensional array. The test cell will print your snowflake for you.

## Example input
```python
5
```

## Example return value
```python
[['*', '.', '*', '.', '*'], ['.', '*', '*', '*', '.'], ['*', '*', '*', '*', '*'], ['.', '*', '*', '*', '.'], ['*', '.', '*', '.', '*']]
```

**NB** This will result in the following being printed in the test cell (you don't have to implement this part):
```python
* . * . *
. * * * .
* * * * *
. * * * .
* . * . *
```

<!-- grade_id: assignment-09-02-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 9.2

def assignment_9_2(n):
    if n == 1: 
        return [["*"]]
    
    pattern = [["." for h in range(n)] for i in range(n)]
    
    for i in range(0, n, n//2): # first and last row
        pattern[0][i] = "*" 
        pattern[-1][i] = "*"
    
    pattern[n//2] = ["*" for i in range(n)] # middle row 
    
    for i in range(1,n//2): # Top half 
        pattern[i][i] = "*" 
        pattern[i][-(i+1)] = "*"
        pattern[i][n//2] = "*"
    
    for i in range(n//2 +1, n): # cloning inverted top half 
        pattern[i] = pattern[n-i-1]
        
    return pattern

```

```python
# Test for assignment 9.2

test_cell_9_2(assignment_9_2)
```

Output:
```
Please enter an odd integer (Press Enter for default value 5):  27

* . . . . . . . . . . . . * . . . . . . . . . . . . * 
. * . . . . . . . . . . . * . . . . . . . . . . . * . 
. . * . . . . . . . . . . * . . . . . . . . . . * . . 
. . . * . . . . . . . . . * . . . . . . . . . * . . . 
. . . . * . . . . . . . . * . . . . . . . . * . . . . 
. . . . . * . . . . . . . * . . . . . . . * . . . . . 
. . . . . . * . . . . . . * . . . . . . * . . . . . . 
. . . . . . . * . . . . . * . . . . . * . . . . . . . 
. . . . . . . . * . . . . * . . . . * . . . . . . . . 
. . . . . . . . . * . . . * . . . * . . . . . . . . . 
. . . . . . . . . . * . . * . . * . . . . . . . . . . 
. . . . . . . . . . . * . * . * . . . . . . . . . . . 
. . . . . . . . . . . . * * * . . . . . . . . . . . . 
* * * * * * * * * * * * * * * * * * * * * * * * * * * 
. . . . . . . . . . . . * * * . . . . . . . . . . . . 
. . . . . . . . . . . * . * . * . . . . . . . . . . . 
. . . . . . . . . . * . . * . . * . . . . . . . . . 
```

<!-- grade_id: assignment-09-02-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 9.2

check_assertion(assignment_9_2)
```

Output:
```
Call assignment_9_2(5) ... OK
Call assignment_9_2(7) ... OK
Call assignment_9_2(9) ... OK
Call assignment_9_2(11) ... OK
Call assignment_9_2(3) ... OK
Call assignment_9_2(1) ... OK

Well done! You seem to have solved assignment_9_2!
```

<br />
<br />
<br />
<br />

# 9.3 Checkered flag

Given two numbers `rows` and `cols`. Create a two-dimensional array of size `rows` x `cols` and populate it with the characters `'.'` and `'*'` to create a checkered flag (see image). The top left corner should have the character `'.'`. Return the two-dimensional array. The test cell will print the checkered flag.

<img src="figures/checkered_flag.png" alt="Checkered flag">

## Example input
```python
3 4
```

## Example return value
```python
[['.', '*', '.', '*'], ['*', '.', '*', '.'], ['.', '*', '.', '*']]
```

**NB** This will result in the following being printed in the test cell (you don't have to implement this part):
```python
. * . *
* . * .
. * . *
```

<!-- grade_id: assignment-09-03-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 9.3

def assignment_9_3(rows, cols):
    checkered_flag = [["." for column in range(cols)] for row in range(rows)] #Creating the initial checkered flag using only "." to be changed to "*"
    
    for row in range(rows):
        if row % 2 == 0:
            for column in range(1, cols, 2):
                checkered_flag[row][column] = "*"
        else:
            for column in range(0, cols, 2):
                checkered_flag[row][column] = "*"
    
    return checkered_flag
            
```

```python
# Test for assignment 9.3

test_cell_9_3(assignment_9_3)
```

Output:
```
Please enter a series of integers (rows and columns), separated by spaces (Press Enter for default values 3 4):  

. * . * 
* . * . 
. * . *
```

<!-- grade_id: assingment-9-3-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 9.3

check_assertion(assignment_9_3)
```

Output:
```
Call assignment_9_3(3, 4) ... OK
Call assignment_9_3(6, 8) ... OK
Call assignment_9_3(8, 3) ... OK
Call assignment_9_3(7, 5) ... OK
Call assignment_9_3(1, 6) ... OK
Call assignment_9_3(8, 1) ... OK
Call assignment_9_3(1, 1) ... OK

Well done! You seem to have solved assignment_9_3!
```

<br />
<br />
<br />
<br />

# 9.4 The diagonals parallel to the main

Given an integer `size`, produce a two-dimensional array of size `size` times `size` and complete it according to the following rules:
* On the main diagonal, which is from top-left to bottom-right corner, have numbers `0`
* On the diagonals adjacent to the main, have numbers `1`
* On the next adjacent diagonals have numbers `2`, and so forth

Return the resulting two-dimensional array.

## Example input
```python
5
```

## Example return value
```python
[[0, 1, 2, 3, 4],
 [1, 0, 1, 2, 3],
 [2, 1, 0, 1, 2],
 [3, 2, 1, 0, 1],
 [4, 3, 2, 1, 0]]
```

**NB** This will result in the following being printed in the test cell (you don't have to implement this part):
```python
0 1 2 3 4
1 0 1 2 3
2 1 0 1 2
3 2 1 0 1
4 3 2 1 0
```

<!-- grade_id: assignment-09-04-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 9.4

def assignment_9_4(size):
    #Creating initial matrix populated of zeroes
    matrix = [[0 for column in range(size)] for row in range(size)]
    
    #Creating the right half of the matrix
    counter = 0
    for row in matrix:
        row[counter:] = [x for x in range(size - counter)]
        counter += 1
        

        
    #Creating the left half of the matrix
    counter = size
    for row in matrix[::-1]: 
        row[:counter] = [x for x in range(counter - 1, -1, -1)]
        counter -= 1
        
    return matrix
        
        
```

```python
# Test for assignment 9.4

test_cell_9_4(assignment_9_4)
```

Output:
```
Please enter the dimension of the matrix (Press Enter for default value 5):  

0 1 2 3 4 
1 0 1 2 3 
2 1 0 1 2 
3 2 1 0 1 
4 3 2 1 0
```

<!-- grade_id: assignment-09-04-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 9.4

check_assertion(assignment_9_4)
```

Output:
```
Call assignment_9_4(5) ... OK
Call assignment_9_4(6) ... OK
Call assignment_9_4(7) ... OK
Call assignment_9_4(4) ... OK
Call assignment_9_4(3) ... OK
Call assignment_9_4(2) ... OK
Call assignment_9_4(1) ... OK

Well done! You seem to have solved assignment_9_4!
```

<br />
<br />
<br />
<br />

# 9.5 Minor diagonal

Given an integer `size`, create a two-dimensional array of size `size` x `size` and populate it as follows:

* The elements on the minor diagonal (from the lower left corner to the upper right corner) should have `1`s
* The elements above this diagonal should have `0`s
* The elements below the diagonal should have `2`s

Return the resulting two-dimensional array.

## Example input
```python
4
```

## Example return value
```
[[0, 0, 0, 1], [0, 0, 1, 2], [0, 1, 2, 2], [1, 2, 2, 2]]
```

**NB** This will result in the following print in the test cell (you don't have to implement this part):
```python
0 0 0 1
0 0 1 2
0 1 2 2
1 2 2 2
```


<!-- grade_id: assignment-09-05-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 9.5

def assignment_9_5(size):
    #Creating initial matrix populated of zeroes
    matrix = [[1 for column in range(size)] for row in range(size)]
    
    #Creating the right half of the matrix
    counter = 1
    for row in matrix[1:]:
        row[-counter:] = [2 for x in range(counter)]
        counter += 1
        
    #Creating the left half of the matrix
    counter = 1
    for row in matrix[:-1]:  
        row[:-counter] = [0 for x in range(size - counter)]
        counter += 1
        
    return matrix
```

```python
# Test for assignment 9.5

test_cell_9_5(assignment_9_5)
```

Output:
```
Please enter the dimension of the matrix (Press Enter for default value 4):  

0 0 0 1 
0 0 1 2 
0 1 2 2 
1 2 2 2
```

<!-- grade_id: assignment-09-05-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 9.5

check_assertion(assignment_9_5)
```

Output:
```
Call assignment_9_5(4) ... OK
Call assignment_9_5(1) ... OK
Call assignment_9_5(2) ... OK
Call assignment_9_5(3) ... OK
Call assignment_9_5(5) ... OK
Call assignment_9_5(6) ... OK
Call assignment_9_5(7) ... OK
Call assignment_9_5(8) ... OK

Well done! You seem to have solved assignment_9_5!
```

<br />
<br />
<br />
<br />

# 9.6 Swap the columns

Given a matrix `matrix` of dimension `rows` x `cols`, write a function `swap_columns(matrix, col_idx_1, col_idx_2)` that will swap column with index `col_idx_1` with column with index `col_idx_2` in `matrix`.

In the test cell, first enter `rows` and `cols` on a single line, separated by spaces. Then enter the elements for matrix `matrix` numbers for a single row separated by spaced. Finally, enter `col_idx_1` and `col_idx_2`, separated by spaces.

## Example input
```python
3 4
11 12 13 14
21 22 23 24
31 32 33 34
0 1
```

## Example return value
```python
[[12, 11, 13, 14], [22, 21, 23, 24], [32, 31, 33, 34]]
```

**NB** This will result in the following print in the test cell (you don't have to implement this part):

```python
12 11 13 14
22 21 23 24
32 31 33 34
```

<!-- grade_id: assignment-09-06-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 9.6

# YOUR CODE HERE
def swap_columns(matrix, col_idx_1, col_idx_2): 
    
    for row in matrix: 
        transit = row[col_idx_1]
        row[col_idx_1] = row[col_idx_2]
        row[col_idx_2] = transit
        
    return matrix
    


        
    
```

```python
# Test for assignment 9.6

test_cell_9_6(swap_columns)
```

Output:
```
Please enter a series of integer numbers, separated by spaces (Press Enter for default values 11 12 13 14):  

Working with default matrix: [[11, 12, 13, 14], [21, 22, 23, 24], [31, 32, 33, 34]]
Please enter a series of indexes of two columns, separated by spaces (Press Enter for default values 0 1):  

12 11 13 14 
22 21 23 24 
32 31 33 34
```

<!-- grade_id: assignment-09-06-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 9.6

check_assertion(swap_columns)
```

Output:
```
Call swap_columns([[11, 12, 13, 14], [21, 22, 23, 24], [31, 32, 33, 34]], 0, 1) ... OK
Call swap_columns([[1, 2], [3, 4]], 0, 1) ... OK
Call swap_columns([[1], [2], [3], [4], [5], [6], [7], [8], [9], [10]], 0, 0) ... OK
Call swap_columns([[1, 2, 3], [4, 5, 6], [7, 8, 9], [10, 11, 12]], 1, 2) ... OK
Call swap_columns([[1, 2, 3, 4, 5]], 1, 4) ... OK
Call swap_columns([[179]], 0, 0) ... OK
Call swap_columns([[11, 12, 13, 14, 15, 16], [21, 22, 23, 24, 25, 26], [31, 32, 33, 34, 35, 36], [41, 42, 43, 44, 45, 46]], 0, 5) ... OK
Call swap_columns([[11, 12, 13, 14, 15, 16], [21, 22, 23, 24, 25, 26], [31, 32, 33, 34, 35, 36], [41, 42, 43, 44, 45, 46]], 5, 0) ... OK
Call swap_columns([[11, 12, 13, 14, 15, 16], [21, 22, 23, 24, 25, 26], [31, 32, 33, 34, 35, 36], [41, 42, 43, 44, 45, 46]], 4, 5) ... OK
Call swap_columns([[11, 12, 13, 14, 15, 16], [21, 22, 23, 24, 25, 26], [31, 32, 33, 34, 35, 36], [41, 42, 43, 44, 45, 46]], 1, 4) ... OK
Call swap_columns([[11, 12, 13, 14, 15, 16], [21, 22, 23, 24, 
```

<br />
<br />
<br />
<br />

# 9.7 Scale a matrix

Given a matrix `matrix` of dimension `rows` x `cols`, finalize `assignment_9_7(matrix, constant)` in such a way that it returns a two-dimensional list where all elements in `matrix` have been multiplied by `constant`.

## Example input
```python
3 4
11 12 13 14
21 22 23 24
31 32 33 34
2
```

## Example return value
```python
[[22, 24, 26, 28], [42, 44, 46, 48], [62, 64, 66, 68]]
```

**NB** This will result in the following print in the test cell (you don't have to implement this part):
```python
22 24 26 28
42 44 46 48
62 64 66 68
```


<!-- grade_id: assignment-09-07-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 9.7

def assignment_9_7(matrix, constant):
    matrix = [[constant * column for column in row] for row in matrix]
    return matrix
```

```python
# Test for assignment 9.7

test_cell_9_7(assignment_9_7)
```

Output:
```
Please enter a series of integer numbers, separated by spaces (Press Enter for default values 11 12 13 14):  

Working with default matrix: [[11, 12, 13, 14], [21, 22, 23, 24], [31, 32, 33, 34]]
Please enter multiplication constant (Press Enter for default value 2):  

22 24 26 28 
42 44 46 48 
62 64 66 68
```

<!-- grade_id: assignment-09-07-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 9.7

check_assertion(assignment_9_7)
```

Output:
```
Call assignment_9_7([[11, 12, 13, 14], [21, 22, 23, 24], [31, 32, 33, 34]], 2) ... OK
Call assignment_9_7([[1, 2], [3, 4]], 0) ... OK
Call assignment_9_7([[1], [2], [3], [4], [5], [6], [7], [8], [9], [10]], -1) ... OK
Call assignment_9_7([[1, 2, 3], [4, 5, 6], [7, 8, 9], [10, 11, 12]], 1) ... OK
Call assignment_9_7([[9]], 9) ... OK

Well done! You seem to have solved assignment_9_7!
```

<br />
<br />
<br />
<br />

# 9.8 Multiply two matrices

Given three positive integers $l$, $m$ and $n$, where matrix $A$ has dimensions $l \times m$, and matrix $B$ has dimensions $m \times n$. Multiplication of $A \times B$ results in matrix $C$ with dimensions $l \times n$ (see [Wikipedia](https://en.wikipedia.org/wiki/Matrix_multiplication)).

Finalize the function `assignment_9_8(a, b)` in such a way that it will return $C$, a two-dimensional list with the result of the multiplication $A \times B$.

The test cell will request $l$, $m$ and $n$ seperated by spaces first. After that it will request $l$ rows of $m$ numbers for $A$, separated by spaces. After that, it will request $m$ rows of $n$ elements for $B$, separated by spaces.

**Hint** The function `assignment_9_8(a, b)` only takes the 2 two-dimensional lists as inputs. You can obtain $l$, $m$ and $n$ and create an "empty" $C$ inside the function with:
```python
l = len(a)
m = len(a[0])
n = len(b[0])
c = [[None for j in range(n)] for i in range(l)]
```

## Example input

You will be requested to enter two matrices A and B in the test cell. If you press Enter on all input questions you will be using the default values below:

```python
3 4               # <--- dimension matrix A
0 1 2 3           
4 5 6 7
8 9 10 11         

4 2               # <--- dimension matrix B
2 3               
0 4
5 -1
1 1               
```

This corresponds to:
$$A \times B = \begin{bmatrix} 0 & 1 & 2 & 3 \\ 4 & 5 & 6 & 7 \\ 8 & 9 & 10 & 11 \end{bmatrix} \times \begin{bmatrix} 2 & 3 \\ 0 & 4 \\ 5 & -1 \\ 1 & 1 \end{bmatrix}$$

## Example return value
```python
[[13, 5], [45, 33], [77, 61]]
```

This corresponds to:
$$C = A \times B = \begin{bmatrix} 13 & 5 \\ 45 & 33 \\ 77 & 61 \end{bmatrix}$$

**NB** The test cell will print:
```python
13 5
45 33
77 61
```

<!-- grade_id: assignment-09-08-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 9.8

def assignment_9_8(a, b):
    l = len(a)
    m = len(a[0])
    n = len(b[0])
    c = [[0 for j in range(n)] for i in range(l)]
    
    row_idx_c = -1
    
    for row_a in a: 
        row_idx_c +=1
        
        for col_idx_b in range(n):
            col_idx_c = col_idx_b
            col_idx_a = 0
            for row_b in b:
                c[row_idx_c][col_idx_c] += row_a[col_idx_a] * row_b[col_idx_b]
                col_idx_a += 1
            
    return c    
```

```python
# Test for assignment 9.8

test_cell_9_8(assignment_9_8)
```

Output:
```
--- Matrix A ---

Please enter a series of integer numbers, separated by spaces (Press Enter for default values 0 1 2 3):  

Working with default matrix: [[0, 1, 2, 3], [4, 5, 6, 7], [8, 9, 10, 11]]

--- Matrix B ---

Please enter a series of integer numbers, separated by spaces (Press Enter for default values 2 3):  

Working with default matrix: [[2, 3], [0, 4], [5, -1], [1, 1]]

The product of these two matrices is:
13 5 
45 33 
77 61
```

<!-- grade_id: assignment-09-08-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 9.8

check_assertion(assignment_9_8)
```

Output:
```
Call assignment_9_8([[0, 1, 2, 3], [4, 5, 6, 7], [8, 9, 10, 11]], [[2, 3], [0, 4], [5, -1], [1, 1]]) ... OK
Call assignment_9_8([[1, 2], [3, 4]], [[5, 4], [3, 2]]) ... OK
Call assignment_9_8([[5]], [[7]]) ... OK
Call assignment_9_8([[1, 2, 3], [4, 5, 6], [7, 8, 9], [10, 11, 12]], [[1], [2], [3]]) ... OK

Well done! You seem to have solved assignment_9_8!
```