# 1. Input, print and numbers

This Jupyter Notebook contains the [problems, or assignments, for Lesson 1](https://snakify.org/en/lessons/print_input_numbers/problems/), gratefully imported from the online course on [Snakify](https://snakify.org). If you get stuck, you can either ask the teacher and the teaching assistents, or skim through the [theory](https://snakify.org/en/lessons/print_input_numbers/) and [explanation of the steps](https://snakify.org/en/lessons/print_input_numbers/steps/1/) on Snakify.

```python
# Please run the contents of this cell.
# Select this cell and press Ctrl-Enter (or Command-Enter on mac).

from assertions import *
from test_cells import *
```

Output:
```
Assertions imported OK
Test cells imported OK
```

<br />
<br />
<br />
<br />

# 1.1 Sum of three numbers

Finalize the function below that has three integer arguments `a`, `b` and `c`, and return their sum.

## Example input

```python
2
3
6
```

## Example return value

```python
11
```

<!-- grade_id: assignment-01-01-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 1.1

def assignment_1_1(a, b, c):
    # Hint: add c to "return a + b"
    
    # YOUR CODE HERE
    return a + b + c
```

```python
# Test for assignment 1.1

test_cell_1_1(assignment_1_1)
```

Output:
```
Please enter integer value 1 (Press Enter for default value 2):  2
Please enter integer value 2 (Press Enter for default value 3):  3
Please enter integer value 3 (Press Enter for default value 6):  4

The sum of the integers 2, 3 and 4 is:
9
```

<!-- grade_id: assignment-01-01-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 1.1

check_assertion(assignment_1_1)
```

Output:
```
Call assignment_1_1(2, 3, 6) ... OK
Call assignment_1_1(0, 20, 300) ... OK
Call assignment_1_1(-5, 180, -17) ... OK

Well done! You seem to have solved assignment_1_1!
```

<br />
<br />
<br />
<br />

# 1.2 Hi John

Finalize the function below that takes 1 string argument `name` in such a way that it returns **"Hi "** followed by the name of the person. See the example below.

## Example input

```python
John
```

## Example return value

```python
Hi John
```

<!-- grade_id: assignment-01-02-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 1.2

def assignment_1_2(name):
    # Hint: Do not forget the space between 'Hi' and name
    
    # YOUR CODE HERE
    return f"Hi {name}"
```

```python
# Test for assignment 1.2

test_cell_1_2(assignment_1_2)
```

Output:
```
Please enter a name (Press Enter for default "John"):  Majdi

Hi Majdi
```

<!-- grade_id: assignment-01-02-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 1.2

check_assertion(assignment_1_2)
```

Output:
```
Call assignment_1_2('John') ... OK
Call assignment_1_2('Alice') ... OK
Call assignment_1_2('Bob') ... OK
Call assignment_1_2('Tanner') ... OK
Call assignment_1_2('Jake') ... OK

Well done! You seem to have solved assignment_1_2!
```

<br />
<br />
<br />
<br />

# 1.3 Square

Finalize the function below that has 1 integer argument `number` in such a way that it returns the square of `number`.

## Example input

```python
5
```

## Example return value

```python
25
```

<!-- grade_id: assignment-01-03-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 1.3

def assignment_1_3(number):
    # YOUR CODE HERE
    return number * number
```

```python
# Test for assignment 1.3

test_cell_1_3(assignment_1_3)
```

Output:
```
Please enter an integer to be squared (Press Enter for default value 5):  9

The square of the integer 9 is:
81
```

<!-- grade_id: assignment-01-03-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 1.3

check_assertion(assignment_1_3)
```

Output:
```
Call assignment_1_3(5) ... OK
Call assignment_1_3(6) ... OK
Call assignment_1_3(2) ... OK
Call assignment_1_3(1) ... OK
Call assignment_1_3(10) ... OK
Call assignment_1_3(9) ... OK
Call assignment_1_3(7) ... OK
Call assignment_1_3(200) ... OK

Well done! You seem to have solved assignment_1_3!
```

<br />
<br />
<br />
<br />

# 1.4 Area of right-angled triangle

Finalize the function below that has two arguments `base` and `height` representing the length of the base and the height of a right-angled triangle and make it return the area. 

Use formula:

```
area = 0.5 * base * height
```

## Example input

```python
3
5
```

## Example return value

```python
7.5
```

<!-- grade_id: assignment-01-04-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 1.4

def assignment_1_4(base, height):
    # YOUR CODE HERE
    return 0.5 * base * height
```

```python
# Test for assignment 1.4

test_cell_1_4(assignment_1_4)
```

Output:
```
Please enter the base (Press Enter for default value 3):  
Please enter the height (Press Enter for default value 5):  

The area of a triangle with base 3 and height 5 is:
7.5
```

<!-- grade_id: assignment-01-04-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 1.4

check_assertion(assignment_1_4)
```

Output:
```
Call assignment_1_4(3, 5) ... OK
Call assignment_1_4(10, 2) ... OK
Call assignment_1_4(179, 1534) ... OK
Call assignment_1_4(1543, 57) ... OK

Well done! You seem to have solved assignment_1_4!
```

<br />
<br />
<br />
<br />

# 1.5 Hello, Harry!

Finalize the function below that has 1 string argument `name` in such a way that it returns a string combining the word "Hello", a comma, the name of the user and an exclamation mark after it. See the examples below.

**Warning.** Your program's return value should strictly match the desired one, character by character. There shouldn't be any space between the name and the exclamation mark. You can use `+` operator to concatenate two strings. See the lesson for details.

## Example input

```python
Harry
```

## Example return value

```python
Hello, Harry!
```

<!-- grade_id: assignment-01-05-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 1.5

def assignment_1_5(name):
    # YOUR CODE HERE
    return f"Hello, {name}!" 
```

```python
# Test for assignment 1.5

test_cell_1_5(assignment_1_5)
```

Output:
```
Please enter a name (Press Enter for default "Harry"):  

Hello, Harry!
```

<!-- grade_id: assignment-01-05-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 1.5

check_assertion(assignment_1_5)
```

Output:
```
Call assignment_1_5('Harry') ... OK
Call assignment_1_5('Mr. Potter') ... OK
Call assignment_1_5('Lord Voldemort') ... OK

Well done! You seem to have solved assignment_1_5!
```

<br />
<br />
<br />
<br />

# 1.6 Apple sharing

A number of students (`students`) take a number of apples (`apples`) and distribute them among each other evenly. The remaining (the undivisible) part remains in the basket. How many apples will each single student get? How many apples will remain in the basket?

The function below has two integer arguments `students` and `apples`. It should return the two answers for the questions above. You can separate both values following your return statement with a comma, like:
```python
return a, b
```

**NB** It will return the values `a` and `b` as a tuple (to be recognized by the parentheses), which is a collection datatype. Collection datatypes will be explained in session 2 of this course.

## Example input

```python
6
50
```

## Example return value

```python
(8, 2)
```

<!-- grade_id: assignment-01-06-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 1.6

def assignment_1_6(students, apples):
    # Hint 1: use integer division with // and remainder with %
    # Hint 2: use return a, b to return both 
    
    # YOUR CODE HERE
    return (apples // students, apples % students) 
```

```python
# Test for assignment 1.6

test_cell_1_6(assignment_1_6)
```

Output:
```
Please enter number of students (Press Enter for default value 6):  
Please enter number of apples (Press Enter for default value 50):  

The number of apples per student, and the number of remaining appels, are:
(8, 2)
```

<!-- grade_id: assignment-01-06-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 1.6

check_assertion(assignment_1_6)
```

Output:
```
Call assignment_1_6(6, 50) ... OK
Call assignment_1_6(1, 10) ... OK
Call assignment_1_6(5, 25) ... OK
Call assignment_1_6(4, 2) ... OK

Well done! You seem to have solved assignment_1_6!
```

<br />
<br />
<br />
<br />

# 1.7 Previous and next
Finalize the function below that has 1 integer argument `number` and make it **print** its previous and next numbers shaped in a nice sentence. See the examples below for the exact format your answers should take. There shouldn't be a space before the period.

Remember that you can convert the numbers to strings using the function `str`.

**NB** Please note that your function should **print** the strings rather than **return** the values.

## Example input

```python
179
```

## Example output (**printed**)

```python
The number before 179 is 178.
The number after 179 is 180.
```

<!-- grade_id: assignment-01-07-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 1.7

def assignment_1_7(number):
    # YOUR CODE HERE
    next_number = number + 1
    previous_number = number - 1 
    print(f"The number before {number} is {previous_number}.")
    print(f"The number after {number} is {next_number}.")
```

```python
# Test for assignment 1.7

test_cell_1_7(assignment_1_7)
```

Output:
```
Please enter an integer value (Press Enter for default value 179):  

The number before 179 is 178.
The number after 179 is 180.
```

<!-- grade_id: assignment-01-07-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 1.7

check_assertion(assignment_1_7)
```

Output:
```
Call assignment_1_7(179) ... OK
Call assignment_1_7(0) ... OK
Call assignment_1_7(100) ... OK

Well done! You seem to have solved assignment_1_7!
```

<br />
<br />
<br />
<br />

# 1.8 Two timestamps

A timestamp is three numbers: a number of hours, minutes and seconds. Given two timestamps, calculate how many seconds is between them and return that value. The moment of the first timestamp occurred before the moment of the second timestamp.

## Example input

```python
1
1
1
2
2
2
```

## Example return value

```python
3661
```

<!-- grade_id: assignment-01-08-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 1.8

def assignment_1_8(hour1, minute1, second1, hour2, minute2, second2):
    # Hint: convert both h1:m1:s1 and h2:m2:s2 to seconds and calculate the difference
    # YOUR CODE HERE
    return (hour2 * 3600 + minute2 * 60 + second2) - (hour1 * 3600 + minute1 * 60 + second1)
```

```python
# Test for assignment 1.8

test_cell_1_8(assignment_1_8)
```

Output:
```
Please enter hours for time 1 (Press Enter for default value 1):  
Please enter minutes for time 1 (Press Enter for default value 1):  
Please enter seconds for time 1 (Press Enter for default value 1):  
Please enter hours for time 2 (Press Enter for default value 2):  
Please enter minutes for time 2 (Press Enter for default value 2):  
Please enter seconds for time 2 (Press Enter for default value 2):  

The difference, in seconds, between 01:01:01 and 02:02:02 is:
3661
```

<!-- grade_id: assignment-01-08-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 1.8

check_assertion(assignment_1_8)
```

Output:
```
Call assignment_1_8(1, 1, 1, 2, 2, 2) ... OK
Call assignment_1_8(6, 1, 30, 6, 2, 10) ... OK
Call assignment_1_8(10, 8, 2, 10, 8, 2) ... OK
Call assignment_1_8(2, 3, 4, 4, 7, 28) ... OK
Call assignment_1_8(5, 50, 48, 10, 20, 35) ... OK
Call assignment_1_8(1, 15, 39, 6, 15, 39) ... OK
Call assignment_1_8(11, 7, 29, 11, 30, 29) ... OK
Call assignment_1_8(5, 34, 4, 5, 34, 50) ... OK
Call assignment_1_8(4, 19, 21, 7, 39, 45) ... OK

Well done! You seem to have solved assignment_1_8!
```

<br />
<br />
<br />
<br />

# 1.9 School desks

A school decided to replace the desks in three classrooms. Each desk sits two students. Given the number of students in each class, return the smallest possible number of desks that can be purchased.

The function below uses three integer argumentes: the number of students in each of the three classes, `students_a`, `students_b` and `students_c` respectively.

In the example below there are three groups. The first group has 20 students and thus needs 10 desks. The second group has 21 students, so they can get by with no fewer than 11 desks. 11 desks is also enough for the third group of 22 students. So we need 32 desks in total.

## Example input

```python
20
21
22
```

## Example return value

```python
32
```

<!-- grade_id: assignment-01-09-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 1.9

def assignment_1_9(students_a, students_b, students_c):
    # Hint: use integer division with // and remainder with %
    
    # YOUR CODE HERE
    return (students_a // 2 + students_a % 2) + (students_b // 2 + students_b % 2) + (students_c // 2 + students_c % 2) 
```

```python
# Test for assignment 1.9

test_cell_1_9(assignment_1_9)
```

Output:
```
Please enter number of students in class A (Press Enter for default value 20):  
Please enter number of students in class B (Press Enter for default value 21):  
Please enter number of students in class C (Press Enter for default value 22):  

The total number of desks needed is:
32
```

<!-- grade_id: assignment-01-09-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 1.9

check_assertion(assignment_1_9)
```

Output:
```
Call assignment_1_9(20, 21, 22) ... OK
Call assignment_1_9(1, 1, 1) ... OK
Call assignment_1_9(26, 20, 16) ... OK
Call assignment_1_9(25, 21, 23) ... OK
Call assignment_1_9(17, 19, 18) ... OK

Well done! You seem to have solved assignment_1_9!
```