# 4. For loop with range

This Jupyter Notebook contains the [problems, or assignments, for Lesson 4](https://snakify.org/en/lessons/for_loop_range/problems/), gratefully imported from the online course on [Snakify](https://snakify.org). If you get stuck, you can either ask the teacher and the teaching assistents, or skim through the [theory](https://snakify.org/en/lessons/for_loop_range) and [explanation of the steps](https://snakify.org/en/lessons/for_loop_range/steps/1/) on Snakify.

```python
# Please run the contents of this cell.
# Select this cell and press Ctrl-Enter (or Command-Enter on mac).

from assertions import *
from test_cells import *
```

<br />
<br />
<br />
<br />

# 4.1 Series - 1

Given two integers `start` and `end` (`start` ≤ `end`). Return a list `result` that contains all numbers from `start` to `end`, including `start` and `end` themselves.

## Example input

```python
1
10
```

## Example return value

```python
[1, 2, 3, 4, 5, 6, 7, 8, 9, 10]
```

<!-- grade_id: assignment-04-01-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 4.1

def assignment_4_1(start, end):
    list = []
    for i in range(start, end + 1):
        list.append(i) 
    return list 
        
```

```python
# Test for assignment 4.1

test_cell_4_1(assignment_4_1)
```

Output:
```
Please enter value for start (Press Enter for default value 1):  
Please enter value for end (Press Enter for default value 10):  

The list with all numbers from 1 to 10 (inclusive) is: 
[1, 2, 3, 4, 5, 6, 7, 8, 9, 10]
```

<!-- grade_id: assignment-04-01-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 4.1

check_assertion(assignment_4_1)
```

Output:
```
Call assignment_4_1(1, 10) ... OK
Call assignment_4_1(-3, 14) ... OK
Call assignment_4_1(0, 0) ... OK
Call assignment_4_1(20, 20) ... OK

Well done! You seem to have solved assignment_4_1!
```

<br />
<br />
<br />
<br />

# 4.2 Series - 2

Given two integers `start` and `end`. Create a list `result` that contains all numbers from `start` to `end` in ascending order, if `start` < `end`, or in descending order, if `start` ≥ `end`. The range should include both `start` and `end`. Return `result` at the end of your function.

## Example input

```python
10
1
```

## Example return value

```python
[10, 9, 8, 7, 6, 5, 4, 3, 2, 1]
```

<!-- grade_id: assignment-04-02-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 4.2

def assignment_4_2(start, end):
    list = [] 
    if start < end:
        for i in range(start, end + 1): 
            list.append(i)
    else: 
        for i in range (start, end - 1, -1):
            list.append(i) 
    return list
```

```python
# Test for assignment 4.2

test_cell_4_2(assignment_4_2)
```

Output:
```
Please enter value for start (Press Enter for default value 10):  
Please enter value for end (Press Enter for default value 1):  

The list with all numbers from 10 to 1 (inclusive) are: 
[10, 9, 8, 7, 6, 5, 4, 3, 2, 1]
```

<!-- grade_id: assignment-04-02-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 4.2

check_assertion(assignment_4_2)
```

Output:
```
Call assignment_4_2(1, 10) ... OK
Call assignment_4_2(10, 1) ... OK
Call assignment_4_2(179, 179) ... OK
Call assignment_4_2(-14, 7) ... OK
Call assignment_4_2(12, -5) ... OK
Call assignment_4_2(-3, -7) ... OK

Well done! You seem to have solved assignment_4_2!
```

<br />
<br />
<br />
<br />

# 4.3 Sum of ten numbers

The starting number `start` of a range is given. Return the sum of the 10 numbers in the range:

`start`, `start + 1`, ... , `start + 9`. 

## Example input

```python
0
```

## Example return value

```python
45
```

<!-- grade_id: assignment-04-03-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 4.3

def assignment_4_3(start):
    return int((10*(2 * start + 9))/2)
```

```python
# Test for assignment 4.3

test_cell_4_3(assignment_4_3)
```

Output:
```
Please enter start number of the range (Press Enter for default value 0):  

The sum of 0 and the 9 following numbers is:
45
```

<!-- grade_id: assignment-04-03-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 4.3

check_assertion(assignment_4_3)
```

Output:
```
Call assignment_4_3(0) ... OK
Call assignment_4_3(3) ... OK
Call assignment_4_3(-8) ... OK
Call assignment_4_3(203) ... OK

Well done! You seem to have solved assignment_4_3!
```

<br />
<br />
<br />
<br />

# 4.4 Sum of N numbers

The starting number `start` of a range is given, together with `number`, the number of numbers in the sum. Return the sum of the `number` numbers in the range `start`, `start + 1`, ... , `start + (number - 1)`.

## Example input

```python
0
10
```

## Example return value

```python
45
```

<!-- grade_id: assignment-04-04-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 4.4

def assignment_4_4(start, number):
    return int((number*(2 * start + (number-1)))/2)
```

```python
# Test for assignment 4.4

test_cell_4_4(assignment_4_4)
```

Output:
```
Please enter the start number of the range (start) (Press Enter for default value 0):  
Please enter the number of numbers in the range (number) (Press Enter for default value 10):  

The sum of 0 and the 9 following numbers is:
45
```

<!-- grade_id: assignment-04-04-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 4.4

check_assertion(assignment_4_4)
```

Output:
```
Call assignment_4_4(0, 10) ... OK
Call assignment_4_4(10, 20) ... OK
Call assignment_4_4(-12, 80) ... OK
Call assignment_4_4(390, 15) ... OK
Call assignment_4_4(12, 1) ... OK

Well done! You seem to have solved assignment_4_4!
```

<br />
<br />
<br />
<br />

# 4.5 Sum of cubes

For the given integer $n$ calculate the following sum:

${1}^{3} + {2}^{3} + \cdots + {n}^{3}$

## Example input

```python
3
```

## Example return value

```python
36
```

<!-- grade_id: assignment-04-05-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 4.5

def assignment_4_5(n):
    sum = 0
    for i in range(n+1): 
        sum += i**3 
    return sum 
        
```

```python
# Test for assignment 4.5

test_cell_4_5(assignment_4_5)
```

Output:
```
Please enter the value for n (Press Enter for default value 3):  

The total sum for the cubes from 1 to 3 is: 
36
```

<!-- grade_id: assignment-04-05-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 4.5

check_assertion(assignment_4_5)
```

Output:
```
Call assignment_4_5(3) ... OK
Call assignment_4_5(1) ... OK
Call assignment_4_5(2) ... OK
Call assignment_4_5(4) ... OK
Call assignment_4_5(5) ... OK
Call assignment_4_5(6) ... OK
Call assignment_4_5(7) ... OK
Call assignment_4_5(8) ... OK
Call assignment_4_5(9) ... OK
Call assignment_4_5(20) ... OK
Call assignment_4_5(30) ... OK

Well done! You seem to have solved assignment_4_5!
```

<br />
<br />
<br />
<br />

# 4.6 Factorial

In mathematics, the factorial of a positive integer $n$, denoted by $n!$ is the following product:

$n! = 1 \times 2 \times \cdots \times n$

By convention, $0!$ is defined to be $1$. For the given integer n calculate the value $n!$.

Don't use math module in this exercise.

## Example input

```python
4
```

## Example return value

```python
24
```

<!-- grade_id: assignment-04-06-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 4.6

def assignment_4_6(n):
    if n == 0:
        return 1 
    else: 
        result = 1 
        i = n 
        while i != 1: 
            result = result * i
            i = i - 1 
        return result 
```

```python
# Test for assignment 4.6

test_cell_4_6(assignment_4_6)
```

Output:
```
Please enter the value for n (Press Enter for default value 4):  

The value of 4! is:
24
```

<!-- grade_id: assignment-04-06-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 4.6

check_assertion(assignment_4_6)
```

Output:
```
Call assignment_4_6(4) ... OK
Call assignment_4_6(1) ... OK
Call assignment_4_6(2) ... OK
Call assignment_4_6(3) ... OK
Call assignment_4_6(5) ... OK
Call assignment_4_6(6) ... OK
Call assignment_4_6(7) ... OK
Call assignment_4_6(8) ... OK
Call assignment_4_6(9) ... OK
Call assignment_4_6(10) ... OK
Call assignment_4_6(11) ... OK
Call assignment_4_6(12) ... OK
Call assignment_4_6(0) ... OK

Well done! You seem to have solved assignment_4_6!
```

<br />
<br />
<br />
<br />

## 4.7 The number of multiples of 3

Given the start of range with `start`, containing `number` numbers: `start`, `start + 1`, ..., `start + (number - 1)`, count how many of these are multiples of 3. Examples are 3, 6, 9, 12, etc.

## Example input

```python
10
20
```

## Example return value

```python
6
```


<!-- grade_id: assignment-04-07-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 4.7

def assignment_4_7(start, number):
    mod_3 = 0
    for i in range (start, start + number): 
        if i % 3 == 0: 
            mod_3 += 1 
    return mod_3
```

```python
# Test for assignment 4.7

test_cell_4_7(assignment_4_7)
```

Output:
```
Please enter the start number of the range (start) (Press Enter for default value 10):  
Please enter the number of numbers in the range (number) (Press Enter for default value 20):  

The number of multiples of 3 for a range of 20 numbers starting with 10 is:
6
```

<!-- grade_id: assignment-04-07-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 4.7

check_assertion(assignment_4_7)
```

Output:
```
Call assignment_4_7(10, 20) ... OK
Call assignment_4_7(8, 13) ... OK
Call assignment_4_7(0, 0) ... OK
Call assignment_4_7(25, 250) ... OK
Call assignment_4_7(-100, 300) ... OK

Well done! You seem to have solved assignment_4_7!
```

<br />
<br />
<br />
<br />

# 4.8 Adding factorials

Given a positive integer $n$, print the sum $1! + 2! + 3! + \cdots +n!$.

This problem has a solution with only one loop, so try to discover it. And don't use the math library :)

## Example input

```python
4
```

## Example return value

```python
33
```

<!-- grade_id: assignment-04-08-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 4.8

def assignment_4_8(n):
    sum = 0 
    fac_i = 1
    for i in range (1, n + 1): 
        fac_i = fac_i * i
        sum += fac_i
    return sum 
```

```python
# Test for assignment 4.8

test_cell_4_8(assignment_4_8)
```

Output:
```
Please enter integer value n (Press Enter for default value 4):  

The value of the sum of the factorials for 4 is:
33
```

<!-- grade_id: assignment-04-08-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 4.8

check_assertion(assignment_4_8)
```

Output:
```
Call assignment_4_8(1) ... OK
Call assignment_4_8(2) ... OK
Call assignment_4_8(3) ... OK
Call assignment_4_8(4) ... OK
Call assignment_4_8(5) ... OK
Call assignment_4_8(6) ... OK
Call assignment_4_8(7) ... OK
Call assignment_4_8(8) ... OK
Call assignment_4_8(9) ... OK
Call assignment_4_8(10) ... OK

Well done! You seem to have solved assignment_4_8!
```

<br />
<br />
<br />
<br />

# 4.9 Ladder

For given integer `steps` ≤ 9 return a list with "ladders" of `steps` steps. The ladder always starts with **1**, the next step is achieved by adding a 2 to the end, resulting in **12**. The third step is adding a 3 to the end, resulting in **123**.

Store all ladders in a list called `result` and return that.

**NB** Please note that all items in the list should be integers!

## Example input

```python
3
```

## Example return value

```python
[1, 12, 123]
```

<!-- grade_id: assignment-04-09-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 4.9

def assignment_4_9(steps):
    list = [] 
    num_str = ""
    for i in range (1, steps+1): 
        num_str = num_str + str(i) 
        list.append(int(num_str))
    return list 
```

```python
# Test for assignment 4.9

test_cell_4_9(assignment_4_9)
```

Output:
```
Please enter an integer for steps (≤ 9) (Press Enter for default value 3):  

The output for 3 steps is: 
[1, 12, 123]
```

<!-- grade_id: assignment-04-09-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 4.9

check_assertion(assignment_4_9)
```

Output:
```
Call assignment_4_9(3) ... OK
Call assignment_4_9(4) ... OK
Call assignment_4_9(2) ... OK
Call assignment_4_9(5) ... OK
Call assignment_4_9(1) ... OK
Call assignment_4_9(7) ... OK
Call assignment_4_9(8) ... OK
Call assignment_4_9(9) ... OK
Call assignment_4_9(6) ... OK

Well done! You seem to have solved assignment_4_9!
```

<br />
<br />
<br />
<br />

## 4.10 Lost card

There was a set of cards with numbers from 1 to `top`. One of the card is now lost. Determine the number on that lost card given the numbers for the remaining cards.

Given a number `top`, followed by `top - 1` integers - representing the numbers on the remaining cards (distinct integers in the range from 1 to `top`). Find and return the number on the lost card.

## Example input

```python
5

1
2
3
4
```

## Example return value

```python
5
```

**Hint** Make use of the fact that the number on the missing card is the difference between the sum of *all* cards minus the sum of the cards entered.


<!-- grade_id: assignment-4-10-code, points: 0, grade: false, solution: true -->
```python
# Code for assignment 4.10

def assignment_4_10(list_of_numbers):
    n = list_of_numbers[0]
    cards_present = list_of_numbers[1:]
    sum = 0
    for i in cards_present:
        sum += i
    card_lost = int(((n*(n+1))/2) - sum)
    return card_lost 
        
        
    
```

```python
# Test for assignment 4.10

test_cell_4_10(assignment_4_10)
```

Output:
```
Please enter the number of cards (top) (Press Enter for default value 5):  
Please enter the value for card 1 (Press Enter for default value 1):  
Please enter the value for card 2 (Press Enter for default value 2):  
Please enter the value for card 3 (Press Enter for default value 3):  
Please enter the value for card 4 (Press Enter for default value 4):  

The missing card has face value: 
5
```

<!-- grade_id: assignment-4-10-assert, points: 1, grade: true, solution: false -->
```python
# Assertion for assignment 4.10

check_assertion(assignment_4_10)
```

Output:
```
Call assignment_4_10([5, 1, 2, 3, 4]) ... OK
Call assignment_4_10([5, 3, 5, 2, 1]) ... OK
Call assignment_4_10([4, 2, 3, 4]) ... OK
Call assignment_4_10([3, 1, 2]) ... OK
Call assignment_4_10([3, 3, 2]) ... OK
Call assignment_4_10([3, 3, 1]) ... OK
Call assignment_4_10([1]) ... OK
Call assignment_4_10([10, 4, 1, 7, 8, 3, 5, 9, 10, 6]) ... OK

Well done! You seem to have solved assignment_4_10!
```